{
 "cells": [
  {
   "cell_type": "code",
   "execution_count": 1,
   "metadata": {},
   "outputs": [
    {
     "name": "stderr",
     "output_type": "stream",
     "text": [
      "/home/apluska/foreduce/foreduce/data/data.py:64: FutureWarning: You are using `torch.load` with `weights_only=False` (the current default value), which uses the default pickle module implicitly. It is possible to construct malicious pickle data which will execute arbitrary code during unpickling (See https://github.com/pytorch/pytorch/blob/main/SECURITY.md#untrusted-models for more details). In a future release, the default value for `weights_only` will be flipped to `True`. This limits the functions that could be executed during unpickling. Arbitrary objects will no longer be allowed to be loaded via this mode unless they are explicitly allowlisted by the user via `torch.serialization.add_safe_globals`. We recommend you start setting `weights_only=True` for any use case where you don't have full control of the loaded file. Please open an issue on GitHub for any issues related to this experimental feature.\n",
      "  max_arity, data, labels = torch.load(path)\n"
     ]
    }
   ],
   "source": [
    "from foreduce.data.data import GraphDataset\n",
    "\n",
    "trainset = GraphDataset.load(f'./data/666_train.pt')"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 27,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/plain": [
       "tensor(1.1953)"
      ]
     },
     "execution_count": 27,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "import torch\n",
    "\n",
    "x = torch.tensor([1.0, 0.0, 1.0])\n",
    "y = torch.tensor([0.5, 0.0, 0.5])\n",
    "\n",
    "alpha = 1.0\n",
    "\n",
    "torch.nn.functional.cross_entropy(x, y, label_smoothing=alpha)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 2,
   "metadata": {},
   "outputs": [],
   "source": [
    "graph, labels = trainset[0]"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 3,
   "metadata": {},
   "outputs": [],
   "source": [
    "from torch_geometric.loader import DataLoader\n",
    "\n",
    "loader  = DataLoader(trainset, batch_size=1, shuffle=True)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 5,
   "metadata": {},
   "outputs": [
    {
     "name": "stderr",
     "output_type": "stream",
     "text": [
      "/home/apluska/miniconda3/envs/foreduce/lib/python3.12/site-packages/tqdm/auto.py:21: TqdmWarning: IProgress not found. Please update jupyter and ipywidgets. See https://ipywidgets.readthedocs.io/en/stable/user_install.html\n",
      "  from .autonotebook import tqdm as notebook_tqdm\n"
     ]
    }
   ],
   "source": [
    "from foreduce.transformer.model import GraphModel\n",
    "from foreduce.data.data import _type_mapping\n",
    "\n",
    "\n",
    "model = GraphModel(len(_type_mapping), trainset.max_arity, 6, 64)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 10,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/plain": [
       "tensor([0.0000, 0.3333, 0.3333, 0.3333, 0.0000, 0.0000, 0.0000, 0.0000, 0.0000,\n",
       "        0.0000, 0.0000])"
      ]
     },
     "execution_count": 10,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "b.to(torch.float).reshape(-1) / torch.sum(b)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 16,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/plain": [
       "11"
      ]
     },
     "execution_count": 16,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "b.shape[1]"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 17,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/plain": [
       "tensor(0.2224, grad_fn=<DivBackward0>)"
      ]
     },
     "execution_count": 17,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "import torch\n",
    "\n",
    "torch.nn.functional.cross_entropy(model(a), b.to(torch.float).reshape(-1) / torch.sum(b)) / b.shape[1]"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 67,
   "metadata": {},
   "outputs": [],
   "source": [
    "from foreduce.vampire.vampire import VampireInteractive\n",
    "import os\n",
    "import torch\n",
    "\n",
    "VAMPIRE = os.getenv(\"VAMPIRE\")\n",
    "\n",
    "\n",
    "interactive = VampireInteractive(VAMPIRE, f'./problems/SYN/SYN986+1.005.p')\n",
    "interactive.__enter__()\n",
    "symbols = interactive.problem.function_symbols() | interactive.problem.predicate_symbols()\n",
    "permutation = torch.randperm(32)\n",
    "name = {s: permutation[i] + 1 for i, s in enumerate(symbols)}"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 68,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/plain": [
       "goal_0 | ~r(zero, X4, X5) | ~r(zero, X3, X4) | ~r(zero, X2, X3) | ~r(zero, X1, X2) | ~r(zero, X0, X1) | ~r(zero, zero, X0)\n",
       "r(X6, zero, succ(X6))\n",
       "r(X7, succ(X8), X10) | ~r(X9, X8, X10) | ~r(X7, X8, X9)\n",
       "~goal_0"
      ]
     },
     "execution_count": 68,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "interactive.problem"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 69,
   "metadata": {},
   "outputs": [],
   "source": [
    "graphs = []\n",
    "datalist = []"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 70,
   "metadata": {},
   "outputs": [],
   "source": [
    "from torch_geometric.utils.convert import from_networkx\n",
    "from torch_geometric.data import Batch\n",
    "\n",
    "from foreduce.data.data import _type_mapping\n"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 71,
   "metadata": {},
   "outputs": [],
   "source": [
    "from foreduce.transformer.model import Model\n",
    "\n",
    "model = Model.load_from_checkpoint(f'./models/test.ckpt')"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 72,
   "metadata": {},
   "outputs": [],
   "source": [
    "active = torch.zeros(1024, dtype=torch.bool)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 73,
   "metadata": {},
   "outputs": [],
   "source": [
    "import sys\n",
    "sys.setrecursionlimit(10000)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 138,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/plain": [
       "\u001b[1m 0: goal_0 | ~r(zero, X4, X5) | ~r(zero, X3, X4) | ~r(zero, X2, X3) | ~r(zero, X1, X2) | ~r(zero, X0, X1) | ~r(zero, zero, X0)\u001b[0m\n",
       "\u001b[1m 1: r(X6, zero, succ(X6))\u001b[0m\n",
       "\u001b[1m 2: r(X7, succ(X8), X10) | ~r(X9, X8, X10) | ~r(X7, X8, X9)\u001b[0m\n",
       "\u001b[1m 3: ~goal_0\u001b[0m\n",
       "\u001b[1m 4: ~r(zero, X0, X1) | ~r(zero, X2, X0) | ~r(zero, X3, X2) | ~r(zero, X4, X3) | ~r(zero, succ(zero), X4) | goal_0\u001b[0m\n",
       "\u001b[1m 5: r(X0, succ(zero), succ(X1)) | ~r(X0, zero, X1)\u001b[0m\n",
       "\u001b[1m 6: ~r(zero, zero, X0) | ~r(zero, X1, X2) | ~r(zero, X3, X1) | ~r(zero, succ(X0), X3) | ~r(zero, X2, X4) | goal_0\u001b[0m\n",
       "\u001b[1m 7: ~r(X5, zero, X6) | r(X7, succ(succ(zero)), succ(X6)) | ~r(X7, succ(zero), X5)\u001b[0m\n",
       "\u001b[1m 8: ~r(zero, zero, X0) | ~r(zero, X1, X2) | ~r(zero, X3, X1) | ~r(zero, succ(X0), X3) | ~r(zero, X2, X4)\u001b[0m\n",
       "\u001b[1m 9: r(X0, succ(succ(zero)), succ(X1)) | ~r(succ(X2), zero, X1) | ~r(X0, zero, X2)\u001b[0m\n",
       "\u001b[1m 10: r(X0, succ(succ(zero)), succ(succ(succ(X1)))) | ~r(X0, zero, X1)\u001b[0m\n",
       "\u001b[1m 11: ~r(X0, zero, X1) | r(X2, succ(succ(succ(zero))), succ(succ(succ(X1)))) | ~r(X2, succ(succ(zero)), X0)\u001b[0m\n",
       "\u001b[1m 12: r(X0, succ(succ(succ(zero))), succ(succ(succ(X1)))) | ~r(succ(succ(succ(X2))), zero, X1) | ~r(X0, zero, X2)\u001b[0m\n",
       "\u001b[1m 13: r(X0, succ(succ(succ(zero))), succ(succ(succ(succ(succ(succ(succ(X1)))))))) | ~r(X0, zero, X1)\u001b[0m\n",
       "\u001b[1m 14: ~r(zero, X0, X1) | ~r(zero, succ(X2), X0) | ~r(zero, zero, zero) | ~r(zero, X1, X3) | ~r(zero, zero, X2)\u001b[0m\n",
       "\u001b[1m 15: ~r(zero, X4, X5) | ~r(zero, succ(succ(succ(X6))), X4) | ~r(zero, zero, succ(zero)) | ~r(zero, X5, X7) | ~r(zero, zero, X6)\u001b[0m\n",
       "\u001b[1m 16: ~r(zero, X4, X5) | ~r(zero, succ(succ(succ(X6))), X4) | ~r(zero, X5, X7) | ~r(zero, zero, X6)\u001b[0m\n",
       "\u001b[1m 17: ~r(zero, succ(X0), X1) | ~r(zero, zero, zero) | ~r(zero, X1, X2) | ~r(zero, zero, zero) | ~r(zero, zero, X0)\u001b[0m\n",
       "\u001b[1m 18: ~r(zero, succ(succ(succ(X3))), X4) | ~r(zero, zero, zero) | ~r(zero, X4, X5) | ~r(zero, zero, succ(zero)) | ~r(zero, zero, X3)\u001b[0m\n",
       "\u001b[1m 19: ~r(zero, succ(X0), X1) | ~r(zero, zero, zero) | ~r(zero, X1, X2) | ~r(zero, zero, X0)\u001b[0m\n",
       "\u001b[1m 20: ~r(zero, succ(succ(succ(X3))), X4) | ~r(zero, zero, zero) | ~r(zero, X4, X5) | ~r(zero, zero, X3)\u001b[0m\n",
       "\u001b[1m 21: ~r(zero, zero, zero) | ~r(zero, succ(X0), X1) | ~r(zero, zero, zero) | ~r(zero, zero, X0)\u001b[0m\n",
       "\u001b[1m 22: ~r(zero, zero, zero) | ~r(zero, succ(succ(succ(X2))), X3) | ~r(zero, zero, succ(zero)) | ~r(zero, zero, X2)\u001b[0m\n",
       "\u001b[1m 23: ~r(zero, zero, zero) | ~r(zero, succ(X0), X1) | ~r(zero, zero, X0)\u001b[0m\n",
       "\u001b[1m 24: ~r(zero, zero, zero) | ~r(zero, succ(succ(succ(X2))), X3) | ~r(zero, zero, X2)\u001b[0m\n",
       "\u001b[1m 25: ~r(zero, zero, zero) | ~r(zero, zero, zero) | ~r(zero, zero, X0)\u001b[0m\n",
       "\u001b[1m 26: ~r(zero, zero, zero) | ~r(zero, zero, succ(zero)) | ~r(zero, zero, X1)\u001b[0m\n",
       "\u001b[1m 27: ~r(zero, zero, zero) | ~r(zero, zero, X0)\u001b[0m\n",
       "\u001b[1m 28: ~r(zero, zero, zero) | ~r(zero, zero, X1)\u001b[0m\n",
       "\u001b[1m 29: ~r(zero, zero, X0) | ~r(zero, zero, zero) | ~r(zero, zero, zero)\u001b[0m\n",
       "\u001b[1m 30: ~r(zero, zero, X1) | ~r(zero, zero, zero) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(X1))))))), X2) | ~r(zero, zero, zero)\u001b[0m\n",
       "\u001b[1m 31: ~r(zero, zero, X3) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(X3))))))), X4) | ~r(zero, X4, X5) | ~r(zero, zero, zero)\u001b[0m\n",
       "\u001b[1m 32: ~r(zero, zero, X6) | ~r(zero, zero, zero) | ~r(zero, zero, succ(succ(zero)))\u001b[0m\n",
       "\u001b[1m 33: ~r(zero, zero, X7) | ~r(zero, zero, zero) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(X7))))))), X8) | ~r(zero, zero, succ(succ(zero)))\u001b[0m\n",
       "\u001b[1m 34: ~r(zero, zero, X9) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(X9))))))), X10) | ~r(zero, zero, zero) | ~r(zero, X10, X11) | ~r(zero, zero, succ(succ(zero)))\u001b[0m\n",
       "\u001b[1m 35: ~r(zero, zero, X12) | ~r(zero, X13, X14) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(X12))))))), X13) | ~r(zero, zero, succ(succ(zero))) | ~r(zero, X14, X15)\u001b[0m\n",
       "\u001b[1m 36: ~r(X16, zero, X17) | r(X18, succ(succ(succ(succ(zero)))), succ(succ(succ(succ(succ(succ(succ(X17)))))))) | ~r(X18, succ(succ(succ(zero))), X16)\u001b[0m\n",
       "\u001b[1m 37: ~r(zero, zero, X1) | ~r(zero, zero, zero) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(X1))))))), X2)\u001b[0m\n",
       "\u001b[1m 38: ~r(zero, zero, X0) | ~r(zero, zero, zero)\u001b[0m\n",
       "\u001b[1m 39: r(X0, succ(succ(succ(succ(zero)))), succ(succ(succ(succ(succ(succ(succ(X1)))))))) | ~r(succ(succ(succ(succ(succ(succ(succ(X2))))))), zero, X1) | ~r(X0, zero, X2)\u001b[0m\n",
       "\u001b[1m 40: r(X0, succ(succ(succ(succ(zero)))), succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))) | ~r(X0, zero, X1)\u001b[0m\n",
       "\u001b[1m 41: ~r(zero, zero, X0) | ~r(zero, zero, zero) | ~r(zero, zero, succ(zero))\u001b[0m\n",
       "\u001b[1m 42: ~r(zero, zero, X1) | ~r(zero, zero, zero) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1))))))))))))))), X2) | ~r(zero, zero, succ(zero))\u001b[0m\n",
       "\u001b[1m 43: ~r(zero, zero, X3) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X3))))))))))))))), X4) | ~r(zero, X4, X5) | ~r(zero, zero, succ(zero))\u001b[0m\n",
       "\u001b[1m 44: ~r(zero, zero, X6) | ~r(zero, zero, zero) | ~r(zero, zero, succ(succ(succ(zero))))\u001b[0m\n",
       "\u001b[1m 45: ~r(zero, zero, X7) | ~r(zero, zero, zero) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X7))))))))))))))), X8) | ~r(zero, zero, succ(succ(succ(zero))))\u001b[0m\n",
       "\u001b[1m 46: ~r(zero, zero, X9) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X9))))))))))))))), X10) | ~r(zero, zero, zero) | ~r(zero, X10, X11) | ~r(zero, zero, succ(succ(succ(zero))))\u001b[0m\n",
       "\u001b[1m 47: ~r(zero, zero, X12) | ~r(zero, X13, X14) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X12))))))))))))))), X13) | ~r(zero, zero, succ(succ(succ(zero)))) | ~r(zero, X14, X15)\u001b[0m\n",
       "\u001b[1m 48: ~r(X16, zero, X17) | r(X18, succ(succ(succ(succ(succ(zero))))), succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X17)))))))))))))))) | ~r(X18, succ(succ(succ(succ(zero)))), X16)\u001b[0m\n",
       "\u001b[1m 49: ~r(zero, zero, X3) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X3))))))))))))))), X4) | ~r(zero, X4, X5)\u001b[0m\n",
       "\u001b[1m 50: r(X0, succ(succ(succ(succ(succ(zero))))), succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))) | ~r(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X2))))))))))))))), zero, X1) | ~r(X0, zero, X2)\u001b[0m\n",
       "\u001b[1m 51: r(X0, succ(succ(succ(succ(succ(zero))))), succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))) | ~r(X0, zero, X1)\u001b[0m\n",
       "\u001b[1m 52: ~r(zero, zero, X0) | ~r(zero, zero, zero) | ~r(zero, zero, succ(succ(zero)))\u001b[0m\n",
       "\u001b[1m 53: ~r(zero, zero, X1) | ~r(zero, zero, zero) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1))))))))))))))))))))))))))))))), X2) | ~r(zero, zero, succ(succ(zero)))\u001b[0m\n",
       "\u001b[1m 54: ~r(zero, zero, X3) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X3))))))))))))))))))))))))))))))), X4) | ~r(zero, X4, X5) | ~r(zero, zero, succ(succ(zero)))\u001b[0m\n",
       "\u001b[1m 55: ~r(zero, zero, X6) | ~r(zero, zero, zero) | ~r(zero, zero, succ(succ(succ(succ(zero)))))\u001b[0m\n",
       "\u001b[1m 56: ~r(zero, zero, X7) | ~r(zero, zero, zero) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X7))))))))))))))))))))))))))))))), X8) | ~r(zero, zero, succ(succ(succ(succ(zero)))))\u001b[0m\n",
       "\u001b[1m 57: ~r(zero, zero, X9) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X9))))))))))))))))))))))))))))))), X10) | ~r(zero, zero, zero) | ~r(zero, X10, X11) | ~r(zero, zero, succ(succ(succ(succ(zero)))))\u001b[0m\n",
       "\u001b[1m 58: ~r(zero, zero, X12) | ~r(zero, X13, X14) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X12))))))))))))))))))))))))))))))), X13) | ~r(zero, zero, succ(succ(succ(succ(zero))))) | ~r(zero, X14, X15)\u001b[0m\n",
       "\u001b[1m 59: ~r(X16, zero, X17) | r(X18, succ(succ(succ(succ(succ(succ(zero)))))), succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X17)))))))))))))))))))))))))))))))) | ~r(X18, succ(succ(succ(succ(succ(zero))))), X16)\u001b[0m\n",
       "\u001b[1m 60: r(X0, succ(succ(succ(succ(succ(succ(zero)))))), succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))) | ~r(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X2))))))))))))))))))))))))))))))), zero, X1) | ~r(X0, zero, X2)\u001b[0m\n",
       "\u001b[1m 61: r(X0, succ(succ(succ(succ(succ(succ(zero)))))), succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(X0, zero, X1)\u001b[0m\n",
       "\u001b[1m 62: ~r(zero, zero, X0) | ~r(zero, zero, zero) | ~r(zero, zero, succ(succ(succ(zero))))\u001b[0m\n",
       "63: ~r(zero, zero, X1) | ~r(zero, zero, zero) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))), X2) | ~r(zero, zero, succ(succ(succ(zero))))\n",
       "\u001b[1m 64: ~r(zero, zero, X3) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X3))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))), X4) | ~r(zero, X4, X5) | ~r(zero, zero, succ(succ(succ(zero))))\u001b[0m\n",
       "\u001b[1m 65: ~r(zero, zero, X6) | ~r(zero, zero, zero) | ~r(zero, zero, succ(succ(succ(succ(succ(zero))))))\u001b[0m\n",
       "66: ~r(zero, zero, X7) | ~r(zero, zero, zero) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X7))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))), X8) | ~r(zero, zero, succ(succ(succ(succ(succ(zero))))))\n",
       "67: ~r(zero, zero, X9) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X9))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))), X10) | ~r(zero, zero, zero) | ~r(zero, X10, X11) | ~r(zero, zero, succ(succ(succ(succ(succ(zero))))))\n",
       "68: ~r(zero, zero, X12) | ~r(zero, X13, X14) | ~r(zero, succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X12))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))), X13) | ~r(zero, zero, succ(succ(succ(succ(succ(zero)))))) | ~r(zero, X14, X15)\n",
       "69: ~r(X16, zero, X17) | r(X18, succ(succ(succ(succ(succ(succ(succ(zero))))))), succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X17)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(X18, succ(succ(succ(succ(succ(succ(zero)))))), X16)"
      ]
     },
     "execution_count": 138,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "graphs = interactive.problem.extend(graphs, len(graphs), depth=8)\n",
    "for graph in graphs[len(datalist):]:\n",
    "    data = from_networkx(graph)\n",
    "    data.type = torch.tensor([_type_mapping[t] for t in data.type], dtype=torch.int)\n",
    "    data.name = torch.tensor([name[s] if s is not None else 0 for s in data.name], dtype=torch.int)\n",
    "    data.arity = torch.tensor([min(4 + 1, a + 1) if a is not None else 0 for a in data.arity], dtype=torch.int)\n",
    "    data.pos = torch.tensor([min(4 + 1, a + 1) if a is not None else 0 for a in data.pos], dtype=torch.int)\n",
    "    datalist.append(data)\n",
    "score = model(Batch.from_data_list(datalist))\n",
    "given = torch.where(active[:len(graphs)], -10000, score).argmax().item()\n",
    "active[given] = True\n",
    "interactive.step(given)\n",
    "interactive"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 57,
   "metadata": {},
   "outputs": [
    {
     "data": {
      "text/plain": [
       "1272"
      ]
     },
     "execution_count": 57,
     "metadata": {},
     "output_type": "execute_result"
    }
   ],
   "source": [
    "sum(len(g) for g in graphs)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 56,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "% Running in auto input_syntax mode. Trying TPTP\n",
      "[SA] new: 1. goal_0 | ~r(zero,X4,X5) | ~r(zero,X3,X4) | ~r(zero,X2,X3) | ~r(zero,X1,X2) | ~r(zero,X0,X1) | ~r(zero,zero,X0) [input]\n",
      "[SA] new: 2. r(X6,zero,succ(X6)) [input]\n",
      "[SA] new: 3. r(X7,succ(X8),X10) | ~r(X9,X8,X10) | ~r(X7,X8,X9) [input]\n",
      "[SA] new: 4. ~goal_0 [input]\n",
      "[SA] new: 5. r(X0,succ(zero),succ(X1)) | ~r(X0,zero,X1) [resolution 3,2]\n",
      "[SA] new: 6. ~r(X0,zero,X1) | r(X2,succ(succ(zero)),succ(X1)) | ~r(X2,succ(zero),X0) [resolution 5,3]\n",
      "[SA] new: 7. r(X0,succ(succ(zero)),succ(X1)) | ~r(succ(X2),zero,X1) | ~r(X0,zero,X2) [resolution 6,5]\n",
      "[SA] new: 8. r(X0,succ(succ(zero)),succ(succ(succ(X1)))) | ~r(X0,zero,X1) [resolution 7,2]\n",
      "[SA] new: 9. ~r(X0,zero,X1) | r(X2,succ(succ(succ(zero))),succ(succ(succ(X1)))) | ~r(X2,succ(succ(zero)),X0) [resolution 8,3]\n",
      "[SA] new: 10. r(X0,succ(succ(succ(zero))),succ(succ(succ(X1)))) | ~r(succ(succ(succ(X2))),zero,X1) | ~r(X0,zero,X2) [resolution 9,8]\n",
      "[SA] new: 11. r(X0,succ(succ(succ(zero))),succ(succ(succ(succ(succ(succ(succ(X1)))))))) | ~r(X0,zero,X1) [resolution 10,2]\n",
      "[SA] new: 12. ~r(X0,zero,X1) | r(X2,succ(succ(succ(succ(zero)))),succ(succ(succ(succ(succ(succ(succ(X1)))))))) | ~r(X2,succ(succ(succ(zero))),X0) [resolution 11,3]\n",
      "[SA] new: 13. r(X0,succ(succ(succ(succ(zero)))),succ(succ(succ(succ(succ(succ(succ(X1)))))))) | ~r(succ(succ(succ(succ(succ(succ(succ(X2))))))),zero,X1) | ~r(X0,zero,X2) [resolution 12,11]\n",
      "[SA] new: 14. r(X0,succ(succ(succ(succ(zero)))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))) | ~r(X0,zero,X1) [resolution 13,2]\n",
      "[SA] new: 15. ~r(X0,zero,X1) | r(X2,succ(succ(succ(succ(succ(zero))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))) | ~r(X2,succ(succ(succ(succ(zero)))),X0) [resolution 14,3]\n",
      "[SA] new: 16. r(X0,succ(succ(succ(succ(succ(zero))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))) | ~r(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X2))))))))))))))),zero,X1) | ~r(X0,zero,X2) [resolution 15,14]\n",
      "[SA] new: 17. r(X0,succ(succ(succ(succ(succ(zero))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))) | ~r(X0,zero,X1) [resolution 16,2]\n",
      "[SA] new: 18. ~r(X0,zero,X1) | r(X2,succ(succ(succ(succ(succ(succ(zero)))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))) | ~r(X2,succ(succ(succ(succ(succ(zero))))),X0) [resolution 17,3]\n",
      "[SA] new: 19. r(X0,succ(succ(succ(succ(succ(succ(zero)))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))) | ~r(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X2))))))))))))))))))))))))))))))),zero,X1) | ~r(X0,zero,X2) [resolution 18,17]\n",
      "[SA] new: 20. r(X0,succ(succ(succ(succ(succ(succ(zero)))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(X0,zero,X1) [resolution 19,2]\n",
      "[SA] new: 21. ~r(X0,zero,X1) | r(X2,succ(succ(succ(succ(succ(succ(succ(zero))))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(X2,succ(succ(succ(succ(succ(succ(zero)))))),X0) [resolution 20,3]\n",
      "[SA] new: 22. r(X0,succ(succ(succ(succ(succ(succ(succ(zero))))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X2))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))),zero,X1) | ~r(X0,zero,X2) [resolution 21,20]\n",
      "[SA] new: 23. r(X0,succ(succ(succ(succ(succ(succ(succ(zero))))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(X0,zero,X1) [resolution 22,2]\n",
      "[SA] new: 24. ~r(X0,zero,X1) | r(X2,succ(succ(succ(succ(succ(succ(succ(succ(zero)))))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(X2,succ(succ(succ(succ(succ(succ(succ(zero))))))),X0) [resolution 23,3]\n",
      "[SA] new: 25. r(X0,succ(succ(succ(succ(succ(succ(succ(succ(zero)))))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X2))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))),zero,X1) | ~r(X0,zero,X2) [resolution 24,23]\n",
      "[SA] new: 26. r(X0,succ(succ(succ(succ(succ(succ(succ(succ(zero)))))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(X0,zero,X1) [resolution 25,2]\n",
      "[SA] new: 27. ~r(X0,zero,X1) | r(X2,succ(succ(succ(succ(succ(succ(succ(succ(succ(zero))))))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(X2,succ(succ(succ(succ(succ(succ(succ(succ(zero)))))))),X0) [resolution 26,3]\n",
      "[SA] new: 28. r(X0,succ(succ(succ(succ(succ(succ(succ(succ(succ(zero))))))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X2))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))),zero,X1) | ~r(X0,zero,X2) [resolution 27,26]\n",
      "[SA] new: 29. r(X0,succ(succ(succ(succ(succ(succ(succ(succ(succ(zero))))))))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X1)))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))))) | ~r(X0,zero,X1) [resolution 28,2]\n",
      "\n"
     ]
    }
   ],
   "source": [
    "print(interactive.proof)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 145,
   "metadata": {},
   "outputs": [
    {
     "ename": "RuntimeError",
     "evalue": "The size of tensor a (1024) must match the size of tensor b (1108) at non-singleton dimension 0",
     "output_type": "error",
     "traceback": [
      "\u001b[0;31m---------------------------------------------------------------------------\u001b[0m",
      "\u001b[0;31mRuntimeError\u001b[0m                              Traceback (most recent call last)",
      "Cell \u001b[0;32mIn[145], line 20\u001b[0m\n\u001b[1;32m     18\u001b[0m     datalist\u001b[38;5;241m.\u001b[39mappend(data)\n\u001b[1;32m     19\u001b[0m score \u001b[38;5;241m=\u001b[39m model(Batch\u001b[38;5;241m.\u001b[39mfrom_data_list(datalist))\n\u001b[0;32m---> 20\u001b[0m given \u001b[38;5;241m=\u001b[39m torch\u001b[38;5;241m.\u001b[39mwhere(active[:\u001b[38;5;28mlen\u001b[39m(graphs)], \u001b[38;5;241m-\u001b[39m\u001b[38;5;241m10000\u001b[39m, score)\u001b[38;5;241m.\u001b[39margmax()\u001b[38;5;241m.\u001b[39mitem()\n\u001b[1;32m     21\u001b[0m active[given] \u001b[38;5;241m=\u001b[39m \u001b[38;5;28;01mTrue\u001b[39;00m\n\u001b[1;32m     22\u001b[0m interactive\u001b[38;5;241m.\u001b[39mstep(given)\n",
      "\u001b[0;31mRuntimeError\u001b[0m: The size of tensor a (1024) must match the size of tensor b (1108) at non-singleton dimension 0"
     ]
    }
   ],
   "source": [
    "graphs = []\n",
    "datalist = []\n",
    "active = torch.zeros(1024, dtype=torch.bool)\n",
    "with VampireInteractive(VAMPIRE, f'./problems/GRP/GRP001-2.p') as interactive:\n",
    "    symbols = interactive.problem.function_symbols() | interactive.problem.predicate_symbols()\n",
    "    permutation = torch.randperm(32)\n",
    "    name = {s: permutation[i] + 1 for i, s in enumerate(symbols)}\n",
    "    while not interactive.finished and interactive.step_count < 128 and (\n",
    "        sum(len(g) for g in graphs) < 2**20\n",
    "    ) and (len(graphs) == 0 or not active[:len(graphs)].all()):\n",
    "        graphs = interactive.problem.extend(graphs, len(graphs), depth=4)\n",
    "        for graph in graphs[len(datalist):]:\n",
    "            data = from_networkx(graph)\n",
    "            data.type = torch.tensor([_type_mapping[t] for t in data.type], dtype=torch.int)\n",
    "            data.name = torch.tensor([name[s] if s is not None else 0 for s in data.name], dtype=torch.int)\n",
    "            data.arity = torch.tensor([min(8 + 1, a + 1) if a is not None else 0 for a in data.arity], dtype=torch.int)\n",
    "            data.pos = torch.tensor([min(8 + 1, a + 1) if a is not None else 0 for a in data.pos], dtype=torch.int)\n",
    "            datalist.append(data)\n",
    "        score = model(Batch.from_data_list(datalist))\n",
    "        given = torch.where(active[:len(graphs)], -10000, score).argmax().item()\n",
    "        active[given] = True\n",
    "        interactive.step(given)"
   ]
  },
  {
   "cell_type": "code",
   "execution_count": 146,
   "metadata": {},
   "outputs": [
    {
     "name": "stdout",
     "output_type": "stream",
     "text": [
      "% Running in auto input_syntax mode. Trying TPTP\n",
      "[SA] new: 1. multiply(X0,identity) = X0 [input]\n",
      "[SA] new: 2. identity = multiply(X1,inverse(X1)) [input]\n",
      "[SA] new: 3. identity = multiply(X2,X2) [input]\n",
      "[SA] new: 4. multiply(a,b) = c [input]\n",
      "[SA] new: 5. goal_0 | c != multiply(b,a) [input]\n",
      "[SA] new: 6. multiply(identity,X3) = X3 [input]\n",
      "[SA] new: 7. identity = multiply(inverse(X4),X4) [input]\n",
      "[SA] new: 8. multiply(multiply(X5,X6),X7) = multiply(X5,multiply(X6,X7)) [input]\n",
      "[SA] new: 9. ~goal_0 [input]\n",
      "[SA] new: 10. identity = inverse(identity) [superposition 1,7]\n",
      "[SA] new: 11. identity = inverse(identity) [superposition 7,1]\n",
      "[SA] new: 12. identity = multiply(identity,identity) [superposition 7,10]\n",
      "[SA] new: 13. multiply(inverse(X0),multiply(X0,X1)) = multiply(identity,X1) [superposition 8,7]\n",
      "[SA] new: 14. multiply(identity,X3) = multiply(X2,multiply(X2,X3)) [superposition 8,3]\n",
      "[SA] new: 15. multiply(X4,X5) = multiply(identity,multiply(X4,X5)) [superposition 8,6]\n",
      "[SA] new: 16. multiply(multiply(X6,X7),multiply(X8,X9)) = multiply(multiply(X6,multiply(X7,X8)),X9) [superposition 8,8]\n",
      "[SA] new: 17. multiply(X10,multiply(identity,X11)) = multiply(X10,X11) [superposition 8,1]\n",
      "[SA] new: 18. multiply(a,multiply(b,X12)) = multiply(c,X12) [superposition 8,4]\n",
      "[SA] new: 19. identity = multiply(X0,multiply(X1,multiply(X0,X1))) [superposition 8,3]\n",
      "[SA] new: 20. multiply(X5,X6) = multiply(X5,multiply(X6,identity)) [superposition 8,1]\n",
      "[SA] new: 21. identity = multiply(X0,multiply(X1,multiply(X0,X1))) [superposition 3,8]\n",
      "[SA] new: 22. multiply(X2,X3) = multiply(X2,multiply(X3,identity)) [superposition 1,8]\n",
      "[SA] new: 23. multiply(inverse(X0),multiply(X0,X1)) = X1 [forward demodulation 13,6]\n",
      "[SA] new: 24. multiply(X2,multiply(X2,X3)) = X3 [forward demodulation 14,6]\n",
      "[SA] new: 25. multiply(multiply(X6,X7),multiply(X8,X9)) = multiply(X6,multiply(multiply(X7,X8),X9)) [forward demodulation 16,8]\n",
      "[SA] new: 26. multiply(multiply(X6,X7),multiply(X8,X9)) = multiply(X6,multiply(X7,multiply(X8,X9))) [forward demodulation 25,8]\n",
      "[SA] new: 27. multiply(identity,multiply(identity,X0)) = X0 [superposition 23,10]\n",
      "[SA] new: 28. multiply(inverse(inverse(X0)),identity) = X0 [superposition 23,7]\n",
      "[SA] new: 29. multiply(inverse(X1),identity) = X1 [superposition 23,3]\n",
      "[SA] new: 30. multiply(inverse(identity),X2) = X2 [superposition 23,6]\n",
      "[SA] new: 31. multiply(inverse(multiply(X3,X4)),multiply(X3,multiply(X4,X5))) = X5 [superposition 23,8]\n",
      "[SA] new: 32. identity = multiply(inverse(X6),X6) [superposition 23,1]\n",
      "[SA] new: 33. multiply(X7,X8) = multiply(inverse(inverse(X7)),X8) [superposition 23,23]\n",
      "[SA] new: 34. b = multiply(inverse(a),c) [superposition 23,4]\n",
      "[SA] new: 35. multiply(X1,X2) = multiply(inverse(X0),multiply(multiply(X0,X1),X2)) [superposition 8,23]\n",
      "[SA] new: 36. multiply(identity,X0) = X0 [forward demodulation 27,6]\n",
      "[SA] new: 37. multiply(identity,X2) = X2 [forward demodulation 30,10]\n",
      "[SA] new: 38. multiply(X1,X2) = multiply(inverse(X0),multiply(X0,multiply(X1,X2))) [forward demodulation 35,8]\n",
      "[SA] new: 39. multiply(identity,X0) = multiply(inverse(identity),X0) [superposition 33,10]\n",
      "[SA] new: 40. identity = multiply(X0,inverse(X0)) [superposition 33,7]\n",
      "[SA] new: 41. identity = multiply(X1,inverse(inverse(X1))) [superposition 33,3]\n",
      "[SA] new: 42. inverse(inverse(X4)) = multiply(X4,identity) [superposition 33,1]\n",
      "[SA] new: 43. multiply(X5,multiply(inverse(X5),X6)) = X6 [superposition 33,23]\n",
      "[SA] new: 44. identity = multiply(X0,inverse(X0)) [superposition 7,33]\n",
      "[SA] new: 45. identity = multiply(X1,inverse(inverse(X1))) [superposition 3,33]\n",
      "[SA] new: 46. multiply(inverse(inverse(inverse(X2))),multiply(X2,X3)) = X3 [superposition 23,33]\n",
      "[SA] new: 47. multiply(inverse(inverse(X4)),multiply(X5,X6)) = multiply(multiply(X4,X5),X6) [superposition 8,33]\n",
      "[SA] new: 48. inverse(inverse(X7)) = multiply(X7,identity) [superposition 1,33]\n",
      "[SA] new: 49. multiply(X8,multiply(inverse(X8),X9)) = X9 [superposition 23,33]\n",
      "[SA] new: 50. inverse(inverse(X4)) = X4 [forward demodulation 42,1]\n",
      "[SA] new: 51. multiply(inverse(inverse(X4)),multiply(X5,X6)) = multiply(X4,multiply(X5,X6)) [forward demodulation 47,8]\n",
      "[SA] new: 52. inverse(inverse(X7)) = X7 [forward demodulation 48,1]\n",
      "[SA] new: 53. identity = inverse(identity) [superposition 50,10]\n",
      "[SA] new: 54. multiply(X0,multiply(inverse(X0),X1)) = X1 [superposition 23,50]\n",
      "[SA] new: 55. identity = multiply(X2,inverse(X2)) [superposition 7,50]\n",
      "[SA] new: 56. multiply(inverse(X2),multiply(X2,X3)) = X3 [forward demodulation 46,50]\n",
      "[SA] new: 57. multiply(identity,multiply(identity,X0)) = X0 [superposition 49,10]\n",
      "[SA] new: 58. multiply(inverse(X1),multiply(X1,X2)) = X2 [superposition 49,50]\n",
      "[SA] new: 59. multiply(X0,identity) = X0 [superposition 49,7]\n",
      "[SA] new: 60. inverse(X1) = multiply(X1,identity) [superposition 49,3]\n",
      "[SA] new: 61. identity = multiply(X2,inverse(X2)) [superposition 49,1]\n",
      "[SA] new: 62. multiply(X5,X6) = multiply(inverse(inverse(X5)),X6) [superposition 49,49]\n",
      "[SA] new: 63. multiply(inverse(identity),X2) = X2 [superposition 49,6]\n",
      "[SA] new: 64. multiply(X3,multiply(X4,multiply(inverse(multiply(X3,X4)),X5))) = X5 [superposition 49,8]\n",
      "[SA] new: 65. multiply(inverse(identity),X0) = X0 [superposition 6,49]\n",
      "[SA] new: 66. multiply(X4,X5) = multiply(X3,multiply(multiply(inverse(X3),X4),X5)) [superposition 8,49]\n",
      "[SA] new: 67. multiply(X6,multiply(X7,multiply(inverse(multiply(X6,X7)),X8))) = X8 [superposition 8,49]\n",
      "[SA] new: 68. multiply(identity,X0) = X0 [forward demodulation 57,6]\n",
      "[SA] new: 69. inverse(X1) = X1 [forward demodulation 60,1]\n",
      "[SA] new: 70. multiply(identity,X2) = X2 [forward demodulation 63,10]\n",
      "[SA] new: 71. multiply(identity,X0) = X0 [forward demodulation 65,10]\n",
      "[SA] new: 72. multiply(X4,X5) = multiply(X3,multiply(inverse(X3),multiply(X4,X5))) [forward demodulation 66,8]\n",
      "[SA] new: 73. multiply(inverse(X0),multiply(X0,multiply(inverse(identity),X1))) = X1 [superposition 64,7]\n",
      "[SA] new: 74. multiply(X2,multiply(X2,multiply(inverse(identity),X3))) = X3 [superposition 64,3]\n",
      "[SA] new: 75. multiply(identity,multiply(X4,multiply(inverse(X4),X5))) = X5 [superposition 64,6]\n",
      "[SA] new: 76. multiply(multiply(X6,X7),multiply(X8,multiply(inverse(multiply(X6,multiply(X7,X8))),X9))) = X9 [superposition 64,8]\n",
      "[SA] new: 77. multiply(X10,multiply(identity,multiply(inverse(X10),X11))) = X11 [superposition 64,1]\n",
      "[SA] new: 78. multiply(inverse(X12),multiply(multiply(X12,X13),multiply(inverse(X13),X14))) = X14 [superposition 64,23]\n",
      "[SA] new: 79. multiply(X15,multiply(multiply(inverse(X15),X16),multiply(inverse(X16),X17))) = X17 [superposition 64,49]\n",
      "[SA] new: 80. multiply(X18,multiply(multiply(X19,multiply(inverse(multiply(X18,X19)),X20)),multiply(inverse(X20),X21))) = X21 [superposition 64,64]\n",
      "[SA] new: 81. multiply(a,multiply(b,multiply(inverse(c),X22))) = X22 [superposition 64,4]\n",
      "[SA] new: 82. multiply(X0,X1) = multiply(X0,multiply(X1,identity)) [superposition 64,7]\n",
      "[SA] new: 83. multiply(X2,multiply(X3,identity)) = inverse(multiply(X2,X3)) [superposition 64,3]\n",
      "[SA] new: 84. identity = multiply(X4,multiply(X5,inverse(multiply(X4,X5)))) [superposition 64,1]\n",
      "[SA] new: 85. multiply(multiply(X6,X7),X8) = multiply(X6,multiply(X7,X8)) [superposition 64,23]\n",
      "[SA] new: 86. multiply(inverse(inverse(multiply(X9,X10))),X11) = multiply(X9,multiply(X10,X11)) [superposition 64,49]\n",
      "[SA] new: 87. multiply(X14,multiply(inverse(multiply(inverse(multiply(X12,X13)),X14)),X15)) = multiply(X12,multiply(X13,X15)) [superposition 64,64]\n",
      "[SA] new: 88. multiply(X1,X2) = multiply(inverse(multiply(X0,inverse(multiply(X1,X0)))),X2) [superposition 64,64]\n",
      "[SA] new: 89. multiply(X3,multiply(inverse(multiply(X3,identity)),X4)) = X4 [superposition 64,6]\n",
      "[SA] new: 90. multiply(X7,multiply(X5,multiply(X6,multiply(inverse(multiply(X7,multiply(X5,X6))),X8)))) = X8 [superposition 64,8]\n",
      "[SA] new: 91. multiply(inverse(multiply(inverse(X3),X3)),X4) = X4 [superposition 64,23]\n",
      "[SA] new: 92. multiply(inverse(multiply(X5,inverse(X5))),X6) = X6 [superposition 64,49]\n",
      "[SA] new: 93. multiply(X7,multiply(inverse(multiply(identity,X7)),X8)) = X8 [superposition 64,6]\n",
      "[SA] new: 94. multiply(X9,multiply(X10,multiply(X11,multiply(inverse(multiply(multiply(X9,X10),X11)),X12)))) = X12 [superposition 64,8]\n",
      "[SA] new: 95. multiply(inverse(multiply(inverse(X0),X0)),X1) = X1 [superposition 23,64]\n",
      "[SA] new: 96. multiply(inverse(multiply(X2,inverse(X2))),X3) = X3 [superposition 49,64]\n",
      "[SA] new: 97. multiply(X4,multiply(inverse(multiply(identity,X4)),X5)) = X5 [superposition 6,64]\n",
      "[SA] new: 98. multiply(X7,multiply(inverse(multiply(X6,X7)),X8)) = multiply(inverse(X6),X8) [superposition 23,64]\n",
      "[SA] new: 99. multiply(X9,multiply(multiply(X10,multiply(inverse(multiply(X9,X10)),X11)),X12)) = multiply(X11,X12) [superposition 8,64]\n",
      "[SA] new: 100. multiply(X13,X15) = multiply(X14,multiply(inverse(multiply(inverse(X13),X14)),X15)) [superposition 49,64]\n",
      "[SA] new: 101. multiply(X16,multiply(X17,multiply(X18,multiply(inverse(multiply(multiply(X16,X17),X18)),X19)))) = X19 [superposition 8,64]\n",
      "[SA] new: 102. multiply(inverse(X0),multiply(X0,multiply(identity,X1))) = X1 [forward demodulation 73,10]\n",
      "[SA] new: 103. multiply(inverse(X0),multiply(X0,X1)) = X1 [forward demodulation 102,6]\n",
      "[SA] new: 104. multiply(X2,multiply(X2,multiply(identity,X3))) = X3 [forward demodulation 74,10]\n",
      "[SA] new: 105. multiply(X2,multiply(X2,X3)) = X3 [forward demodulation 104,6]\n",
      "[SA] new: 106. multiply(identity,X5) = X5 [forward demodulation 75,49]\n",
      "[SA] new: 107. multiply(X10,multiply(inverse(X10),X11)) = X11 [forward demodulation 77,6]\n",
      "[SA] new: 108. multiply(inverse(X12),multiply(X12,multiply(X13,multiply(inverse(X13),X14)))) = X14 [forward demodulation 78,8]\n",
      "[SA] new: 109. multiply(inverse(X12),multiply(X12,X14)) = X14 [forward demodulation 108,49]\n",
      "[SA] new: 110. multiply(X15,multiply(inverse(X15),multiply(X16,multiply(inverse(X16),X17)))) = X17 [forward demodulation 79,8]\n",
      "[SA] new: 111. multiply(X15,multiply(inverse(X15),X17)) = X17 [forward demodulation 110,49]\n",
      "[SA] new: 112. multiply(X18,multiply(X19,multiply(multiply(inverse(multiply(X18,X19)),X20),multiply(inverse(X20),X21)))) = X21 [forward demodulation 80,8]\n",
      "[SA] new: 113. multiply(X18,multiply(X19,multiply(inverse(multiply(X18,X19)),multiply(X20,multiply(inverse(X20),X21))))) = X21 [forward demodulation 112,8]\n",
      "[SA] new: 114. multiply(X18,multiply(X19,multiply(inverse(multiply(X18,X19)),X21))) = X21 [forward demodulation 113,49]\n",
      "[SA] new: 115. multiply(X2,X3) = inverse(multiply(X2,X3)) [forward demodulation 83,1]\n",
      "[SA] new: 116. multiply(X9,multiply(X10,X11)) = multiply(multiply(X9,X10),X11) [forward demodulation 86,50]\n",
      "[SA] new: 117. multiply(X3,multiply(inverse(X3),X4)) = X4 [forward demodulation 89,1]\n",
      "[SA] new: 118. multiply(inverse(identity),X4) = X4 [forward demodulation 91,7]\n",
      "[SA] new: 119. multiply(identity,X4) = X4 [forward demodulation 118,10]\n",
      "[SA] new: 120. multiply(X7,multiply(inverse(X7),X8)) = X8 [forward demodulation 93,6]\n",
      "[SA] new: 121. multiply(X9,multiply(X10,multiply(X11,multiply(inverse(multiply(X9,multiply(X10,X11))),X12)))) = X12 [forward demodulation 94,8]\n",
      "[SA] new: 122. multiply(inverse(identity),X1) = X1 [forward demodulation 95,7]\n",
      "[SA] new: 123. multiply(identity,X1) = X1 [forward demodulation 122,10]\n",
      "[SA] new: 124. multiply(X4,multiply(inverse(X4),X5)) = X5 [forward demodulation 97,6]\n",
      "[SA] new: 125. multiply(X11,X12) = multiply(X9,multiply(X10,multiply(multiply(inverse(multiply(X9,X10)),X11),X12))) [forward demodulation 99,8]\n",
      "[SA] new: 126. multiply(X11,X12) = multiply(X9,multiply(X10,multiply(inverse(multiply(X9,X10)),multiply(X11,X12)))) [forward demodulation 125,8]\n",
      "[SA] new: 127. multiply(X16,multiply(X17,multiply(X18,multiply(inverse(multiply(X16,multiply(X17,X18))),X19)))) = X19 [forward demodulation 101,8]\n",
      "[SA] new: 128. multiply(X3,multiply(inverse(X3),X5)) = X5 [backward demodulation 64,98]\n",
      "[SA] new: 129. multiply(inverse(inverse(X0)),X1) = multiply(X0,multiply(inverse(identity),X1)) [superposition 98,7]\n",
      "[SA] new: 130. multiply(X2,multiply(inverse(identity),X3)) = multiply(inverse(X2),X3) [superposition 98,3]\n",
      "[SA] new: 131. multiply(X4,multiply(inverse(X4),X5)) = multiply(inverse(identity),X5) [superposition 98,6]\n",
      "[SA] new: 132. multiply(X8,multiply(inverse(multiply(X6,multiply(X7,X8))),X9)) = multiply(inverse(multiply(X6,X7)),X9) [superposition 98,8]\n",
      "[SA] new: 133. multiply(inverse(X10),X11) = multiply(identity,multiply(inverse(X10),X11)) [superposition 98,1]\n",
      "[SA] new: 134. multiply(multiply(X12,X13),multiply(inverse(X13),X14)) = multiply(inverse(inverse(X12)),X14) [superposition 98,23]\n",
      "[SA] new: 135. multiply(multiply(inverse(X15),X16),multiply(inverse(X16),X17)) = multiply(inverse(X15),X17) [superposition 98,49]\n",
      "[SA] new: 136. multiply(inverse(X18),X21) = multiply(multiply(inverse(multiply(X19,X18)),X20),multiply(inverse(multiply(inverse(X19),X20)),X21)) [superposition 98,98]\n",
      "[SA] new: 137. multiply(b,multiply(inverse(c),X22)) = multiply(inverse(a),X22) [superposition 98,4]\n",
      "[SA] new: 138. multiply(inverse(X0),multiply(X0,X1)) = multiply(X1,identity) [superposition 98,7]\n",
      "[SA] new: 139. multiply(X3,identity) = multiply(inverse(X2),inverse(multiply(X2,X3))) [superposition 98,3]\n",
      "[SA] new: 140. multiply(X5,inverse(multiply(X4,X5))) = multiply(inverse(X4),identity) [superposition 98,1]\n",
      "[SA] new: 141. multiply(X7,X8) = multiply(inverse(X6),multiply(multiply(X6,X7),X8)) [superposition 98,23]\n",
      "[SA] new: 142. multiply(X10,X11) = multiply(inverse(X9),multiply(inverse(inverse(multiply(X9,X10))),X11)) [superposition 98,49]\n",
      "[SA] new: 143. multiply(inverse(X12),multiply(inverse(multiply(X14,inverse(multiply(X12,X13)))),X15)) = multiply(X13,multiply(inverse(X14),X15)) [superposition 98,98]\n",
      "[SA] new: 144. multiply(inverse(X3),X4) = multiply(inverse(multiply(X3,identity)),X4) [superposition 98,6]\n",
      "[SA] new: 145. multiply(X5,multiply(X6,multiply(inverse(multiply(X7,multiply(X5,X6))),X8))) = multiply(inverse(X7),X8) [superposition 98,8]\n",
      "[SA] new: 146. multiply(inverse(X0),X1) = multiply(inverse(multiply(X0,identity)),X1) [superposition 6,98]\n",
      "[SA] new: 147. multiply(inverse(multiply(X3,X2)),X4) = multiply(inverse(X2),multiply(inverse(X3),X4)) [superposition 23,98]\n",
      "[SA] new: 148. multiply(X5,multiply(multiply(inverse(multiply(X6,X5)),X7),X8)) = multiply(multiply(inverse(X6),X7),X8) [superposition 8,98]\n",
      "[SA] new: 149. multiply(inverse(multiply(X10,inverse(X9))),X11) = multiply(X9,multiply(inverse(X10),X11)) [superposition 49,98]\n",
      "[SA] new: 150. multiply(inverse(X14),X15) = multiply(X12,multiply(X13,multiply(inverse(multiply(X14,multiply(X12,X13))),X15))) [superposition 8,98]\n",
      "[SA] new: 151. multiply(inverse(inverse(X0)),X1) = multiply(X0,multiply(identity,X1)) [forward demodulation 129,10]\n",
      "[SA] new: 152. multiply(X0,X1) = multiply(inverse(inverse(X0)),X1) [forward demodulation 151,6]\n",
      "[SA] new: 153. multiply(X2,multiply(identity,X3)) = multiply(inverse(X2),X3) [forward demodulation 130,10]\n",
      "[SA] new: 154. multiply(X2,X3) = multiply(inverse(X2),X3) [forward demodulation 153,6]\n",
      "[SA] new: 155. multiply(X4,multiply(inverse(X4),X5)) = multiply(identity,X5) [forward demodulation 131,10]\n",
      "[SA] new: 156. multiply(X4,multiply(inverse(X4),X5)) = X5 [forward demodulation 155,6]\n",
      "[SA] new: 157. multiply(multiply(X12,X13),multiply(inverse(X13),X14)) = multiply(X12,X14) [forward demodulation 134,50]\n",
      "[SA] new: 158. multiply(inverse(X0),multiply(X0,X1)) = X1 [forward demodulation 138,1]\n",
      "[SA] new: 159. multiply(inverse(X2),inverse(multiply(X2,X3))) = X3 [forward demodulation 139,1]\n",
      "[SA] new: 160. inverse(X4) = multiply(X5,inverse(multiply(X4,X5))) [forward demodulation 140,1]\n",
      "[SA] new: 161. multiply(X7,X8) = multiply(inverse(X6),multiply(X6,multiply(X7,X8))) [forward demodulation 141,8]\n",
      "[SA] new: 162. multiply(X10,X11) = multiply(inverse(X9),multiply(multiply(X9,X10),X11)) [forward demodulation 142,50]\n",
      "[SA] new: 163. multiply(X10,X11) = multiply(inverse(X9),multiply(X9,multiply(X10,X11))) [forward demodulation 162,8]\n",
      "[SA] new: 164. multiply(X5,multiply(multiply(inverse(multiply(X6,X5)),X7),X8)) = multiply(inverse(X6),multiply(X7,X8)) [forward demodulation 148,8]\n",
      "[SA] new: 165. multiply(inverse(X6),multiply(X7,X8)) = multiply(X5,multiply(inverse(multiply(X6,X5)),multiply(X7,X8))) [forward demodulation 164,8]\n",
      "[SA] new: 166. multiply(identity,X0) = multiply(X1,multiply(inverse(multiply(identity,X1)),X0)) [superposition 100,10]\n",
      "[SA] new: 167. multiply(inverse(X2),X3) = multiply(X4,multiply(inverse(multiply(X2,X4)),X3)) [superposition 100,50]\n",
      "[SA] new: 168. multiply(X0,X1) = multiply(X0,multiply(inverse(identity),X1)) [superposition 100,7]\n",
      "[SA] new: 169. multiply(X2,X3) = multiply(inverse(X2),multiply(inverse(identity),X3)) [superposition 100,3]\n",
      "[SA] new: 170. multiply(X4,X5) = multiply(identity,multiply(inverse(inverse(X4)),X5)) [superposition 100,1]\n",
      "[SA] new: 171. multiply(X6,X8) = multiply(multiply(X6,X7),multiply(inverse(X7),X8)) [superposition 100,23]\n",
      "[SA] new: 172. multiply(X9,X11) = multiply(multiply(inverse(inverse(X9)),X10),multiply(inverse(X10),X11)) [superposition 100,49]\n",
      "[SA] new: 173. multiply(X12,X15) = multiply(multiply(inverse(multiply(X13,inverse(X12))),X14),multiply(inverse(multiply(inverse(X13),X14)),X15)) [superposition 100,98]\n",
      "[SA] new: 174. multiply(X16,X19) = multiply(multiply(inverse(multiply(inverse(X17),inverse(X16))),X18),multiply(inverse(multiply(X17,X18)),X19)) [superposition 100,100]\n",
      "[SA] new: 175. multiply(X0,multiply(inverse(X0),X1)) = multiply(X1,identity) [superposition 100,7]\n",
      "[SA] new: 176. multiply(X3,identity) = multiply(X2,inverse(multiply(inverse(X2),X3))) [superposition 100,3]\n",
      "[SA] new: 177. multiply(X4,identity) = multiply(X5,inverse(multiply(inverse(X4),X5))) [superposition 100,1]\n",
      "[SA] new: 178. multiply(X7,X8) = multiply(X6,multiply(multiply(inverse(X6),X7),X8)) [superposition 100,23]\n",
      "[SA] new: 179. multiply(X10,X11) = multiply(X9,multiply(inverse(inverse(multiply(inverse(X9),X10))),X11)) [superposition 100,49]\n",
      "[SA] new: 180. multiply(X13,multiply(inverse(X14),X15)) = multiply(X12,multiply(inverse(multiply(X14,inverse(multiply(inverse(X12),X13)))),X15)) [superposition 100,98]\n",
      "[SA] new: 181. multiply(X16,multiply(inverse(multiply(inverse(X18),inverse(multiply(inverse(X16),X17)))),X19)) = multiply(X17,multiply(X18,X19)) [superposition 100,100]\n",
      "[SA] new: 182. multiply(X4,X5) = multiply(inverse(inverse(X4)),X5) [superposition 100,98]\n",
      "[SA] new: 183. multiply(X6,X7) = multiply(inverse(multiply(inverse(X6),identity)),X7) [superposition 100,6]\n",
      "[SA] new: 184. multiply(X10,X11) = multiply(X8,multiply(X9,multiply(inverse(multiply(inverse(X10),multiply(X8,X9))),X11))) [superposition 100,8]\n",
      "[SA] new: 185. multiply(X1,X2) = multiply(inverse(inverse(X1)),X2) [superposition 98,100]\n",
      "[SA] new: 186. multiply(X3,X4) = multiply(inverse(multiply(inverse(X3),identity)),X4) [superposition 6,100]\n",
      "[SA] new: 187. multiply(inverse(X5),X8) = multiply(multiply(inverse(multiply(inverse(X6),X5)),X7),multiply(inverse(multiply(X6,X7)),X8)) [superposition 98,100]\n",
      "[SA] new: 188. multiply(inverse(multiply(inverse(X10),X9)),X11) = multiply(inverse(X9),multiply(X10,X11)) [superposition 23,100]\n",
      "[SA] new: 189. multiply(X12,multiply(multiply(inverse(multiply(inverse(X13),X12)),X14),X15)) = multiply(multiply(X13,X14),X15) [superposition 8,100]\n",
      "[SA] new: 190. multiply(X16,multiply(X17,X18)) = multiply(inverse(multiply(inverse(X17),inverse(X16))),X18) [superposition 49,100]\n",
      "[SA] new: 191. multiply(X21,X22) = multiply(X19,multiply(X20,multiply(inverse(multiply(inverse(X21),multiply(X19,X20))),X22))) [superposition 8,100]\n",
      "[SA] new: 192. multiply(inverse(X23),multiply(inverse(multiply(inverse(X25),inverse(multiply(X23,X24)))),X26)) = multiply(X24,multiply(X25,X26)) [superposition 98,100]\n",
      "[SA] new: 193. multiply(identity,X0) = multiply(X1,multiply(inverse(X1),X0)) [forward demodulation 166,6]\n",
      "[SA] new: 194. multiply(identity,X0) = X0 [forward demodulation 193,49]\n",
      "[SA] new: 195. multiply(X0,X1) = multiply(X0,multiply(identity,X1)) [forward demodulation 168,10]\n",
      "[SA] new: 196. multiply(X2,X3) = multiply(inverse(X2),multiply(identity,X3)) [forward demodulation 169,10]\n",
      "[SA] new: 197. multiply(X2,X3) = multiply(inverse(X2),X3) [forward demodulation 196,6]\n",
      "[SA] new: 198. multiply(X4,X5) = multiply(identity,multiply(X4,X5)) [forward demodulation 170,50]\n",
      "[SA] new: 199. multiply(X9,X11) = multiply(multiply(X9,X10),multiply(inverse(X10),X11)) [forward demodulation 172,50]\n",
      "[SA] new: 200. multiply(X0,multiply(inverse(X0),X1)) = X1 [forward demodulation 175,1]\n",
      "[SA] new: 201. multiply(X2,inverse(multiply(inverse(X2),X3))) = X3 [forward demodulation 176,1]\n",
      "[SA] new: 202. multiply(X5,inverse(multiply(inverse(X4),X5))) = X4 [forward demodulation 177,1]\n",
      "[SA] new: 203. multiply(X7,X8) = multiply(X6,multiply(inverse(X6),multiply(X7,X8))) [forward demodulation 178,8]\n",
      "[SA] new: 204. multiply(X10,X11) = multiply(X9,multiply(multiply(inverse(X9),X10),X11)) [forward demodulation 179,50]\n",
      "[SA] new: 205. multiply(X10,X11) = multiply(X9,multiply(inverse(X9),multiply(X10,X11))) [forward demodulation 204,8]\n",
      "[SA] new: 206. multiply(X6,X7) = multiply(inverse(inverse(X6)),X7) [forward demodulation 183,1]\n",
      "[SA] new: 207. multiply(X3,X4) = multiply(inverse(inverse(X3)),X4) [forward demodulation 186,1]\n",
      "[SA] new: 208. multiply(X12,multiply(multiply(inverse(multiply(inverse(X13),X12)),X14),X15)) = multiply(X13,multiply(X14,X15)) [forward demodulation 189,8]\n",
      "[SA] new: 209. multiply(X13,multiply(X14,X15)) = multiply(X12,multiply(inverse(multiply(inverse(X13),X12)),multiply(X14,X15))) [forward demodulation 208,8]\n",
      "[SA] new: 210. multiply(X1,multiply(inverse(X0),multiply(X0,multiply(inverse(multiply(X1,identity)),X2)))) = X2 [superposition 121,7]\n",
      "[SA] new: 211. multiply(X4,multiply(X3,multiply(X3,multiply(inverse(multiply(X4,identity)),X5)))) = X5 [superposition 121,3]\n",
      "[SA] new: 212. multiply(X7,multiply(identity,multiply(X6,multiply(inverse(multiply(X7,X6)),X8)))) = X8 [superposition 121,6]\n",
      "[SA] new: 213. multiply(X12,multiply(multiply(X9,X10),multiply(X11,multiply(inverse(multiply(X12,multiply(X9,multiply(X10,X11)))),X13)))) = X13 [superposition 121,8]\n",
      "[SA] new: 214. multiply(X15,multiply(X14,multiply(identity,multiply(inverse(multiply(X15,X14)),X16)))) = X16 [superposition 121,1]\n",
      "[SA] new: 215. multiply(X19,multiply(inverse(X17),multiply(multiply(X17,X18),multiply(inverse(multiply(X19,X18)),X20)))) = X20 [superposition 121,23]\n",
      "[SA] new: 216. multiply(X23,multiply(X21,multiply(multiply(inverse(X21),X22),multiply(inverse(multiply(X23,X22)),X24)))) = X24 [superposition 121,49]\n",
      "[SA] new: 217. multiply(X28,multiply(X25,multiply(multiply(inverse(multiply(X26,X25)),X27),multiply(inverse(multiply(X28,multiply(inverse(X26),X27))),X29)))) = X29 [superposition 121,98]\n",
      "[SA] new: 218. multiply(X33,multiply(X30,multiply(multiply(inverse(multiply(inverse(X31),X30)),X32),multiply(inverse(multiply(X33,multiply(X31,X32))),X34)))) = X34 [superposition 121,100]\n",
      "[SA] new: 219. multiply(X39,multiply(X35,multiply(multiply(X36,multiply(X37,multiply(inverse(multiply(X35,multiply(X36,X37))),X38))),multiply(inverse(multiply(X39,X38)),X40)))) = X40 [superposition 121,121]\n",
      "[SA] new: 220. multiply(X41,multiply(a,multiply(b,multiply(inverse(multiply(X41,c)),X42)))) = X42 [superposition 121,4]\n",
      "[SA] new: 221. multiply(inverse(X0),multiply(X0,multiply(X1,multiply(inverse(X1),X2)))) = X2 [superposition 121,23]\n",
      "[SA] new: 222. multiply(X3,multiply(inverse(X3),multiply(X4,multiply(inverse(X4),X5)))) = X5 [superposition 121,49]\n",
      "[SA] new: 223. multiply(X6,multiply(inverse(multiply(X7,X6)),multiply(X8,multiply(inverse(multiply(inverse(X7),X8)),X9)))) = X9 [superposition 121,98]\n",
      "[SA] new: 224. multiply(X10,multiply(inverse(multiply(inverse(X11),X10)),multiply(X12,multiply(inverse(multiply(X11,X12)),X13)))) = X13 [superposition 121,100]\n",
      "[SA] new: 225. multiply(X14,multiply(X15,multiply(multiply(X16,multiply(inverse(multiply(X14,multiply(X15,X16))),X17)),multiply(inverse(X17),X18)))) = X18 [superposition 121,121]\n",
      "[SA] new: 226. multiply(inverse(multiply(X19,X20)),multiply(X19,multiply(X20,multiply(inverse(identity),X21)))) = X21 [superposition 121,7]\n",
      "[SA] new: 227. multiply(multiply(X22,X23),multiply(X22,multiply(X23,multiply(inverse(identity),X24)))) = X24 [superposition 121,3]\n",
      "[SA] new: 228. multiply(identity,multiply(X25,multiply(X26,multiply(inverse(multiply(X25,X26)),X27)))) = X27 [superposition 121,6]\n",
      "[SA] new: 229. multiply(multiply(X28,X29),multiply(X30,multiply(X31,multiply(inverse(multiply(X28,multiply(X29,multiply(X30,X31)))),X32)))) = X32 [superposition 121,8]\n",
      "[SA] new: 230. multiply(X0,multiply(X1,X2)) = multiply(X0,multiply(X1,multiply(X2,identity))) [superposition 121,7]\n",
      "[SA] new: 231. inverse(multiply(X3,multiply(X4,X5))) = multiply(X3,multiply(X4,multiply(X5,identity))) [superposition 121,3]\n",
      "[SA] new: 232. identity = multiply(X6,multiply(X7,multiply(X8,inverse(multiply(X6,multiply(X7,X8)))))) [superposition 121,1]\n",
      "[SA] new: 233. multiply(multiply(X9,multiply(X10,X11)),X12) = multiply(X9,multiply(X10,multiply(X11,X12))) [superposition 121,23]\n",
      "[SA] new: 234. multiply(inverse(inverse(multiply(X13,multiply(X14,X15)))),X16) = multiply(X13,multiply(X14,multiply(X15,X16))) [superposition 121,49]\n",
      "[SA] new: 235. multiply(inverse(multiply(X20,inverse(multiply(X17,multiply(X18,X19))))),X21) = multiply(X17,multiply(X18,multiply(X19,multiply(inverse(X20),X21)))) [superposition 121,98]\n",
      "[SA] new: 236. multiply(inverse(multiply(inverse(X25),inverse(multiply(X22,multiply(X23,X24))))),X26) = multiply(X22,multiply(X23,multiply(X24,multiply(X25,X26)))) [superposition 121,100]\n",
      "[SA] new: 237. multiply(X30,multiply(X31,multiply(inverse(multiply(inverse(multiply(X27,multiply(X28,X29))),multiply(X30,X31))),X32))) = multiply(X27,multiply(X28,multiply(X29,X32))) [superposition 121,121]\n",
      "[SA] new: 238. multiply(X3,multiply(inverse(multiply(X0,multiply(inverse(multiply(X1,multiply(X2,X0))),X3))),X4)) = multiply(X1,multiply(X2,X4)) [superposition 121,121]\n",
      "[SA] new: 239. multiply(X5,multiply(X6,multiply(inverse(multiply(X5,multiply(X6,identity))),X7))) = X7 [superposition 121,6]\n",
      "[SA] new: 240. multiply(X10,multiply(X11,multiply(X8,multiply(X9,multiply(inverse(multiply(X10,multiply(X11,multiply(X8,X9)))),X12))))) = X12 [superposition 121,8]\n",
      "[SA] new: 241. multiply(X2,X3) = multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X2,multiply(X0,X1)))))),X3) [superposition 121,121]\n",
      "[SA] new: 242. multiply(X5,multiply(inverse(multiply(X5,multiply(inverse(X4),X4))),X6)) = X6 [superposition 121,23]\n",
      "[SA] new: 243. multiply(X8,multiply(inverse(multiply(X8,multiply(X7,inverse(X7)))),X9)) = X9 [superposition 121,49]\n",
      "[SA] new: 244. multiply(X12,multiply(inverse(X11),multiply(inverse(multiply(X12,multiply(X10,inverse(multiply(X11,X10))))),X13))) = X13 [superposition 121,98]\n",
      "[SA] new: 245. multiply(X16,multiply(X15,multiply(inverse(multiply(X16,multiply(X14,inverse(multiply(inverse(X15),X14))))),X17))) = X17 [superposition 121,100]\n",
      "[SA] new: 246. multiply(X19,multiply(X18,multiply(inverse(multiply(X19,multiply(identity,X18))),X20))) = X20 [superposition 121,6]\n",
      "[SA] new: 247. multiply(X24,multiply(X21,multiply(X22,multiply(X23,multiply(inverse(multiply(X24,multiply(multiply(X21,X22),X23))),X25))))) = X25 [superposition 121,8]\n",
      "[SA] new: 248. multiply(X5,multiply(inverse(multiply(inverse(X4),multiply(X4,X5))),X6)) = X6 [superposition 121,23]\n",
      "[SA] new: 249. multiply(X8,multiply(inverse(multiply(X7,multiply(inverse(X7),X8))),X9)) = X9 [superposition 121,49]\n",
      "[SA] new: 250. multiply(inverse(X11),multiply(X12,multiply(inverse(multiply(X10,multiply(inverse(multiply(X11,X10)),X12))),X13))) = X13 [superposition 121,98]\n",
      "[SA] new: 251. multiply(X15,multiply(X16,multiply(inverse(multiply(X14,multiply(inverse(multiply(inverse(X15),X14)),X16))),X17))) = X17 [superposition 121,100]\n",
      "[SA] new: 252. multiply(X18,multiply(X19,multiply(inverse(multiply(identity,multiply(X18,X19))),X20))) = X20 [superposition 121,6]\n",
      "[SA] new: 253. multiply(X21,multiply(X22,multiply(X23,multiply(X24,multiply(inverse(multiply(multiply(X21,X22),multiply(X23,X24))),X25))))) = X25 [superposition 121,8]\n",
      "[SA] new: 254. multiply(X1,multiply(inverse(multiply(inverse(X0),multiply(X0,X1))),X2)) = X2 [superposition 23,121]\n",
      "[SA] new: 255. multiply(X4,multiply(inverse(multiply(X3,multiply(inverse(X3),X4))),X5)) = X5 [superposition 49,121]\n",
      "[SA] new: 256. multiply(inverse(X7),multiply(X8,multiply(inverse(multiply(X6,multiply(inverse(multiply(X7,X6)),X8))),X9))) = X9 [superposition 98,121]\n",
      "[SA] new: 257. multiply(X11,multiply(X12,multiply(inverse(multiply(X10,multiply(inverse(multiply(inverse(X11),X10)),X12))),X13))) = X13 [superposition 100,121]\n",
      "[SA] new: 258. multiply(X14,multiply(X15,multiply(inverse(multiply(identity,multiply(X14,X15))),X16))) = X16 [superposition 6,121]\n",
      "[SA] new: 259. multiply(inverse(X17),X21) = multiply(multiply(X18,multiply(X19,multiply(inverse(multiply(X17,multiply(X18,X19))),X20))),multiply(inverse(X20),X21)) [superposition 98,121]\n",
      "[SA] new: 260. multiply(X23,multiply(X24,multiply(inverse(multiply(X22,multiply(X23,X24))),X25))) = multiply(inverse(X22),X25) [superposition 23,121]\n",
      "[SA] new: 261. multiply(X26,multiply(multiply(X27,multiply(X28,multiply(inverse(multiply(X26,multiply(X27,X28))),X29))),X30)) = multiply(X29,X30) [superposition 8,121]\n",
      "[SA] new: 262. multiply(X31,X35) = multiply(multiply(X32,multiply(X33,multiply(inverse(multiply(inverse(X31),multiply(X32,X33))),X34))),multiply(inverse(X34),X35)) [superposition 100,121]\n",
      "[SA] new: 263. multiply(X37,multiply(X38,multiply(inverse(multiply(inverse(X36),multiply(X37,X38))),X39))) = multiply(X36,X39) [superposition 49,121]\n",
      "[SA] new: 264. multiply(X40,multiply(X41,multiply(X42,multiply(X43,multiply(inverse(multiply(multiply(X40,X41),multiply(X42,X43))),X44))))) = X44 [superposition 8,121]\n",
      "[SA] new: 265. multiply(inverse(X45),multiply(X47,multiply(X48,multiply(inverse(multiply(inverse(multiply(X45,X46)),multiply(X47,X48))),X49)))) = multiply(X46,X49) [superposition 98,121]\n",
      "[SA] new: 266. multiply(X50,multiply(X52,multiply(X53,multiply(inverse(multiply(inverse(multiply(inverse(X50),X51)),multiply(X52,X53))),X54)))) = multiply(X51,X54) [superposition 100,121]\n",
      "[SA] new: 267. multiply(X1,multiply(inverse(multiply(X1,identity)),X2)) = X2 [forward demodulation 210,23]\n",
      "[SA] new: 268. multiply(X1,multiply(inverse(X1),X2)) = X2 [forward demodulation 267,1]\n",
      "[SA] new: 269. multiply(X4,multiply(X3,multiply(X3,multiply(inverse(X4),X5)))) = X5 [forward demodulation 211,1]\n",
      "[SA] new: 270. multiply(X7,multiply(X6,multiply(inverse(multiply(X7,X6)),X8))) = X8 [forward demodulation 212,6]\n",
      "[SA] new: 271. multiply(X7,multiply(inverse(X7),X8)) = X8 [forward demodulation 270,98]\n",
      "[SA] new: 272. multiply(X12,multiply(X9,multiply(X10,multiply(X11,multiply(inverse(multiply(X12,multiply(X9,multiply(X10,X11)))),X13))))) = X13 [forward demodulation 213,8]\n",
      "[SA] new: 273. multiply(X15,multiply(X14,multiply(inverse(multiply(X15,X14)),X16))) = X16 [forward demodulation 214,6]\n",
      "[SA] new: 274. multiply(X15,multiply(inverse(X15),X16)) = X16 [forward demodulation 273,98]\n",
      "[SA] new: 275. multiply(X19,multiply(inverse(X17),multiply(X17,multiply(X18,multiply(inverse(multiply(X19,X18)),X20))))) = X20 [forward demodulation 215,8]\n",
      "[SA] new: 276. multiply(X19,multiply(X18,multiply(inverse(multiply(X19,X18)),X20))) = X20 [forward demodulation 275,23]\n",
      "[SA] new: 277. multiply(X19,multiply(inverse(X19),X20)) = X20 [forward demodulation 276,98]\n",
      "[SA] new: 278. multiply(X23,multiply(X21,multiply(inverse(X21),multiply(X22,multiply(inverse(multiply(X23,X22)),X24))))) = X24 [forward demodulation 216,8]\n",
      "[SA] new: 279. multiply(X23,multiply(X22,multiply(inverse(multiply(X23,X22)),X24))) = X24 [forward demodulation 278,49]\n",
      "[SA] new: 280. multiply(X23,multiply(inverse(X23),X24)) = X24 [forward demodulation 279,98]\n",
      "[SA] new: 281. multiply(X28,multiply(X25,multiply(inverse(multiply(X26,X25)),multiply(X27,multiply(inverse(multiply(X28,multiply(inverse(X26),X27))),X29))))) = X29 [forward demodulation 217,8]\n",
      "[SA] new: 282. multiply(X28,multiply(inverse(X26),multiply(X27,multiply(inverse(multiply(X28,multiply(inverse(X26),X27))),X29)))) = X29 [forward demodulation 281,98]\n",
      "[SA] new: 283. multiply(X33,multiply(X30,multiply(inverse(multiply(inverse(X31),X30)),multiply(X32,multiply(inverse(multiply(X33,multiply(X31,X32))),X34))))) = X34 [forward demodulation 218,8]\n",
      "[SA] new: 284. multiply(X33,multiply(inverse(inverse(X31)),multiply(X32,multiply(inverse(multiply(X33,multiply(X31,X32))),X34)))) = X34 [forward demodulation 283,98]\n",
      "[SA] new: 285. multiply(X33,multiply(X31,multiply(X32,multiply(inverse(multiply(X33,multiply(X31,X32))),X34)))) = X34 [forward demodulation 284,50]\n",
      "[SA] new: 286. multiply(X39,multiply(X35,multiply(X36,multiply(multiply(X37,multiply(inverse(multiply(X35,multiply(X36,X37))),X38)),multiply(inverse(multiply(X39,X38)),X40))))) = X40 [forward demodulation 219,8]\n",
      "[SA] new: 287. multiply(X39,multiply(X35,multiply(X36,multiply(X37,multiply(multiply(inverse(multiply(X35,multiply(X36,X37))),X38),multiply(inverse(multiply(X39,X38)),X40)))))) = X40 [forward demodulation 286,8]\n",
      "[SA] new: 288. multiply(X39,multiply(X35,multiply(X36,multiply(X37,multiply(inverse(multiply(X35,multiply(X36,X37))),multiply(X38,multiply(inverse(multiply(X39,X38)),X40))))))) = X40 [forward demodulation 287,8]\n",
      "[SA] new: 289. multiply(X39,multiply(X38,multiply(inverse(multiply(X39,X38)),X40))) = X40 [forward demodulation 288,121]\n",
      "[SA] new: 290. multiply(X39,multiply(inverse(X39),X40)) = X40 [forward demodulation 289,98]\n",
      "[SA] new: 291. multiply(inverse(X0),multiply(X0,X2)) = X2 [forward demodulation 221,49]\n",
      "[SA] new: 292. multiply(X3,multiply(inverse(X3),X5)) = X5 [forward demodulation 222,49]\n",
      "[SA] new: 293. multiply(X6,multiply(inverse(multiply(X7,X6)),multiply(inverse(inverse(X7)),X9))) = X9 [forward demodulation 223,98]\n",
      "[SA] new: 294. multiply(X6,multiply(inverse(multiply(X7,X6)),multiply(X7,X9))) = X9 [forward demodulation 293,50]\n",
      "[SA] new: 295. multiply(X10,multiply(inverse(multiply(inverse(X11),X10)),multiply(inverse(X11),X13))) = X13 [forward demodulation 224,98]\n",
      "[SA] new: 296. multiply(X14,multiply(X15,multiply(X16,multiply(multiply(inverse(multiply(X14,multiply(X15,X16))),X17),multiply(inverse(X17),X18))))) = X18 [forward demodulation 225,8]\n",
      "[SA] new: 297. multiply(X14,multiply(X15,multiply(X16,multiply(inverse(multiply(X14,multiply(X15,X16))),multiply(X17,multiply(inverse(X17),X18)))))) = X18 [forward demodulation 296,8]\n",
      "[SA] new: 298. multiply(X14,multiply(X15,multiply(X16,multiply(inverse(multiply(X14,multiply(X15,X16))),X18)))) = X18 [forward demodulation 297,49]\n",
      "[SA] new: 299. multiply(inverse(multiply(X19,X20)),multiply(X19,multiply(X20,multiply(identity,X21)))) = X21 [forward demodulation 226,10]\n",
      "[SA] new: 300. multiply(inverse(multiply(X19,X20)),multiply(X19,multiply(X20,X21))) = X21 [forward demodulation 299,6]\n",
      "[SA] new: 301. multiply(multiply(X22,X23),multiply(X22,multiply(X23,multiply(identity,X24)))) = X24 [forward demodulation 227,10]\n",
      "[SA] new: 302. multiply(multiply(X22,X23),multiply(X22,multiply(X23,X24))) = X24 [forward demodulation 301,6]\n",
      "[SA] new: 303. multiply(identity,multiply(X25,multiply(inverse(X25),X27))) = X27 [forward demodulation 228,98]\n",
      "[SA] new: 304. multiply(identity,X27) = X27 [forward demodulation 303,49]\n",
      "[SA] new: 305. multiply(X3,multiply(X4,X5)) = inverse(multiply(X3,multiply(X4,X5))) [forward demodulation 231,1]\n",
      "[SA] new: 306. multiply(X13,multiply(X14,multiply(X15,X16))) = multiply(multiply(X13,multiply(X14,X15)),X16) [forward demodulation 234,50]\n",
      "[SA] new: 307. multiply(X5,multiply(X6,multiply(inverse(multiply(X5,X6)),X7))) = X7 [forward demodulation 239,1]\n",
      "[SA] new: 308. multiply(X5,multiply(inverse(X5),X7)) = X7 [forward demodulation 307,98]\n",
      "[SA] new: 309. multiply(X5,multiply(inverse(multiply(X5,identity)),X6)) = X6 [forward demodulation 242,7]\n",
      "[SA] new: 310. multiply(X5,multiply(inverse(X5),X6)) = X6 [forward demodulation 309,1]\n",
      "[SA] new: 311. multiply(X19,multiply(X18,multiply(inverse(multiply(X19,X18)),X20))) = X20 [forward demodulation 246,6]\n",
      "[SA] new: 312. multiply(X19,multiply(inverse(X19),X20)) = X20 [forward demodulation 311,98]\n",
      "[SA] new: 313. multiply(X24,multiply(X21,multiply(X22,multiply(X23,multiply(inverse(multiply(X24,multiply(X21,multiply(X22,X23)))),X25))))) = X25 [forward demodulation 247,8]\n",
      "[SA] new: 314. multiply(X5,multiply(inverse(X5),X6)) = X6 [forward demodulation 248,23]\n",
      "[SA] new: 315. multiply(X8,multiply(inverse(X8),X9)) = X9 [forward demodulation 249,49]\n",
      "[SA] new: 316. multiply(inverse(X11),multiply(X12,multiply(inverse(multiply(inverse(X11),X12)),X13))) = X13 [forward demodulation 250,98]\n",
      "[SA] new: 317. multiply(inverse(X11),multiply(inverse(inverse(X11)),X13)) = X13 [forward demodulation 316,98]\n",
      "[SA] new: 318. multiply(X15,multiply(X16,multiply(inverse(multiply(inverse(inverse(X15)),X16)),X17))) = X17 [forward demodulation 251,98]\n",
      "[SA] new: 319. multiply(X15,multiply(inverse(inverse(inverse(X15))),X17)) = X17 [forward demodulation 318,98]\n",
      "[SA] new: 320. multiply(X15,multiply(inverse(X15),X17)) = X17 [forward demodulation 319,50]\n",
      "[SA] new: 321. multiply(X18,multiply(X19,multiply(inverse(multiply(X18,X19)),X20))) = X20 [forward demodulation 252,6]\n",
      "[SA] new: 322. multiply(X18,multiply(inverse(X18),X20)) = X20 [forward demodulation 321,98]\n",
      "[SA] new: 323. multiply(X21,multiply(X22,multiply(X23,multiply(X24,multiply(inverse(multiply(X21,multiply(X22,multiply(X23,X24)))),X25))))) = X25 [forward demodulation 253,8]\n",
      "[SA] new: 324. multiply(X1,multiply(inverse(X1),X2)) = X2 [forward demodulation 254,23]\n",
      "[SA] new: 325. multiply(X4,multiply(inverse(X4),X5)) = X5 [forward demodulation 255,49]\n",
      "[SA] new: 326. multiply(inverse(X7),multiply(X8,multiply(inverse(multiply(inverse(X7),X8)),X9))) = X9 [forward demodulation 256,98]\n",
      "[SA] new: 327. multiply(inverse(X7),multiply(inverse(inverse(X7)),X9)) = X9 [forward demodulation 326,98]\n",
      "[SA] new: 328. multiply(X11,multiply(X12,multiply(inverse(multiply(inverse(inverse(X11)),X12)),X13))) = X13 [forward demodulation 257,98]\n",
      "[SA] new: 329. multiply(X11,multiply(inverse(inverse(inverse(X11))),X13)) = X13 [forward demodulation 328,98]\n",
      "[SA] new: 330. multiply(X11,multiply(inverse(X11),X13)) = X13 [forward demodulation 329,50]\n",
      "[SA] new: 331. multiply(X14,multiply(X15,multiply(inverse(multiply(X14,X15)),X16))) = X16 [forward demodulation 258,6]\n",
      "[SA] new: 332. multiply(X14,multiply(inverse(X14),X16)) = X16 [forward demodulation 331,98]\n",
      "[SA] new: 333. multiply(X29,X30) = multiply(X26,multiply(X27,multiply(multiply(X28,multiply(inverse(multiply(X26,multiply(X27,X28))),X29)),X30))) [forward demodulation 261,8]\n",
      "[SA] new: 334. multiply(X29,X30) = multiply(X26,multiply(X27,multiply(X28,multiply(multiply(inverse(multiply(X26,multiply(X27,X28))),X29),X30)))) [forward demodulation 333,8]\n",
      "[SA] new: 335. multiply(X29,X30) = multiply(X26,multiply(X27,multiply(X28,multiply(inverse(multiply(X26,multiply(X27,X28))),multiply(X29,X30))))) [forward demodulation 334,8]\n",
      "[SA] new: 336. multiply(X40,multiply(X41,multiply(X42,multiply(X43,multiply(inverse(multiply(X40,multiply(X41,multiply(X42,X43)))),X44))))) = X44 [forward demodulation 264,8]\n",
      "[SA] new: 337. multiply(X0,multiply(inverse(multiply(X0,multiply(identity,identity))),X1)) = X1 [superposition 243,10]\n",
      "[SA] new: 338. multiply(X3,multiply(inverse(multiply(X3,multiply(inverse(X2),X2))),X4)) = X4 [superposition 243,50]\n",
      "[SA] new: 339. multiply(X0,multiply(inverse(multiply(X0,inverse(identity))),X1)) = X1 [superposition 243,6]\n",
      "[SA] new: 340. multiply(X4,multiply(inverse(multiply(X4,multiply(X2,multiply(X3,inverse(multiply(X2,X3)))))),X5)) = X5 [superposition 243,8]\n",
      "[SA] new: 341. multiply(inverse(X0),multiply(inverse(inverse(X0)),X1)) = X1 [superposition 243,23]\n",
      "[SA] new: 342. multiply(X2,multiply(inverse(inverse(inverse(X2))),X3)) = X3 [superposition 243,49]\n",
      "[SA] new: 343. multiply(X4,multiply(inverse(multiply(inverse(X5),inverse(inverse(multiply(X5,X4))))),X6)) = X6 [superposition 243,98]\n",
      "[SA] new: 344. multiply(X7,multiply(inverse(multiply(X8,inverse(inverse(multiply(inverse(X8),X7))))),X9)) = X9 [superposition 243,100]\n",
      "[SA] new: 345. multiply(X10,multiply(inverse(inverse(inverse(multiply(X10,multiply(X11,inverse(X11)))))),X12)) = X12 [superposition 243,243]\n",
      "[SA] new: 346. multiply(inverse(multiply(X13,inverse(X13))),multiply(inverse(identity),X14)) = X14 [superposition 243,7]\n",
      "[SA] new: 347. multiply(multiply(X15,inverse(X15)),multiply(inverse(identity),X16)) = X16 [superposition 243,3]\n",
      "[SA] new: 348. multiply(identity,multiply(inverse(multiply(X17,inverse(X17))),X18)) = X18 [superposition 243,6]\n",
      "[SA] new: 349. multiply(multiply(X19,X20),multiply(inverse(multiply(X19,multiply(X20,multiply(X21,inverse(X21))))),X22)) = X22 [superposition 243,8]\n",
      "[SA] new: 350. multiply(X0,identity) = multiply(X0,multiply(X1,inverse(X1))) [superposition 243,7]\n",
      "[SA] new: 351. multiply(X2,identity) = inverse(multiply(X2,multiply(X3,inverse(X3)))) [superposition 243,3]\n",
      "[SA] new: 352. identity = multiply(X4,inverse(multiply(X4,multiply(X5,inverse(X5))))) [superposition 243,1]\n",
      "[SA] new: 353. multiply(X6,X8) = multiply(multiply(X6,multiply(X7,inverse(X7))),X8) [superposition 243,23]\n",
      "[SA] new: 354. multiply(X9,X11) = multiply(inverse(inverse(multiply(X9,multiply(X10,inverse(X10))))),X11) [superposition 243,49]\n",
      "[SA] new: 355. multiply(inverse(multiply(X14,inverse(multiply(X12,multiply(X13,inverse(X13)))))),X15) = multiply(X12,multiply(inverse(X14),X15)) [superposition 243,98]\n",
      "[SA] new: 356. multiply(inverse(multiply(inverse(X18),inverse(multiply(X16,multiply(X17,inverse(X17)))))),X19) = multiply(X16,multiply(X18,X19)) [superposition 243,100]\n",
      "[SA] new: 357. multiply(inverse(multiply(inverse(multiply(X20,multiply(X21,inverse(X21)))),multiply(X22,inverse(X22)))),X23) = multiply(X20,X23) [superposition 243,243]\n",
      "[SA] new: 358. multiply(X26,multiply(X27,multiply(inverse(multiply(inverse(multiply(X24,multiply(X25,inverse(X25)))),multiply(X26,X27))),X28))) = multiply(X24,X28) [superposition 243,121]\n",
      "[SA] new: 359. multiply(inverse(multiply(X3,inverse(X3))),X4) = X4 [superposition 243,98]\n",
      "[SA] new: 360. multiply(X7,multiply(inverse(multiply(X5,multiply(inverse(multiply(X5,multiply(X6,inverse(X6)))),X7))),X8)) = X8 [superposition 243,121]\n",
      "[SA] new: 361. multiply(inverse(multiply(identity,multiply(X9,inverse(X9)))),X10) = X10 [superposition 243,6]\n",
      "[SA] new: 362. multiply(X11,multiply(X12,multiply(inverse(multiply(multiply(X11,X12),multiply(X13,inverse(X13)))),X14))) = X14 [superposition 243,8]\n",
      "[SA] new: 363. multiply(inverse(multiply(X0,inverse(X0))),X1) = X1 [superposition 98,243]\n",
      "[SA] new: 364. multiply(X2,multiply(inverse(multiply(X2,multiply(X3,inverse(X3)))),multiply(X4,multiply(inverse(X4),X5)))) = X5 [superposition 121,243]\n",
      "[SA] new: 365. multiply(inverse(X6),multiply(X6,X7)) = X7 [superposition 121,243]\n",
      "[SA] new: 366. multiply(X10,multiply(inverse(multiply(X10,multiply(X8,inverse(multiply(X8,multiply(X9,inverse(X9))))))),X11)) = X11 [superposition 121,243]\n",
      "[SA] new: 367. multiply(X14,multiply(inverse(multiply(X12,multiply(inverse(multiply(X12,multiply(X13,inverse(X13)))),X14))),X15)) = X15 [superposition 121,243]\n",
      "[SA] new: 368. multiply(inverse(multiply(identity,multiply(X16,inverse(X16)))),X17) = X17 [superposition 6,243]\n",
      "[SA] new: 369. multiply(X21,multiply(X18,multiply(multiply(inverse(multiply(X18,multiply(X19,inverse(X19)))),X20),multiply(inverse(multiply(X21,X20)),X22)))) = X22 [superposition 121,243]\n",
      "[SA] new: 370. multiply(inverse(X23),X26) = multiply(multiply(inverse(multiply(X23,multiply(X24,inverse(X24)))),X25),multiply(inverse(X25),X26)) [superposition 98,243]\n",
      "[SA] new: 371. multiply(inverse(multiply(X27,multiply(X28,inverse(X28)))),X29) = multiply(inverse(X27),X29) [superposition 23,243]\n",
      "[SA] new: 372. multiply(X32,X33) = multiply(X30,multiply(multiply(inverse(multiply(X30,multiply(X31,inverse(X31)))),X32),X33)) [superposition 8,243]\n",
      "[SA] new: 373. multiply(X34,X37) = multiply(multiply(inverse(multiply(inverse(X34),multiply(X35,inverse(X35)))),X36),multiply(inverse(X36),X37)) [superposition 100,243]\n",
      "[SA] new: 374. multiply(inverse(multiply(inverse(X38),multiply(X39,inverse(X39)))),X40) = multiply(X38,X40) [superposition 49,243]\n",
      "[SA] new: 375. multiply(X41,multiply(X42,multiply(inverse(multiply(multiply(X41,X42),multiply(X43,inverse(X43)))),X44))) = X44 [superposition 8,243]\n",
      "[SA] new: 376. multiply(inverse(X45),multiply(inverse(multiply(inverse(multiply(X45,X46)),multiply(X47,inverse(X47)))),X48)) = multiply(X46,X48) [superposition 98,243]\n",
      "[SA] new: 377. multiply(X49,multiply(inverse(multiply(inverse(multiply(inverse(X49),X50)),multiply(X51,inverse(X51)))),X52)) = multiply(X50,X52) [superposition 100,243]\n",
      "[SA] new: 378. multiply(inverse(multiply(inverse(multiply(X53,multiply(X54,X55))),multiply(X56,inverse(X56)))),X57) = multiply(X53,multiply(X54,multiply(X55,X57))) [superposition 121,243]\n",
      "[SA] new: 379. multiply(X0,multiply(inverse(multiply(X0,identity)),X1)) = X1 [forward demodulation 337,3]\n",
      "[SA] new: 380. multiply(X0,multiply(inverse(X0),X1)) = X1 [forward demodulation 379,1]\n",
      "[SA] new: 381. multiply(X3,multiply(inverse(multiply(X3,identity)),X4)) = X4 [forward demodulation 338,7]\n",
      "[SA] new: 382. multiply(X3,multiply(inverse(X3),X4)) = X4 [forward demodulation 381,1]\n",
      "[SA] new: 383. multiply(X0,multiply(inverse(multiply(X0,identity)),X1)) = X1 [forward demodulation 339,10]\n",
      "[SA] new: 384. multiply(X0,multiply(inverse(X0),X1)) = X1 [forward demodulation 383,1]\n",
      "[SA] new: 385. multiply(X2,multiply(inverse(X2),X3)) = X3 [forward demodulation 342,50]\n",
      "[SA] new: 386. multiply(X4,multiply(inverse(multiply(inverse(X5),multiply(X5,X4))),X6)) = X6 [forward demodulation 343,50]\n",
      "[SA] new: 387. multiply(X4,multiply(inverse(X4),X6)) = X6 [forward demodulation 386,23]\n",
      "[SA] new: 388. multiply(X7,multiply(inverse(multiply(X8,multiply(inverse(X8),X7))),X9)) = X9 [forward demodulation 344,50]\n",
      "[SA] new: 389. multiply(X7,multiply(inverse(X7),X9)) = X9 [forward demodulation 388,49]\n",
      "[SA] new: 390. multiply(X10,multiply(inverse(multiply(X10,multiply(X11,inverse(X11)))),X12)) = X12 [forward demodulation 345,50]\n",
      "[SA] new: 391. multiply(inverse(multiply(X13,inverse(X13))),multiply(identity,X14)) = X14 [forward demodulation 346,10]\n",
      "[SA] new: 392. multiply(inverse(multiply(X13,inverse(X13))),X14) = X14 [forward demodulation 391,6]\n",
      "[SA] new: 393. multiply(multiply(X15,inverse(X15)),multiply(identity,X16)) = X16 [forward demodulation 347,10]\n",
      "[SA] new: 394. multiply(multiply(X15,inverse(X15)),X16) = X16 [forward demodulation 393,6]\n",
      "[SA] new: 395. multiply(X0,multiply(X1,inverse(X1))) = X0 [forward demodulation 350,1]\n",
      "[SA] new: 396. inverse(multiply(X2,multiply(X3,inverse(X3)))) = X2 [forward demodulation 351,1]\n",
      "[SA] new: 397. multiply(X9,X11) = multiply(multiply(X9,multiply(X10,inverse(X10))),X11) [forward demodulation 354,50]\n",
      "[SA] new: 398. multiply(X7,multiply(inverse(X7),X8)) = X8 [forward demodulation 360,243]\n",
      "[SA] new: 399. multiply(inverse(multiply(X9,inverse(X9))),X10) = X10 [forward demodulation 361,6]\n",
      "[SA] new: 400. multiply(X11,multiply(X12,multiply(inverse(multiply(X11,multiply(X12,multiply(X13,inverse(X13))))),X14))) = X14 [forward demodulation 362,8]\n",
      "[SA] new: 401. multiply(X2,multiply(inverse(multiply(X2,multiply(X3,inverse(X3)))),X5)) = X5 [forward demodulation 364,49]\n",
      "[SA] new: 402. multiply(X14,multiply(inverse(X14),X15)) = X15 [forward demodulation 367,243]\n",
      "[SA] new: 403. multiply(inverse(multiply(X16,inverse(X16))),X17) = X17 [forward demodulation 368,6]\n",
      "[SA] new: 404. multiply(X21,multiply(X18,multiply(inverse(multiply(X18,multiply(X19,inverse(X19)))),multiply(X20,multiply(inverse(multiply(X21,X20)),X22))))) = X22 [forward demodulation 369,8]\n",
      "[SA] new: 405. multiply(X21,multiply(X20,multiply(inverse(multiply(X21,X20)),X22))) = X22 [forward demodulation 404,243]\n",
      "[SA] new: 406. multiply(X21,multiply(inverse(X21),X22)) = X22 [forward demodulation 405,98]\n",
      "[SA] new: 407. multiply(X32,X33) = multiply(X30,multiply(inverse(multiply(X30,multiply(X31,inverse(X31)))),multiply(X32,X33))) [forward demodulation 372,8]\n",
      "[SA] new: 408. multiply(X41,multiply(X42,multiply(inverse(multiply(X41,multiply(X42,multiply(X43,inverse(X43))))),X44))) = X44 [forward demodulation 375,8]\n",
      "[SA] new: 409. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X1,multiply(identity,identity)))))),X2)) = X2 [superposition 366,10]\n",
      "[SA] new: 410. multiply(X4,multiply(inverse(multiply(X4,multiply(X5,inverse(multiply(X5,multiply(inverse(X3),X3)))))),X6)) = X6 [superposition 366,50]\n",
      "[SA] new: 411. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X1,inverse(identity)))))),X2)) = X2 [superposition 366,6]\n",
      "[SA] new: 412. multiply(X5,multiply(inverse(multiply(X5,multiply(X6,inverse(multiply(X6,multiply(X3,multiply(X4,inverse(multiply(X3,X4))))))))),X7)) = X7 [superposition 366,8]\n",
      "[SA] new: 413. multiply(X1,multiply(inverse(multiply(X1,multiply(inverse(X0),inverse(inverse(X0))))),X2)) = X2 [superposition 366,23]\n",
      "[SA] new: 414. multiply(X4,multiply(inverse(multiply(X4,multiply(X3,inverse(inverse(inverse(X3)))))),X5)) = X5 [superposition 366,49]\n",
      "[SA] new: 415. multiply(X8,multiply(inverse(multiply(X8,multiply(X6,inverse(multiply(inverse(X7),inverse(inverse(multiply(X7,X6)))))))),X9)) = X9 [superposition 366,98]\n",
      "[SA] new: 416. multiply(X12,multiply(inverse(multiply(X12,multiply(X10,inverse(multiply(X11,inverse(inverse(multiply(inverse(X11),X10)))))))),X13)) = X13 [superposition 366,100]\n",
      "[SA] new: 417. multiply(X16,multiply(inverse(multiply(X16,multiply(X14,inverse(inverse(inverse(multiply(X14,multiply(X15,inverse(X15))))))))),X17)) = X17 [superposition 366,243]\n",
      "[SA] new: 418. multiply(X21,multiply(inverse(multiply(X21,multiply(X18,inverse(inverse(inverse(multiply(X18,multiply(X19,inverse(multiply(X19,multiply(X20,inverse(X20)))))))))))),X22)) = X22 [superposition 366,366]\n",
      "[SA] new: 419. multiply(X24,multiply(inverse(multiply(X24,multiply(inverse(multiply(X23,inverse(X23))),inverse(identity)))),X25)) = X25 [superposition 366,7]\n",
      "[SA] new: 420. multiply(X27,multiply(inverse(multiply(X27,multiply(multiply(X26,inverse(X26)),inverse(identity)))),X28)) = X28 [superposition 366,3]\n",
      "[SA] new: 421. multiply(X30,multiply(inverse(multiply(X30,multiply(identity,inverse(multiply(X29,inverse(X29)))))),X31)) = X31 [superposition 366,6]\n",
      "[SA] new: 422. multiply(X35,multiply(inverse(multiply(X35,multiply(multiply(X32,X33),inverse(multiply(X32,multiply(X33,multiply(X34,inverse(X34)))))))),X36)) = X36 [superposition 366,8]\n",
      "[SA] new: 423. multiply(X1,multiply(inverse(multiply(X1,inverse(multiply(identity,multiply(X0,inverse(X0)))))),X2)) = X2 [superposition 366,6]\n",
      "[SA] new: 424. multiply(X6,multiply(inverse(multiply(X6,multiply(X3,multiply(X4,inverse(multiply(multiply(X3,X4),multiply(X5,inverse(X5)))))))),X7)) = X7 [superposition 366,8]\n",
      "[SA] new: 425. multiply(inverse(X0),multiply(inverse(inverse(multiply(X0,multiply(X1,inverse(X1))))),X2)) = X2 [superposition 366,23]\n",
      "[SA] new: 426. multiply(X3,multiply(inverse(inverse(multiply(inverse(X3),multiply(X4,inverse(X4))))),X5)) = X5 [superposition 366,49]\n",
      "[SA] new: 427. multiply(X6,multiply(inverse(multiply(inverse(X7),inverse(multiply(inverse(multiply(X7,X6)),multiply(X8,inverse(X8)))))),X9)) = X9 [superposition 366,98]\n",
      "[SA] new: 428. multiply(X10,multiply(inverse(multiply(X11,inverse(multiply(inverse(multiply(inverse(X11),X10)),multiply(X12,inverse(X12)))))),X13)) = X13 [superposition 366,100]\n",
      "[SA] new: 429. multiply(X14,multiply(inverse(inverse(multiply(inverse(multiply(X14,multiply(X15,inverse(X15)))),multiply(X16,inverse(X16))))),X17)) = X17 [superposition 366,243]\n",
      "[SA] new: 430. multiply(X18,multiply(inverse(inverse(multiply(inverse(multiply(X18,multiply(X19,inverse(multiply(X19,multiply(X20,inverse(X20))))))),multiply(X21,inverse(X21))))),X22)) = X22 [superposition 366,366]\n",
      "[SA] new: 431. multiply(inverse(multiply(X23,inverse(multiply(X23,multiply(X24,inverse(X24)))))),multiply(inverse(identity),X25)) = X25 [superposition 366,7]\n",
      "[SA] new: 432. multiply(multiply(X26,inverse(multiply(X26,multiply(X27,inverse(X27))))),multiply(inverse(identity),X28)) = X28 [superposition 366,3]\n",
      "[SA] new: 433. multiply(identity,multiply(inverse(multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30)))))),X31)) = X31 [superposition 366,6]\n",
      "[SA] new: 434. multiply(multiply(X32,X33),multiply(inverse(multiply(X32,multiply(X33,multiply(X34,inverse(multiply(X34,multiply(X35,inverse(X35)))))))),X36)) = X36 [superposition 366,8]\n",
      "[SA] new: 435. multiply(X0,identity) = multiply(X0,multiply(X1,inverse(multiply(X1,multiply(X2,inverse(X2)))))) [superposition 366,7]\n",
      "[SA] new: 436. multiply(X3,identity) = inverse(multiply(X3,multiply(X4,inverse(multiply(X4,multiply(X5,inverse(X5))))))) [superposition 366,3]\n",
      "[SA] new: 437. identity = multiply(X6,inverse(multiply(X6,multiply(X7,inverse(multiply(X7,multiply(X8,inverse(X8)))))))) [superposition 366,1]\n",
      "[SA] new: 438. multiply(multiply(X9,multiply(X10,inverse(multiply(X10,multiply(X11,inverse(X11)))))),X12) = multiply(X9,X12) [superposition 366,23]\n",
      "[SA] new: 439. multiply(inverse(inverse(multiply(X13,multiply(X14,inverse(multiply(X14,multiply(X15,inverse(X15)))))))),X16) = multiply(X13,X16) [superposition 366,49]\n",
      "[SA] new: 440. multiply(inverse(multiply(X20,inverse(multiply(X17,multiply(X18,inverse(multiply(X18,multiply(X19,inverse(X19))))))))),X21) = multiply(X17,multiply(inverse(X20),X21)) [superposition 366,98]\n",
      "[SA] new: 441. multiply(inverse(multiply(inverse(X25),inverse(multiply(X22,multiply(X23,inverse(multiply(X23,multiply(X24,inverse(X24))))))))),X26) = multiply(X22,multiply(X25,X26)) [superposition 366,100]\n",
      "[SA] new: 442. multiply(inverse(multiply(inverse(multiply(X27,multiply(X28,inverse(multiply(X28,multiply(X29,inverse(X29))))))),multiply(X30,inverse(X30)))),X31) = multiply(X27,X31) [superposition 366,243]\n",
      "[SA] new: 443. multiply(inverse(multiply(inverse(multiply(X32,multiply(X33,inverse(multiply(X33,multiply(X34,inverse(X34))))))),multiply(X35,inverse(multiply(X35,multiply(X36,inverse(X36))))))),X37) = multiply(X32,X37) [superposition 366,366]\n",
      "[SA] new: 444. multiply(X41,multiply(X42,multiply(inverse(multiply(inverse(multiply(X38,multiply(X39,inverse(multiply(X39,multiply(X40,inverse(X40))))))),multiply(X41,X42))),X43))) = multiply(X38,X43) [superposition 366,121]\n",
      "[SA] new: 445. multiply(inverse(multiply(X0,inverse(multiply(X0,multiply(X1,inverse(X1)))))),X2) = X2 [superposition 366,98]\n",
      "[SA] new: 446. multiply(X10,multiply(inverse(multiply(X7,multiply(inverse(multiply(X7,multiply(X8,inverse(multiply(X8,multiply(X9,inverse(X9))))))),X10))),X11)) = X11 [superposition 366,121]\n",
      "[SA] new: 447. multiply(inverse(multiply(identity,multiply(X12,inverse(multiply(X12,multiply(X13,inverse(X13))))))),X14) = X14 [superposition 366,6]\n",
      "[SA] new: 448. multiply(X15,multiply(X16,multiply(inverse(multiply(multiply(X15,X16),multiply(X17,inverse(multiply(X17,multiply(X18,inverse(X18))))))),X19))) = X19 [superposition 366,8]\n",
      "[SA] new: 449. multiply(inverse(multiply(X0,inverse(multiply(X0,multiply(X1,inverse(X1)))))),X2) = X2 [superposition 98,366]\n",
      "[SA] new: 450. multiply(X3,multiply(inverse(multiply(X3,multiply(X4,inverse(multiply(X4,multiply(X5,inverse(X5))))))),multiply(X6,multiply(inverse(X6),X7)))) = X7 [superposition 121,366]\n",
      "[SA] new: 451. multiply(inverse(multiply(X8,multiply(X9,inverse(X9)))),multiply(X8,X10)) = X10 [superposition 121,366]\n",
      "[SA] new: 452. multiply(X14,multiply(inverse(multiply(X14,multiply(X11,inverse(multiply(X11,multiply(X12,inverse(multiply(X12,multiply(X13,inverse(X13)))))))))),X15)) = X15 [superposition 121,366]\n",
      "[SA] new: 453. multiply(X19,multiply(inverse(multiply(X16,multiply(inverse(multiply(X16,multiply(X17,inverse(multiply(X17,multiply(X18,inverse(X18))))))),X19))),X20)) = X20 [superposition 121,366]\n",
      "[SA] new: 454. multiply(X21,multiply(inverse(inverse(inverse(multiply(X21,multiply(X22,inverse(multiply(X22,multiply(X23,inverse(X23))))))))),X24)) = X24 [superposition 243,366]\n",
      "[SA] new: 455. multiply(inverse(multiply(identity,multiply(X25,inverse(multiply(X25,multiply(X26,inverse(X26))))))),X27) = X27 [superposition 6,366]\n",
      "[SA] new: 456. multiply(X32,multiply(X28,multiply(multiply(inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30))))))),X31),multiply(inverse(multiply(X32,X31)),X33)))) = X33 [superposition 121,366]\n",
      "[SA] new: 457. multiply(inverse(X34),X38) = multiply(multiply(inverse(multiply(X34,multiply(X35,inverse(multiply(X35,multiply(X36,inverse(X36))))))),X37),multiply(inverse(X37),X38)) [superposition 98,366]\n",
      "[SA] new: 458. multiply(inverse(multiply(X39,multiply(X40,inverse(multiply(X40,multiply(X41,inverse(X41))))))),X42) = multiply(inverse(X39),X42) [superposition 23,366]\n",
      "[SA] new: 459. multiply(X43,multiply(multiply(inverse(multiply(X43,multiply(X44,inverse(multiply(X44,multiply(X45,inverse(X45))))))),X46),X47)) = multiply(X46,X47) [superposition 8,366]\n",
      "[SA] new: 460. multiply(X48,X52) = multiply(multiply(inverse(multiply(inverse(X48),multiply(X49,inverse(multiply(X49,multiply(X50,inverse(X50))))))),X51),multiply(inverse(X51),X52)) [superposition 100,366]\n",
      "[SA] new: 461. multiply(inverse(multiply(inverse(X53),multiply(X54,inverse(multiply(X54,multiply(X55,inverse(X55))))))),X56) = multiply(X53,X56) [superposition 49,366]\n",
      "[SA] new: 462. multiply(X57,multiply(X58,multiply(inverse(multiply(multiply(X57,X58),multiply(X59,inverse(multiply(X59,multiply(X60,inverse(X60))))))),X61))) = X61 [superposition 8,366]\n",
      "[SA] new: 463. multiply(inverse(X62),multiply(inverse(multiply(inverse(multiply(X62,X63)),multiply(X64,inverse(multiply(X64,multiply(X65,inverse(X65))))))),X66)) = multiply(X63,X66) [superposition 98,366]\n",
      "[SA] new: 464. multiply(X67,multiply(inverse(multiply(inverse(multiply(inverse(X67),X68)),multiply(X69,inverse(multiply(X69,multiply(X70,inverse(X70))))))),X71)) = multiply(X68,X71) [superposition 100,366]\n",
      "[SA] new: 465. multiply(inverse(multiply(inverse(multiply(X72,multiply(X73,inverse(X73)))),multiply(X74,inverse(multiply(X74,multiply(X75,inverse(X75))))))),X76) = multiply(X72,X76) [superposition 243,366]\n",
      "[SA] new: 466. multiply(inverse(multiply(inverse(multiply(X77,multiply(X78,X79))),multiply(X80,inverse(multiply(X80,multiply(X81,inverse(X81))))))),X82) = multiply(X77,multiply(X78,multiply(X79,X82))) [superposition 121,366]\n",
      "[SA] new: 467. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X1,identity))))),X2)) = X2 [forward demodulation 409,3]\n",
      "[SA] new: 468. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(X1)))),X2)) = X2 [forward demodulation 467,1]\n",
      "[SA] new: 469. multiply(X4,multiply(inverse(multiply(X4,multiply(X5,inverse(multiply(X5,identity))))),X6)) = X6 [forward demodulation 410,7]\n",
      "[SA] new: 470. multiply(X4,multiply(inverse(multiply(X4,multiply(X5,inverse(X5)))),X6)) = X6 [forward demodulation 469,1]\n",
      "[SA] new: 471. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X1,identity))))),X2)) = X2 [forward demodulation 411,10]\n",
      "[SA] new: 472. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(X1)))),X2)) = X2 [forward demodulation 471,1]\n",
      "[SA] new: 473. multiply(X4,multiply(inverse(multiply(X4,multiply(X3,inverse(X3)))),X5)) = X5 [forward demodulation 414,50]\n",
      "[SA] new: 474. multiply(X8,multiply(inverse(multiply(X8,multiply(X6,inverse(multiply(inverse(X7),multiply(X7,X6)))))),X9)) = X9 [forward demodulation 415,50]\n",
      "[SA] new: 475. multiply(X8,multiply(inverse(multiply(X8,multiply(X6,inverse(X6)))),X9)) = X9 [forward demodulation 474,23]\n",
      "[SA] new: 476. multiply(X12,multiply(inverse(multiply(X12,multiply(X10,inverse(multiply(X11,multiply(inverse(X11),X10)))))),X13)) = X13 [forward demodulation 416,50]\n",
      "[SA] new: 477. multiply(X12,multiply(inverse(multiply(X12,multiply(X10,inverse(X10)))),X13)) = X13 [forward demodulation 476,49]\n",
      "[SA] new: 478. multiply(X16,multiply(inverse(multiply(X16,multiply(X14,inverse(multiply(X14,multiply(X15,inverse(X15))))))),X17)) = X17 [forward demodulation 417,50]\n",
      "[SA] new: 479. multiply(X21,multiply(inverse(multiply(X21,multiply(X18,inverse(multiply(X18,multiply(X19,inverse(multiply(X19,multiply(X20,inverse(X20)))))))))),X22)) = X22 [forward demodulation 418,50]\n",
      "[SA] new: 480. multiply(X24,multiply(inverse(multiply(X24,multiply(inverse(multiply(X23,inverse(X23))),identity))),X25)) = X25 [forward demodulation 419,10]\n",
      "[SA] new: 481. multiply(X24,multiply(inverse(multiply(X24,inverse(multiply(X23,inverse(X23))))),X25)) = X25 [forward demodulation 480,1]\n",
      "[SA] new: 482. multiply(X27,multiply(inverse(multiply(X27,multiply(X26,multiply(inverse(X26),inverse(identity))))),X28)) = X28 [forward demodulation 420,8]\n",
      "[SA] new: 483. multiply(X27,multiply(inverse(multiply(X27,inverse(identity))),X28)) = X28 [forward demodulation 482,49]\n",
      "[SA] new: 484. multiply(X27,multiply(inverse(multiply(X27,identity)),X28)) = X28 [forward demodulation 483,10]\n",
      "[SA] new: 485. multiply(X27,multiply(inverse(X27),X28)) = X28 [forward demodulation 484,1]\n",
      "[SA] new: 486. multiply(X30,multiply(inverse(multiply(X30,inverse(multiply(X29,inverse(X29))))),X31)) = X31 [forward demodulation 421,6]\n",
      "[SA] new: 487. multiply(X35,multiply(inverse(multiply(X35,multiply(X32,multiply(X33,inverse(multiply(X32,multiply(X33,multiply(X34,inverse(X34))))))))),X36)) = X36 [forward demodulation 422,8]\n",
      "[SA] new: 488. multiply(X1,multiply(inverse(multiply(X1,inverse(multiply(X0,inverse(X0))))),X2)) = X2 [forward demodulation 423,6]\n",
      "[SA] new: 489. multiply(X6,multiply(inverse(multiply(X6,multiply(X3,multiply(X4,inverse(multiply(X3,multiply(X4,multiply(X5,inverse(X5))))))))),X7)) = X7 [forward demodulation 424,8]\n",
      "[SA] new: 490. multiply(inverse(X0),multiply(multiply(X0,multiply(X1,inverse(X1))),X2)) = X2 [forward demodulation 425,50]\n",
      "[SA] new: 491. multiply(inverse(X0),multiply(X0,multiply(multiply(X1,inverse(X1)),X2))) = X2 [forward demodulation 490,8]\n",
      "[SA] new: 492. multiply(inverse(X0),multiply(X0,multiply(X1,multiply(inverse(X1),X2)))) = X2 [forward demodulation 491,8]\n",
      "[SA] new: 493. multiply(inverse(X0),multiply(X0,X2)) = X2 [forward demodulation 492,49]\n",
      "[SA] new: 494. multiply(X3,multiply(multiply(inverse(X3),multiply(X4,inverse(X4))),X5)) = X5 [forward demodulation 426,50]\n",
      "[SA] new: 495. multiply(X3,multiply(inverse(X3),multiply(multiply(X4,inverse(X4)),X5))) = X5 [forward demodulation 494,8]\n",
      "[SA] new: 496. multiply(X3,multiply(inverse(X3),multiply(X4,multiply(inverse(X4),X5)))) = X5 [forward demodulation 495,8]\n",
      "[SA] new: 497. multiply(X3,multiply(inverse(X3),X5)) = X5 [forward demodulation 496,49]\n",
      "[SA] new: 498. multiply(X14,multiply(multiply(inverse(multiply(X14,multiply(X15,inverse(X15)))),multiply(X16,inverse(X16))),X17)) = X17 [forward demodulation 429,50]\n",
      "[SA] new: 499. multiply(X14,multiply(inverse(multiply(X14,multiply(X15,inverse(X15)))),multiply(multiply(X16,inverse(X16)),X17))) = X17 [forward demodulation 498,8]\n",
      "[SA] new: 500. multiply(X14,multiply(inverse(multiply(X14,multiply(X15,inverse(X15)))),multiply(X16,multiply(inverse(X16),X17)))) = X17 [forward demodulation 499,8]\n",
      "[SA] new: 501. multiply(X14,multiply(inverse(multiply(X14,multiply(X15,inverse(X15)))),X17)) = X17 [forward demodulation 500,49]\n",
      "[SA] new: 502. multiply(X18,multiply(multiply(inverse(multiply(X18,multiply(X19,inverse(multiply(X19,multiply(X20,inverse(X20))))))),multiply(X21,inverse(X21))),X22)) = X22 [forward demodulation 430,50]\n",
      "[SA] new: 503. multiply(X18,multiply(inverse(multiply(X18,multiply(X19,inverse(multiply(X19,multiply(X20,inverse(X20))))))),multiply(multiply(X21,inverse(X21)),X22))) = X22 [forward demodulation 502,8]\n",
      "[SA] new: 504. multiply(X18,multiply(inverse(multiply(X18,multiply(X19,inverse(multiply(X19,multiply(X20,inverse(X20))))))),multiply(X21,multiply(inverse(X21),X22)))) = X22 [forward demodulation 503,8]\n",
      "[SA] new: 505. multiply(X18,multiply(inverse(multiply(X18,multiply(X19,inverse(multiply(X19,multiply(X20,inverse(X20))))))),X22)) = X22 [forward demodulation 504,49]\n",
      "[SA] new: 506. multiply(inverse(multiply(X23,inverse(multiply(X23,multiply(X24,inverse(X24)))))),multiply(identity,X25)) = X25 [forward demodulation 431,10]\n",
      "[SA] new: 507. multiply(inverse(multiply(X23,inverse(multiply(X23,multiply(X24,inverse(X24)))))),X25) = X25 [forward demodulation 506,6]\n",
      "[SA] new: 508. multiply(multiply(X26,inverse(multiply(X26,multiply(X27,inverse(X27))))),multiply(identity,X28)) = X28 [forward demodulation 432,10]\n",
      "[SA] new: 509. multiply(multiply(X26,inverse(multiply(X26,multiply(X27,inverse(X27))))),X28) = X28 [forward demodulation 508,6]\n",
      "[SA] new: 510. multiply(X0,multiply(X1,inverse(multiply(X1,multiply(X2,inverse(X2)))))) = X0 [forward demodulation 435,1]\n",
      "[SA] new: 511. inverse(multiply(X3,multiply(X4,inverse(multiply(X4,multiply(X5,inverse(X5))))))) = X3 [forward demodulation 436,1]\n",
      "[SA] new: 512. multiply(X13,X16) = multiply(multiply(X13,multiply(X14,inverse(multiply(X14,multiply(X15,inverse(X15)))))),X16) [forward demodulation 439,50]\n",
      "[SA] new: 513. multiply(X10,multiply(inverse(X10),X11)) = X11 [forward demodulation 446,366]\n",
      "[SA] new: 514. multiply(inverse(multiply(X12,inverse(multiply(X12,multiply(X13,inverse(X13)))))),X14) = X14 [forward demodulation 447,6]\n",
      "[SA] new: 515. multiply(X15,multiply(X16,multiply(inverse(multiply(X15,multiply(X16,multiply(X17,inverse(multiply(X17,multiply(X18,inverse(X18)))))))),X19))) = X19 [forward demodulation 448,8]\n",
      "[SA] new: 516. multiply(X3,multiply(inverse(multiply(X3,multiply(X4,inverse(multiply(X4,multiply(X5,inverse(X5))))))),X7)) = X7 [forward demodulation 450,49]\n",
      "[SA] new: 517. multiply(X19,multiply(inverse(X19),X20)) = X20 [forward demodulation 453,366]\n",
      "[SA] new: 518. multiply(X21,multiply(inverse(multiply(X21,multiply(X22,inverse(multiply(X22,multiply(X23,inverse(X23))))))),X24)) = X24 [forward demodulation 454,50]\n",
      "[SA] new: 519. multiply(inverse(multiply(X25,inverse(multiply(X25,multiply(X26,inverse(X26)))))),X27) = X27 [forward demodulation 455,6]\n",
      "[SA] new: 520. multiply(X32,multiply(X28,multiply(inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30))))))),multiply(X31,multiply(inverse(multiply(X32,X31)),X33))))) = X33 [forward demodulation 456,8]\n",
      "[SA] new: 521. multiply(X32,multiply(X31,multiply(inverse(multiply(X32,X31)),X33))) = X33 [forward demodulation 520,366]\n",
      "[SA] new: 522. multiply(X32,multiply(inverse(X32),X33)) = X33 [forward demodulation 521,98]\n",
      "[SA] new: 523. multiply(X46,X47) = multiply(X43,multiply(inverse(multiply(X43,multiply(X44,inverse(multiply(X44,multiply(X45,inverse(X45))))))),multiply(X46,X47))) [forward demodulation 459,8]\n",
      "[SA] new: 524. multiply(X57,multiply(X58,multiply(inverse(multiply(X57,multiply(X58,multiply(X59,inverse(multiply(X59,multiply(X60,inverse(X60)))))))),X61))) = X61 [forward demodulation 462,8]\n",
      "[SA] new: 525. multiply(X0,multiply(inverse(multiply(X0,inverse(multiply(identity,identity)))),X1)) = X1 [superposition 486,10]\n",
      "[SA] new: 526. multiply(X3,multiply(inverse(multiply(X3,inverse(multiply(inverse(X2),X2)))),X4)) = X4 [superposition 486,50]\n",
      "[SA] new: 527. multiply(X0,multiply(inverse(multiply(X0,inverse(inverse(identity)))),X1)) = X1 [superposition 486,6]\n",
      "[SA] new: 528. multiply(X4,multiply(inverse(multiply(X4,inverse(multiply(X2,multiply(X3,inverse(multiply(X2,X3))))))),X5)) = X5 [superposition 486,8]\n",
      "[SA] new: 529. multiply(inverse(inverse(multiply(X0,inverse(X0)))),multiply(inverse(identity),X1)) = X1 [superposition 486,7]\n",
      "[SA] new: 530. multiply(inverse(multiply(X2,inverse(X2))),multiply(inverse(identity),X3)) = X3 [superposition 486,3]\n",
      "[SA] new: 531. multiply(identity,multiply(inverse(inverse(multiply(X4,inverse(X4)))),X5)) = X5 [superposition 486,6]\n",
      "[SA] new: 532. multiply(multiply(X6,X7),multiply(inverse(multiply(X6,multiply(X7,inverse(multiply(X8,inverse(X8)))))),X9)) = X9 [superposition 486,8]\n",
      "[SA] new: 533. multiply(X0,identity) = multiply(X0,inverse(multiply(X1,inverse(X1)))) [superposition 486,7]\n",
      "[SA] new: 534. multiply(X2,identity) = inverse(multiply(X2,inverse(multiply(X3,inverse(X3))))) [superposition 486,3]\n",
      "[SA] new: 535. identity = multiply(X4,inverse(multiply(X4,inverse(multiply(X5,inverse(X5)))))) [superposition 486,1]\n",
      "[SA] new: 536. multiply(X6,X8) = multiply(multiply(X6,inverse(multiply(X7,inverse(X7)))),X8) [superposition 486,23]\n",
      "[SA] new: 537. multiply(X9,X11) = multiply(inverse(inverse(multiply(X9,inverse(multiply(X10,inverse(X10)))))),X11) [superposition 486,49]\n",
      "[SA] new: 538. multiply(X12,multiply(inverse(X14),X15)) = multiply(inverse(multiply(X14,inverse(multiply(X12,inverse(multiply(X13,inverse(X13))))))),X15) [superposition 486,98]\n",
      "[SA] new: 539. multiply(X16,multiply(X18,X19)) = multiply(inverse(multiply(inverse(X18),inverse(multiply(X16,inverse(multiply(X17,inverse(X17))))))),X19) [superposition 486,100]\n",
      "[SA] new: 540. multiply(X20,X23) = multiply(inverse(multiply(inverse(multiply(X20,inverse(multiply(X21,inverse(X21))))),multiply(X22,inverse(X22)))),X23) [superposition 486,243]\n",
      "[SA] new: 541. multiply(inverse(multiply(inverse(multiply(X24,inverse(multiply(X25,inverse(X25))))),inverse(multiply(X26,inverse(X26))))),X27) = multiply(X24,X27) [superposition 486,486]\n",
      "[SA] new: 542. multiply(inverse(multiply(inverse(multiply(X28,inverse(multiply(X29,inverse(X29))))),multiply(X30,inverse(multiply(X30,multiply(X31,inverse(X31))))))),X32) = multiply(X28,X32) [superposition 486,366]\n",
      "[SA] new: 543. multiply(X35,multiply(X36,multiply(inverse(multiply(inverse(multiply(X33,inverse(multiply(X34,inverse(X34))))),multiply(X35,X36))),X37))) = multiply(X33,X37) [superposition 486,121]\n",
      "[SA] new: 544. multiply(multiply(X3,inverse(X3)),X4) = X4 [superposition 486,100]\n",
      "[SA] new: 545. multiply(inverse(inverse(multiply(X5,inverse(X5)))),X6) = X6 [superposition 486,98]\n",
      "[SA] new: 546. multiply(X9,multiply(inverse(multiply(X7,multiply(inverse(multiply(X7,inverse(multiply(X8,inverse(X8))))),X9))),X10)) = X10 [superposition 486,121]\n",
      "[SA] new: 547. multiply(inverse(multiply(identity,inverse(multiply(X11,inverse(X11))))),X12) = X12 [superposition 486,6]\n",
      "[SA] new: 548. multiply(X13,multiply(X14,multiply(inverse(multiply(multiply(X13,X14),inverse(multiply(X15,inverse(X15))))),X16))) = X16 [superposition 486,8]\n",
      "[SA] new: 549. multiply(multiply(X0,inverse(X0)),X1) = X1 [superposition 100,486]\n",
      "[SA] new: 550. multiply(inverse(inverse(multiply(X2,inverse(X2)))),X3) = X3 [superposition 98,486]\n",
      "[SA] new: 551. multiply(X4,multiply(inverse(multiply(X4,inverse(multiply(X5,inverse(X5))))),multiply(X6,multiply(inverse(X6),X7)))) = X7 [superposition 121,486]\n",
      "[SA] new: 552. multiply(X10,multiply(inverse(multiply(X10,multiply(X8,inverse(multiply(X8,inverse(multiply(X9,inverse(X9)))))))),X11)) = X11 [superposition 121,486]\n",
      "[SA] new: 553. multiply(X14,multiply(inverse(multiply(X12,multiply(inverse(multiply(X12,inverse(multiply(X13,inverse(X13))))),X14))),X15)) = X15 [superposition 121,486]\n",
      "[SA] new: 554. multiply(X18,multiply(inverse(multiply(X18,multiply(X16,inverse(inverse(inverse(multiply(X16,inverse(multiply(X17,inverse(X17)))))))))),X19)) = X19 [superposition 366,486]\n",
      "[SA] new: 555. multiply(X20,multiply(inverse(inverse(inverse(multiply(X20,inverse(multiply(X21,inverse(X21))))))),X22)) = X22 [superposition 243,486]\n",
      "[SA] new: 556. multiply(X23,multiply(inverse(inverse(multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),multiply(X25,inverse(X25))))),X26)) = X26 [superposition 366,486]\n",
      "[SA] new: 557. multiply(inverse(multiply(identity,inverse(multiply(X27,inverse(X27))))),X28) = X28 [superposition 6,486]\n",
      "[SA] new: 558. multiply(X32,multiply(X29,multiply(multiply(inverse(multiply(X29,inverse(multiply(X30,inverse(X30))))),X31),multiply(inverse(multiply(X32,X31)),X33)))) = X33 [superposition 121,486]\n",
      "[SA] new: 559. multiply(inverse(X34),X37) = multiply(multiply(inverse(multiply(X34,inverse(multiply(X35,inverse(X35))))),X36),multiply(inverse(X36),X37)) [superposition 98,486]\n",
      "[SA] new: 560. multiply(inverse(multiply(X38,inverse(multiply(X39,inverse(X39))))),X40) = multiply(inverse(X38),X40) [superposition 23,486]\n",
      "[SA] new: 561. multiply(X41,multiply(multiply(inverse(multiply(X41,inverse(multiply(X42,inverse(X42))))),X43),X44)) = multiply(X43,X44) [superposition 8,486]\n",
      "[SA] new: 562. multiply(X45,X48) = multiply(multiply(inverse(multiply(inverse(X45),inverse(multiply(X46,inverse(X46))))),X47),multiply(inverse(X47),X48)) [superposition 100,486]\n",
      "[SA] new: 563. multiply(inverse(multiply(inverse(X49),inverse(multiply(X50,inverse(X50))))),X51) = multiply(X49,X51) [superposition 49,486]\n",
      "[SA] new: 564. multiply(X52,multiply(X53,multiply(inverse(multiply(multiply(X52,X53),inverse(multiply(X54,inverse(X54))))),X55))) = X55 [superposition 8,486]\n",
      "[SA] new: 565. multiply(inverse(X56),multiply(inverse(multiply(inverse(multiply(X56,X57)),inverse(multiply(X58,inverse(X58))))),X59)) = multiply(X57,X59) [superposition 98,486]\n",
      "[SA] new: 566. multiply(X60,multiply(inverse(multiply(inverse(multiply(inverse(X60),X61)),inverse(multiply(X62,inverse(X62))))),X63)) = multiply(X61,X63) [superposition 100,486]\n",
      "[SA] new: 567. multiply(inverse(multiply(inverse(multiply(X64,multiply(X65,inverse(X65)))),inverse(multiply(X66,inverse(X66))))),X67) = multiply(X64,X67) [superposition 243,486]\n",
      "[SA] new: 568. multiply(inverse(multiply(inverse(multiply(X68,multiply(X69,X70))),inverse(multiply(X71,inverse(X71))))),X72) = multiply(X68,multiply(X69,multiply(X70,X72))) [superposition 121,486]\n",
      "[SA] new: 569. multiply(inverse(multiply(inverse(multiply(X73,multiply(X74,inverse(multiply(X74,multiply(X75,inverse(X75))))))),inverse(multiply(X76,inverse(X76))))),X77) = multiply(X73,X77) [superposition 366,486]\n",
      "[SA] new: 570. multiply(X0,multiply(inverse(multiply(X0,inverse(identity))),X1)) = X1 [forward demodulation 525,3]\n",
      "[SA] new: 571. multiply(X0,multiply(inverse(multiply(X0,identity)),X1)) = X1 [forward demodulation 570,10]\n",
      "[SA] new: 572. multiply(X0,multiply(inverse(X0),X1)) = X1 [forward demodulation 571,1]\n",
      "[SA] new: 573. multiply(X3,multiply(inverse(multiply(X3,inverse(identity))),X4)) = X4 [forward demodulation 526,7]\n",
      "[SA] new: 574. multiply(X3,multiply(inverse(multiply(X3,identity)),X4)) = X4 [forward demodulation 573,10]\n",
      "[SA] new: 575. multiply(X3,multiply(inverse(X3),X4)) = X4 [forward demodulation 574,1]\n",
      "[SA] new: 576. multiply(X0,multiply(inverse(multiply(X0,identity)),X1)) = X1 [forward demodulation 527,50]\n",
      "[SA] new: 577. multiply(X0,multiply(inverse(X0),X1)) = X1 [forward demodulation 576,1]\n",
      "[SA] new: 578. multiply(inverse(inverse(multiply(X0,inverse(X0)))),multiply(identity,X1)) = X1 [forward demodulation 529,10]\n",
      "[SA] new: 579. multiply(inverse(inverse(multiply(X0,inverse(X0)))),X1) = X1 [forward demodulation 578,6]\n",
      "[SA] new: 580. multiply(multiply(X0,inverse(X0)),X1) = X1 [forward demodulation 579,50]\n",
      "[SA] new: 581. multiply(inverse(multiply(X2,inverse(X2))),multiply(identity,X3)) = X3 [forward demodulation 530,10]\n",
      "[SA] new: 582. multiply(inverse(multiply(X2,inverse(X2))),X3) = X3 [forward demodulation 581,6]\n",
      "[SA] new: 583. multiply(identity,multiply(multiply(X4,inverse(X4)),X5)) = X5 [forward demodulation 531,50]\n",
      "[SA] new: 584. multiply(identity,multiply(X4,multiply(inverse(X4),X5))) = X5 [forward demodulation 583,8]\n",
      "[SA] new: 585. multiply(identity,X5) = X5 [forward demodulation 584,49]\n",
      "[SA] new: 586. multiply(X0,inverse(multiply(X1,inverse(X1)))) = X0 [forward demodulation 533,1]\n",
      "[SA] new: 587. inverse(multiply(X2,inverse(multiply(X3,inverse(X3))))) = X2 [forward demodulation 534,1]\n",
      "[SA] new: 588. multiply(X9,X11) = multiply(multiply(X9,inverse(multiply(X10,inverse(X10)))),X11) [forward demodulation 537,50]\n",
      "[SA] new: 589. multiply(multiply(X5,inverse(X5)),X6) = X6 [forward demodulation 545,50]\n",
      "[SA] new: 590. multiply(X9,multiply(inverse(X9),X10)) = X10 [forward demodulation 546,486]\n",
      "[SA] new: 591. multiply(inverse(inverse(multiply(X11,inverse(X11)))),X12) = X12 [forward demodulation 547,6]\n",
      "[SA] new: 592. multiply(multiply(X11,inverse(X11)),X12) = X12 [forward demodulation 591,50]\n",
      "[SA] new: 593. multiply(X13,multiply(X14,multiply(inverse(multiply(X13,multiply(X14,inverse(multiply(X15,inverse(X15)))))),X16))) = X16 [forward demodulation 548,8]\n",
      "[SA] new: 594. multiply(multiply(X2,inverse(X2)),X3) = X3 [forward demodulation 550,50]\n",
      "[SA] new: 595. multiply(X4,multiply(inverse(multiply(X4,inverse(multiply(X5,inverse(X5))))),X7)) = X7 [forward demodulation 551,49]\n",
      "[SA] new: 596. multiply(X14,multiply(inverse(X14),X15)) = X15 [forward demodulation 553,486]\n",
      "[SA] new: 597. multiply(X18,multiply(inverse(multiply(X18,multiply(X16,inverse(multiply(X16,inverse(multiply(X17,inverse(X17)))))))),X19)) = X19 [forward demodulation 554,50]\n",
      "[SA] new: 598. multiply(X20,multiply(inverse(multiply(X20,inverse(multiply(X21,inverse(X21))))),X22)) = X22 [forward demodulation 555,50]\n",
      "[SA] new: 599. multiply(X23,multiply(multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),multiply(X25,inverse(X25))),X26)) = X26 [forward demodulation 556,50]\n",
      "[SA] new: 600. multiply(X23,multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),multiply(multiply(X25,inverse(X25)),X26))) = X26 [forward demodulation 599,8]\n",
      "[SA] new: 601. multiply(X23,multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),multiply(X25,multiply(inverse(X25),X26)))) = X26 [forward demodulation 600,8]\n",
      "[SA] new: 602. multiply(X23,multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),X26)) = X26 [forward demodulation 601,49]\n",
      "[SA] new: 603. multiply(inverse(inverse(multiply(X27,inverse(X27)))),X28) = X28 [forward demodulation 557,6]\n",
      "[SA] new: 604. multiply(multiply(X27,inverse(X27)),X28) = X28 [forward demodulation 603,50]\n",
      "[SA] new: 605. multiply(X32,multiply(X29,multiply(inverse(multiply(X29,inverse(multiply(X30,inverse(X30))))),multiply(X31,multiply(inverse(multiply(X32,X31)),X33))))) = X33 [forward demodulation 558,8]\n",
      "[SA] new: 606. multiply(X32,multiply(X31,multiply(inverse(multiply(X32,X31)),X33))) = X33 [forward demodulation 605,486]\n",
      "[SA] new: 607. multiply(X32,multiply(inverse(X32),X33)) = X33 [forward demodulation 606,98]\n",
      "[SA] new: 608. multiply(X43,X44) = multiply(X41,multiply(inverse(multiply(X41,inverse(multiply(X42,inverse(X42))))),multiply(X43,X44))) [forward demodulation 561,8]\n",
      "[SA] new: 609. multiply(X52,multiply(X53,multiply(inverse(multiply(X52,multiply(X53,inverse(multiply(X54,inverse(X54)))))),X55))) = X55 [forward demodulation 564,8]\n",
      "[SA] new: 610. multiply(X1,multiply(inverse(multiply(X1,multiply(inverse(X0),multiply(X0,inverse(identity))))),X2)) = X2 [superposition 340,7]\n",
      "[SA] new: 611. multiply(X4,multiply(inverse(multiply(X4,multiply(X3,multiply(X3,inverse(identity))))),X5)) = X5 [superposition 340,3]\n",
      "[SA] new: 612. multiply(X7,multiply(inverse(multiply(X7,multiply(identity,multiply(X6,inverse(X6))))),X8)) = X8 [superposition 340,6]\n",
      "[SA] new: 613. multiply(X12,multiply(inverse(multiply(X12,multiply(multiply(X9,X10),multiply(X11,inverse(multiply(X9,multiply(X10,X11))))))),X13)) = X13 [superposition 340,8]\n",
      "[SA] new: 614. multiply(X15,multiply(inverse(multiply(X15,multiply(X14,multiply(identity,inverse(X14))))),X16)) = X16 [superposition 340,1]\n",
      "[SA] new: 615. multiply(X19,multiply(inverse(multiply(X19,multiply(inverse(X17),multiply(multiply(X17,X18),inverse(X18))))),X20)) = X20 [superposition 340,23]\n",
      "[SA] new: 616. multiply(X23,multiply(inverse(multiply(X23,multiply(X21,multiply(multiply(inverse(X21),X22),inverse(X22))))),X24)) = X24 [superposition 340,49]\n",
      "[SA] new: 617. multiply(X28,multiply(inverse(multiply(X28,multiply(X25,multiply(multiply(inverse(multiply(X26,X25)),X27),inverse(multiply(inverse(X26),X27)))))),X29)) = X29 [superposition 340,98]\n",
      "[SA] new: 618. multiply(X33,multiply(inverse(multiply(X33,multiply(X30,multiply(multiply(inverse(multiply(inverse(X31),X30)),X32),inverse(multiply(X31,X32)))))),X34)) = X34 [superposition 340,100]\n",
      "[SA] new: 619. multiply(X38,multiply(inverse(multiply(X38,multiply(X35,multiply(multiply(inverse(multiply(X35,multiply(X36,inverse(X36)))),X37),inverse(X37))))),X39)) = X39 [superposition 340,243]\n",
      "[SA] new: 620. multiply(X43,multiply(inverse(multiply(X43,multiply(X40,multiply(multiply(inverse(multiply(X40,inverse(multiply(X41,inverse(X41))))),X42),inverse(X42))))),X44)) = X44 [superposition 340,486]\n",
      "[SA] new: 621. multiply(X49,multiply(inverse(multiply(X49,multiply(X45,multiply(multiply(inverse(multiply(X45,multiply(X46,inverse(multiply(X46,multiply(X47,inverse(X47))))))),X48),inverse(X48))))),X50)) = X50 [superposition 340,366]\n",
      "[SA] new: 622. multiply(X55,multiply(inverse(multiply(X55,multiply(X51,multiply(multiply(inverse(multiply(X51,multiply(X52,multiply(X53,inverse(multiply(X52,X53)))))),X54),inverse(X54))))),X56)) = X56 [superposition 340,340]\n",
      "[SA] new: 623. multiply(X61,multiply(inverse(multiply(X61,multiply(X57,multiply(multiply(X58,multiply(X59,multiply(inverse(multiply(X57,multiply(X58,X59))),X60))),inverse(X60))))),X62)) = X62 [superposition 340,121]\n",
      "[SA] new: 624. multiply(X63,multiply(inverse(multiply(X63,multiply(a,multiply(b,inverse(c))))),X64)) = X64 [superposition 340,4]\n",
      "[SA] new: 625. multiply(X1,multiply(inverse(multiply(X1,multiply(X0,inverse(multiply(X0,identity))))),X2)) = X2 [superposition 340,6]\n",
      "[SA] new: 626. multiply(X6,multiply(inverse(multiply(X6,multiply(X5,multiply(X3,multiply(X4,inverse(multiply(X5,multiply(X3,X4)))))))),X7)) = X7 [superposition 340,8]\n",
      "[SA] new: 627. multiply(X1,multiply(inverse(multiply(X1,inverse(multiply(inverse(X0),X0)))),X2)) = X2 [superposition 340,23]\n",
      "[SA] new: 628. multiply(X4,multiply(inverse(multiply(X4,inverse(multiply(X3,inverse(X3))))),X5)) = X5 [superposition 340,49]\n",
      "[SA] new: 629. multiply(X8,multiply(inverse(multiply(X8,multiply(inverse(X7),inverse(multiply(X6,inverse(multiply(X7,X6))))))),X9)) = X9 [superposition 340,98]\n",
      "[SA] new: 630. multiply(X12,multiply(inverse(multiply(X12,multiply(X11,inverse(multiply(X10,inverse(multiply(inverse(X11),X10))))))),X13)) = X13 [superposition 340,100]\n",
      "[SA] new: 631. multiply(X16,multiply(inverse(multiply(X16,inverse(multiply(X14,inverse(multiply(X14,multiply(X15,inverse(X15)))))))),X17)) = X17 [superposition 340,243]\n",
      "[SA] new: 632. multiply(X20,multiply(inverse(multiply(X20,inverse(multiply(X18,inverse(multiply(X18,inverse(multiply(X19,inverse(X19))))))))),X21)) = X21 [superposition 340,486]\n",
      "[SA] new: 633. multiply(X25,multiply(inverse(multiply(X25,inverse(multiply(X22,inverse(multiply(X22,multiply(X23,inverse(multiply(X23,multiply(X24,inverse(X24))))))))))),X26)) = X26 [superposition 340,366]\n",
      "[SA] new: 634. multiply(X30,multiply(inverse(multiply(X30,inverse(multiply(X27,inverse(multiply(X27,multiply(X28,multiply(X29,inverse(multiply(X28,X29)))))))))),X31)) = X31 [superposition 340,340]\n",
      "[SA] new: 635. multiply(X33,multiply(inverse(multiply(X33,multiply(X32,inverse(multiply(identity,X32))))),X34)) = X34 [superposition 340,6]\n",
      "[SA] new: 636. multiply(X38,multiply(inverse(multiply(X38,multiply(X35,multiply(X36,multiply(X37,inverse(multiply(multiply(X35,X36),X37))))))),X39)) = X39 [superposition 340,8]\n",
      "[SA] new: 637. multiply(inverse(X0),multiply(inverse(multiply(X1,inverse(multiply(X0,X1)))),X2)) = X2 [superposition 340,23]\n",
      "[SA] new: 638. multiply(X3,multiply(inverse(multiply(X4,inverse(multiply(inverse(X3),X4)))),X5)) = X5 [superposition 340,49]\n",
      "[SA] new: 639. multiply(X6,multiply(inverse(multiply(inverse(X7),multiply(X8,inverse(multiply(inverse(multiply(X7,X6)),X8))))),X9)) = X9 [superposition 340,98]\n",
      "[SA] new: 640. multiply(X10,multiply(inverse(multiply(X11,multiply(X12,inverse(multiply(inverse(multiply(inverse(X11),X10)),X12))))),X13)) = X13 [superposition 340,100]\n",
      "[SA] new: 641. multiply(X14,multiply(inverse(multiply(X16,inverse(multiply(inverse(multiply(X14,multiply(X15,inverse(X15)))),X16)))),X17)) = X17 [superposition 340,243]\n",
      "[SA] new: 642. multiply(X18,multiply(inverse(multiply(X20,inverse(multiply(inverse(multiply(X18,inverse(multiply(X19,inverse(X19))))),X20)))),X21)) = X21 [superposition 340,486]\n",
      "[SA] new: 643. multiply(X22,multiply(inverse(multiply(X25,inverse(multiply(inverse(multiply(X22,multiply(X23,inverse(multiply(X23,multiply(X24,inverse(X24))))))),X25)))),X26)) = X26 [superposition 340,366]\n",
      "[SA] new: 644. multiply(X27,multiply(inverse(multiply(X30,inverse(multiply(inverse(multiply(X27,multiply(X28,multiply(X29,inverse(multiply(X28,X29)))))),X30)))),X31)) = X31 [superposition 340,340]\n",
      "[SA] new: 645. multiply(inverse(multiply(X32,multiply(X33,inverse(multiply(X32,X33))))),multiply(inverse(identity),X34)) = X34 [superposition 340,7]\n",
      "[SA] new: 646. multiply(multiply(X35,multiply(X36,inverse(multiply(X35,X36)))),multiply(inverse(identity),X37)) = X37 [superposition 340,3]\n",
      "[SA] new: 647. multiply(identity,multiply(inverse(multiply(X38,multiply(X39,inverse(multiply(X38,X39))))),X40)) = X40 [superposition 340,6]\n",
      "[SA] new: 648. multiply(multiply(X41,X42),multiply(inverse(multiply(X41,multiply(X42,multiply(X43,multiply(X44,inverse(multiply(X43,X44))))))),X45)) = X45 [superposition 340,8]\n",
      "[SA] new: 649. multiply(X0,identity) = multiply(X0,multiply(X1,multiply(X2,inverse(multiply(X1,X2))))) [superposition 340,7]\n",
      "[SA] new: 650. multiply(X3,identity) = inverse(multiply(X3,multiply(X4,multiply(X5,inverse(multiply(X4,X5)))))) [superposition 340,3]\n",
      "[SA] new: 651. identity = multiply(X6,inverse(multiply(X6,multiply(X7,multiply(X8,inverse(multiply(X7,X8))))))) [superposition 340,1]\n",
      "[SA] new: 652. multiply(X9,X12) = multiply(multiply(X9,multiply(X10,multiply(X11,inverse(multiply(X10,X11))))),X12) [superposition 340,23]\n",
      "[SA] new: 653. multiply(X13,X16) = multiply(inverse(inverse(multiply(X13,multiply(X14,multiply(X15,inverse(multiply(X14,X15))))))),X16) [superposition 340,49]\n",
      "[SA] new: 654. multiply(X17,multiply(inverse(X20),X21)) = multiply(inverse(multiply(X20,inverse(multiply(X17,multiply(X18,multiply(X19,inverse(multiply(X18,X19)))))))),X21) [superposition 340,98]\n",
      "[SA] new: 655. multiply(X22,multiply(X25,X26)) = multiply(inverse(multiply(inverse(X25),inverse(multiply(X22,multiply(X23,multiply(X24,inverse(multiply(X23,X24)))))))),X26) [superposition 340,100]\n",
      "[SA] new: 656. multiply(X27,X31) = multiply(inverse(multiply(inverse(multiply(X27,multiply(X28,multiply(X29,inverse(multiply(X28,X29)))))),multiply(X30,inverse(X30)))),X31) [superposition 340,243]\n",
      "[SA] new: 657. multiply(inverse(multiply(inverse(multiply(X32,multiply(X33,multiply(X34,inverse(multiply(X33,X34)))))),inverse(multiply(X35,inverse(X35))))),X36) = multiply(X32,X36) [superposition 340,486]\n",
      "[SA] new: 658. multiply(inverse(multiply(inverse(multiply(X37,multiply(X38,multiply(X39,inverse(multiply(X38,X39)))))),multiply(X40,inverse(multiply(X40,multiply(X41,inverse(X41))))))),X42) = multiply(X37,X42) [superposition 340,366]\n",
      "[SA] new: 659. multiply(inverse(multiply(inverse(multiply(X43,multiply(X44,multiply(X45,inverse(multiply(X44,X45)))))),multiply(X46,multiply(X47,inverse(multiply(X46,X47)))))),X48) = multiply(X43,X48) [superposition 340,340]\n",
      "[SA] new: 660. multiply(X52,multiply(X53,multiply(inverse(multiply(inverse(multiply(X49,multiply(X50,multiply(X51,inverse(multiply(X50,X51)))))),multiply(X52,X53))),X54))) = multiply(X49,X54) [superposition 340,121]\n",
      "[SA] new: 661. multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X0,X1))))),X2) = X2 [superposition 340,98]\n",
      "[SA] new: 662. multiply(X10,multiply(inverse(multiply(X7,multiply(inverse(multiply(X7,multiply(X8,multiply(X9,inverse(multiply(X8,X9)))))),X10))),X11)) = X11 [superposition 340,121]\n",
      "[SA] new: 663. multiply(inverse(multiply(identity,multiply(X12,multiply(X13,inverse(multiply(X12,X13)))))),X14) = X14 [superposition 340,6]\n",
      "[SA] new: 664. multiply(X15,multiply(X16,multiply(inverse(multiply(multiply(X15,X16),multiply(X17,multiply(X18,inverse(multiply(X17,X18)))))),X19))) = X19 [superposition 340,8]\n",
      "[SA] new: 665. multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X0,X1))))),X2) = X2 [superposition 98,340]\n",
      "[SA] new: 666. multiply(X3,multiply(inverse(multiply(X3,multiply(X4,multiply(X5,inverse(multiply(X4,X5)))))),multiply(X6,multiply(inverse(X6),X7)))) = X7 [superposition 121,340]\n",
      "[SA] new: 667. multiply(multiply(X8,inverse(multiply(X9,X8))),multiply(X9,X10)) = X10 [superposition 121,340]\n",
      "[SA] new: 668. multiply(X14,multiply(inverse(multiply(X14,multiply(X11,inverse(multiply(X11,multiply(X12,multiply(X13,inverse(multiply(X12,X13))))))))),X15)) = X15 [superposition 121,340]\n",
      "[SA] new: 669. multiply(X19,multiply(inverse(multiply(X16,multiply(inverse(multiply(X16,multiply(X17,multiply(X18,inverse(multiply(X17,X18)))))),X19))),X20)) = X20 [superposition 121,340]\n",
      "[SA] new: 670. multiply(X24,multiply(inverse(multiply(X24,multiply(X21,inverse(inverse(inverse(multiply(X21,multiply(X22,multiply(X23,inverse(multiply(X22,X23))))))))))),X25)) = X25 [superposition 366,340]\n",
      "[SA] new: 671. multiply(X26,multiply(inverse(inverse(inverse(multiply(X26,multiply(X27,multiply(X28,inverse(multiply(X27,X28)))))))),X29)) = X29 [superposition 243,340]\n",
      "[SA] new: 672. multiply(X30,multiply(inverse(inverse(multiply(inverse(multiply(X30,multiply(X31,multiply(X32,inverse(multiply(X31,X32)))))),multiply(X33,inverse(X33))))),X34)) = X34 [superposition 366,340]\n",
      "[SA] new: 673. multiply(inverse(multiply(identity,multiply(X35,multiply(X36,inverse(multiply(X35,X36)))))),X37) = X37 [superposition 6,340]\n",
      "[SA] new: 674. multiply(X42,multiply(X38,multiply(multiply(inverse(multiply(X38,multiply(X39,multiply(X40,inverse(multiply(X39,X40)))))),X41),multiply(inverse(multiply(X42,X41)),X43)))) = X43 [superposition 121,340]\n",
      "[SA] new: 675. multiply(inverse(X44),X48) = multiply(multiply(inverse(multiply(X44,multiply(X45,multiply(X46,inverse(multiply(X45,X46)))))),X47),multiply(inverse(X47),X48)) [superposition 98,340]\n",
      "[SA] new: 676. multiply(inverse(multiply(X49,multiply(X50,multiply(X51,inverse(multiply(X50,X51)))))),X52) = multiply(inverse(X49),X52) [superposition 23,340]\n",
      "[SA] new: 677. multiply(X56,X57) = multiply(X53,multiply(multiply(inverse(multiply(X53,multiply(X54,multiply(X55,inverse(multiply(X54,X55)))))),X56),X57)) [superposition 8,340]\n",
      "[SA] new: 678. multiply(X58,X62) = multiply(multiply(inverse(multiply(inverse(X58),multiply(X59,multiply(X60,inverse(multiply(X59,X60)))))),X61),multiply(inverse(X61),X62)) [superposition 100,340]\n",
      "[SA] new: 679. multiply(X63,X66) = multiply(inverse(multiply(inverse(X63),multiply(X64,multiply(X65,inverse(multiply(X64,X65)))))),X66) [superposition 49,340]\n",
      "[SA] new: 680. multiply(X67,multiply(X68,multiply(inverse(multiply(multiply(X67,X68),multiply(X69,multiply(X70,inverse(multiply(X69,X70)))))),X71))) = X71 [superposition 8,340]\n",
      "[SA] new: 681. multiply(inverse(X72),multiply(inverse(multiply(inverse(multiply(X72,X73)),multiply(X74,multiply(X75,inverse(multiply(X74,X75)))))),X76)) = multiply(X73,X76) [superposition 98,340]\n",
      "[SA] new: 682. multiply(X77,multiply(inverse(multiply(inverse(multiply(inverse(X77),X78)),multiply(X79,multiply(X80,inverse(multiply(X79,X80)))))),X81)) = multiply(X78,X81) [superposition 100,340]\n",
      "[SA] new: 683. multiply(inverse(multiply(inverse(multiply(X82,multiply(X83,inverse(X83)))),multiply(X84,multiply(X85,inverse(multiply(X84,X85)))))),X86) = multiply(X82,X86) [superposition 243,340]\n",
      "[SA] new: 684. multiply(inverse(multiply(inverse(multiply(X87,inverse(multiply(X88,inverse(X88))))),multiply(X89,multiply(X90,inverse(multiply(X89,X90)))))),X91) = multiply(X87,X91) [superposition 486,340]\n",
      "[SA] new: 685. multiply(inverse(multiply(inverse(multiply(X92,multiply(X93,X94))),multiply(X95,multiply(X96,inverse(multiply(X95,X96)))))),X97) = multiply(X92,multiply(X93,multiply(X94,X97))) [superposition 121,340]\n",
      "[SA] new: 686. multiply(inverse(multiply(inverse(multiply(X98,multiply(X99,inverse(multiply(X99,multiply(X100,inverse(X100))))))),multiply(X101,multiply(X102,inverse(multiply(X101,X102)))))),X103) = multiply(X98,X103) [superposition 366,340]\n",
      "[SA] new: 687. multiply(X1,multiply(inverse(multiply(X1,inverse(identity))),X2)) = X2 [forward demodulation 610,23]\n",
      "[SA] new: 688. multiply(X1,multiply(inverse(multiply(X1,identity)),X2)) = X2 [forward demodulation 687,10]\n",
      "[SA] new: 689. multiply(X1,multiply(inverse(X1),X2)) = X2 [forward demodulation 688,1]\n",
      "[SA] new: 690. multiply(X4,multiply(inverse(multiply(X4,multiply(X3,multiply(X3,identity)))),X5)) = X5 [forward demodulation 611,10]\n",
      "[SA] new: 691. multiply(X4,multiply(inverse(multiply(X4,multiply(X3,X3))),X5)) = X5 [forward demodulation 690,1]\n",
      "[SA] new: 692. multiply(X4,multiply(inverse(multiply(X4,identity)),X5)) = X5 [forward demodulation 691,3]\n",
      "[SA] new: 693. multiply(X4,multiply(inverse(X4),X5)) = X5 [forward demodulation 692,1]\n",
      "[SA] new: 694. multiply(X7,multiply(inverse(multiply(X7,multiply(X6,inverse(X6)))),X8)) = X8 [forward demodulation 612,6]\n",
      "[SA] new: 695. multiply(X12,multiply(inverse(multiply(X12,multiply(X9,multiply(X10,multiply(X11,inverse(multiply(X9,multiply(X10,X11)))))))),X13)) = X13 [forward demodulation 613,8]\n",
      "[SA] new: 696. multiply(X15,multiply(inverse(multiply(X15,multiply(X14,inverse(X14)))),X16)) = X16 [forward demodulation 614,6]\n",
      "[SA] new: 697. multiply(X19,multiply(inverse(multiply(X19,multiply(inverse(X17),multiply(X17,multiply(X18,inverse(X18)))))),X20)) = X20 [forward demodulation 615,8]\n",
      "[SA] new: 698. multiply(X19,multiply(inverse(multiply(X19,multiply(X18,inverse(X18)))),X20)) = X20 [forward demodulation 697,23]\n",
      "[SA] new: 699. multiply(X23,multiply(inverse(multiply(X23,multiply(X21,multiply(inverse(X21),multiply(X22,inverse(X22)))))),X24)) = X24 [forward demodulation 616,8]\n",
      "[SA] new: 700. multiply(X23,multiply(inverse(multiply(X23,multiply(X22,inverse(X22)))),X24)) = X24 [forward demodulation 699,49]\n",
      "[SA] new: 701. multiply(X28,multiply(inverse(multiply(X28,multiply(X25,multiply(inverse(multiply(X26,X25)),multiply(X27,inverse(multiply(inverse(X26),X27))))))),X29)) = X29 [forward demodulation 617,8]\n",
      "[SA] new: 702. multiply(X28,multiply(inverse(multiply(X28,multiply(inverse(X26),multiply(X27,inverse(multiply(inverse(X26),X27)))))),X29)) = X29 [forward demodulation 701,98]\n",
      "[SA] new: 703. multiply(X33,multiply(inverse(multiply(X33,multiply(X30,multiply(inverse(multiply(inverse(X31),X30)),multiply(X32,inverse(multiply(X31,X32))))))),X34)) = X34 [forward demodulation 618,8]\n",
      "[SA] new: 704. multiply(X33,multiply(inverse(multiply(X33,multiply(inverse(inverse(X31)),multiply(X32,inverse(multiply(X31,X32)))))),X34)) = X34 [forward demodulation 703,98]\n",
      "[SA] new: 705. multiply(X33,multiply(inverse(multiply(X33,multiply(X31,multiply(X32,inverse(multiply(X31,X32)))))),X34)) = X34 [forward demodulation 704,50]\n",
      "[SA] new: 706. multiply(X38,multiply(inverse(multiply(X38,multiply(X35,multiply(inverse(multiply(X35,multiply(X36,inverse(X36)))),multiply(X37,inverse(X37)))))),X39)) = X39 [forward demodulation 619,8]\n",
      "[SA] new: 707. multiply(X38,multiply(inverse(multiply(X38,multiply(X37,inverse(X37)))),X39)) = X39 [forward demodulation 706,243]\n",
      "[SA] new: 708. multiply(X43,multiply(inverse(multiply(X43,multiply(X40,multiply(inverse(multiply(X40,inverse(multiply(X41,inverse(X41))))),multiply(X42,inverse(X42)))))),X44)) = X44 [forward demodulation 620,8]\n",
      "[SA] new: 709. multiply(X43,multiply(inverse(multiply(X43,multiply(X42,inverse(X42)))),X44)) = X44 [forward demodulation 708,486]\n",
      "[SA] new: 710. multiply(X49,multiply(inverse(multiply(X49,multiply(X45,multiply(inverse(multiply(X45,multiply(X46,inverse(multiply(X46,multiply(X47,inverse(X47))))))),multiply(X48,inverse(X48)))))),X50)) = X50 [forward demodulation 621,8]\n",
      "[SA] new: 711. multiply(X49,multiply(inverse(multiply(X49,multiply(X48,inverse(X48)))),X50)) = X50 [forward demodulation 710,366]\n",
      "[SA] new: 712. multiply(X55,multiply(inverse(multiply(X55,multiply(X51,multiply(inverse(multiply(X51,multiply(X52,multiply(X53,inverse(multiply(X52,X53)))))),multiply(X54,inverse(X54)))))),X56)) = X56 [forward demodulation 622,8]\n",
      "[SA] new: 713. multiply(X55,multiply(inverse(multiply(X55,multiply(X54,inverse(X54)))),X56)) = X56 [forward demodulation 712,340]\n",
      "[SA] new: 714. multiply(X61,multiply(inverse(multiply(X61,multiply(X57,multiply(X58,multiply(multiply(X59,multiply(inverse(multiply(X57,multiply(X58,X59))),X60)),inverse(X60)))))),X62)) = X62 [forward demodulation 623,8]\n",
      "[SA] new: 715. multiply(X61,multiply(inverse(multiply(X61,multiply(X57,multiply(X58,multiply(X59,multiply(multiply(inverse(multiply(X57,multiply(X58,X59))),X60),inverse(X60))))))),X62)) = X62 [forward demodulation 714,8]\n",
      "[SA] new: 716. multiply(X61,multiply(inverse(multiply(X61,multiply(X57,multiply(X58,multiply(X59,multiply(inverse(multiply(X57,multiply(X58,X59))),multiply(X60,inverse(X60)))))))),X62)) = X62 [forward demodulation 715,8]\n",
      "[SA] new: 717. multiply(X61,multiply(inverse(multiply(X61,multiply(X60,inverse(X60)))),X62)) = X62 [forward demodulation 716,121]\n",
      "[SA] new: 718. multiply(X1,multiply(inverse(multiply(X1,multiply(X0,inverse(X0)))),X2)) = X2 [forward demodulation 625,1]\n",
      "[SA] new: 719. multiply(X1,multiply(inverse(multiply(X1,inverse(identity))),X2)) = X2 [forward demodulation 627,7]\n",
      "[SA] new: 720. multiply(X1,multiply(inverse(multiply(X1,identity)),X2)) = X2 [forward demodulation 719,10]\n",
      "[SA] new: 721. multiply(X1,multiply(inverse(X1),X2)) = X2 [forward demodulation 720,1]\n",
      "[SA] new: 722. multiply(X33,multiply(inverse(multiply(X33,multiply(X32,inverse(X32)))),X34)) = X34 [forward demodulation 635,6]\n",
      "[SA] new: 723. multiply(X38,multiply(inverse(multiply(X38,multiply(X35,multiply(X36,multiply(X37,inverse(multiply(X35,multiply(X36,X37)))))))),X39)) = X39 [forward demodulation 636,8]\n",
      "[SA] new: 724. multiply(inverse(multiply(X32,multiply(X33,inverse(multiply(X32,X33))))),multiply(identity,X34)) = X34 [forward demodulation 645,10]\n",
      "[SA] new: 725. multiply(inverse(multiply(X32,multiply(X33,inverse(multiply(X32,X33))))),X34) = X34 [forward demodulation 724,6]\n",
      "[SA] new: 726. multiply(multiply(X35,multiply(X36,inverse(multiply(X35,X36)))),multiply(identity,X37)) = X37 [forward demodulation 646,10]\n",
      "[SA] new: 727. multiply(multiply(X35,multiply(X36,inverse(multiply(X35,X36)))),X37) = X37 [forward demodulation 726,6]\n",
      "[SA] new: 728. multiply(X0,multiply(X1,multiply(X2,inverse(multiply(X1,X2))))) = X0 [forward demodulation 649,1]\n",
      "[SA] new: 729. inverse(multiply(X3,multiply(X4,multiply(X5,inverse(multiply(X4,X5)))))) = X3 [forward demodulation 650,1]\n",
      "[SA] new: 730. multiply(X13,X16) = multiply(multiply(X13,multiply(X14,multiply(X15,inverse(multiply(X14,X15))))),X16) [forward demodulation 653,50]\n",
      "[SA] new: 731. multiply(X10,multiply(inverse(X10),X11)) = X11 [forward demodulation 662,340]\n",
      "[SA] new: 732. multiply(inverse(multiply(X12,multiply(X13,inverse(multiply(X12,X13))))),X14) = X14 [forward demodulation 663,6]\n",
      "[SA] new: 733. multiply(X15,multiply(X16,multiply(inverse(multiply(X15,multiply(X16,multiply(X17,multiply(X18,inverse(multiply(X17,X18))))))),X19))) = X19 [forward demodulation 664,8]\n",
      "[SA] new: 734. multiply(X3,multiply(inverse(multiply(X3,multiply(X4,multiply(X5,inverse(multiply(X4,X5)))))),X7)) = X7 [forward demodulation 666,49]\n",
      "[SA] new: 735. multiply(X19,multiply(inverse(X19),X20)) = X20 [forward demodulation 669,340]\n",
      "[SA] new: 736. multiply(X24,multiply(inverse(multiply(X24,multiply(X21,inverse(multiply(X21,multiply(X22,multiply(X23,inverse(multiply(X22,X23))))))))),X25)) = X25 [forward demodulation 670,50]\n",
      "[SA] new: 737. multiply(X26,multiply(inverse(multiply(X26,multiply(X27,multiply(X28,inverse(multiply(X27,X28)))))),X29)) = X29 [forward demodulation 671,50]\n",
      "[SA] new: 738. multiply(X30,multiply(multiply(inverse(multiply(X30,multiply(X31,multiply(X32,inverse(multiply(X31,X32)))))),multiply(X33,inverse(X33))),X34)) = X34 [forward demodulation 672,50]\n",
      "[SA] new: 739. multiply(X30,multiply(inverse(multiply(X30,multiply(X31,multiply(X32,inverse(multiply(X31,X32)))))),multiply(multiply(X33,inverse(X33)),X34))) = X34 [forward demodulation 738,8]\n",
      "[SA] new: 740. multiply(X30,multiply(inverse(multiply(X30,multiply(X31,multiply(X32,inverse(multiply(X31,X32)))))),multiply(X33,multiply(inverse(X33),X34)))) = X34 [forward demodulation 739,8]\n",
      "[SA] new: 741. multiply(X30,multiply(inverse(multiply(X30,multiply(X31,multiply(X32,inverse(multiply(X31,X32)))))),X34)) = X34 [forward demodulation 740,49]\n",
      "[SA] new: 742. multiply(inverse(multiply(X35,multiply(X36,inverse(multiply(X35,X36))))),X37) = X37 [forward demodulation 673,6]\n",
      "[SA] new: 743. multiply(X42,multiply(X38,multiply(inverse(multiply(X38,multiply(X39,multiply(X40,inverse(multiply(X39,X40)))))),multiply(X41,multiply(inverse(multiply(X42,X41)),X43))))) = X43 [forward demodulation 674,8]\n",
      "[SA] new: 744. multiply(X42,multiply(X41,multiply(inverse(multiply(X42,X41)),X43))) = X43 [forward demodulation 743,340]\n",
      "[SA] new: 745. multiply(X42,multiply(inverse(X42),X43)) = X43 [forward demodulation 744,98]\n",
      "[SA] new: 746. multiply(X56,X57) = multiply(X53,multiply(inverse(multiply(X53,multiply(X54,multiply(X55,inverse(multiply(X54,X55)))))),multiply(X56,X57))) [forward demodulation 677,8]\n",
      "[SA] new: 747. multiply(X67,multiply(X68,multiply(inverse(multiply(X67,multiply(X68,multiply(X69,multiply(X70,inverse(multiply(X69,X70))))))),X71))) = X71 [forward demodulation 680,8]\n",
      "[SA] new: 748. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X1,inverse(multiply(identity,identity))))))),X2)) = X2 [superposition 552,10]\n",
      "[SA] new: 749. multiply(X4,multiply(inverse(multiply(X4,multiply(X5,inverse(multiply(X5,inverse(multiply(inverse(X3),X3))))))),X6)) = X6 [superposition 552,50]\n",
      "[SA] new: 750. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X1,inverse(inverse(identity))))))),X2)) = X2 [superposition 552,6]\n",
      "[SA] new: 751. multiply(X5,multiply(inverse(multiply(X5,multiply(X6,inverse(multiply(X6,inverse(multiply(X3,multiply(X4,inverse(multiply(X3,X4)))))))))),X7)) = X7 [superposition 552,8]\n",
      "[SA] new: 752. multiply(X1,multiply(inverse(multiply(X1,multiply(inverse(inverse(multiply(X0,inverse(X0)))),inverse(identity)))),X2)) = X2 [superposition 552,7]\n",
      "[SA] new: 753. multiply(X4,multiply(inverse(multiply(X4,multiply(inverse(multiply(X3,inverse(X3))),inverse(identity)))),X5)) = X5 [superposition 552,3]\n",
      "[SA] new: 754. multiply(X7,multiply(inverse(multiply(X7,multiply(identity,inverse(inverse(multiply(X6,inverse(X6))))))),X8)) = X8 [superposition 552,6]\n",
      "[SA] new: 755. multiply(X12,multiply(inverse(multiply(X12,multiply(multiply(X9,X10),inverse(multiply(X9,multiply(X10,inverse(multiply(X11,inverse(X11))))))))),X13)) = X13 [superposition 552,8]\n",
      "[SA] new: 756. multiply(X1,multiply(inverse(multiply(X1,inverse(multiply(identity,inverse(multiply(X0,inverse(X0))))))),X2)) = X2 [superposition 552,6]\n",
      "[SA] new: 757. multiply(X6,multiply(inverse(multiply(X6,multiply(X3,multiply(X4,inverse(multiply(multiply(X3,X4),inverse(multiply(X5,inverse(X5))))))))),X7)) = X7 [superposition 552,8]\n",
      "[SA] new: 758. multiply(inverse(X0),multiply(inverse(inverse(multiply(X0,inverse(multiply(X1,inverse(X1)))))),X2)) = X2 [superposition 552,23]\n",
      "[SA] new: 759. multiply(X3,multiply(inverse(inverse(multiply(inverse(X3),inverse(multiply(X4,inverse(X4)))))),X5)) = X5 [superposition 552,49]\n",
      "[SA] new: 760. multiply(X6,multiply(inverse(multiply(inverse(X7),inverse(multiply(inverse(multiply(X7,X6)),inverse(multiply(X8,inverse(X8))))))),X9)) = X9 [superposition 552,98]\n",
      "[SA] new: 761. multiply(X10,multiply(inverse(multiply(X11,inverse(multiply(inverse(multiply(inverse(X11),X10)),inverse(multiply(X12,inverse(X12))))))),X13)) = X13 [superposition 552,100]\n",
      "[SA] new: 762. multiply(X14,multiply(inverse(inverse(multiply(inverse(multiply(X14,multiply(X15,inverse(X15)))),inverse(multiply(X16,inverse(X16)))))),X17)) = X17 [superposition 552,243]\n",
      "[SA] new: 763. multiply(X18,multiply(inverse(inverse(multiply(inverse(multiply(X18,inverse(multiply(X19,inverse(X19))))),inverse(multiply(X20,inverse(X20)))))),X21)) = X21 [superposition 552,486]\n",
      "[SA] new: 764. multiply(X22,multiply(inverse(inverse(multiply(inverse(multiply(X22,multiply(X23,inverse(multiply(X23,multiply(X24,inverse(X24))))))),inverse(multiply(X25,inverse(X25)))))),X26)) = X26 [superposition 552,366]\n",
      "[SA] new: 765. multiply(X27,multiply(inverse(inverse(multiply(inverse(multiply(X27,multiply(X28,inverse(multiply(X28,inverse(multiply(X29,inverse(X29)))))))),inverse(multiply(X30,inverse(X30)))))),X31)) = X31 [superposition 552,552]\n",
      "[SA] new: 766. multiply(X32,multiply(inverse(inverse(multiply(inverse(multiply(X32,multiply(X33,multiply(X34,inverse(multiply(X33,X34)))))),inverse(multiply(X35,inverse(X35)))))),X36)) = X36 [superposition 552,340]\n",
      "[SA] new: 767. multiply(inverse(multiply(X37,inverse(multiply(X37,inverse(multiply(X38,inverse(X38))))))),multiply(inverse(identity),X39)) = X39 [superposition 552,7]\n",
      "[SA] new: 768. multiply(multiply(X40,inverse(multiply(X40,inverse(multiply(X41,inverse(X41)))))),multiply(inverse(identity),X42)) = X42 [superposition 552,3]\n",
      "[SA] new: 769. multiply(identity,multiply(inverse(multiply(X43,inverse(multiply(X43,inverse(multiply(X44,inverse(X44))))))),X45)) = X45 [superposition 552,6]\n",
      "[SA] new: 770. multiply(multiply(X46,X47),multiply(inverse(multiply(X46,multiply(X47,multiply(X48,inverse(multiply(X48,inverse(multiply(X49,inverse(X49))))))))),X50)) = X50 [superposition 552,8]\n",
      "[SA] new: 771. multiply(X0,identity) = multiply(X0,multiply(X1,inverse(multiply(X1,inverse(multiply(X2,inverse(X2))))))) [superposition 552,7]\n",
      "[SA] new: 772. multiply(X3,identity) = inverse(multiply(X3,multiply(X4,inverse(multiply(X4,inverse(multiply(X5,inverse(X5)))))))) [superposition 552,3]\n",
      "[SA] new: 773. identity = multiply(X6,inverse(multiply(X6,multiply(X7,inverse(multiply(X7,inverse(multiply(X8,inverse(X8))))))))) [superposition 552,1]\n",
      "[SA] new: 774. multiply(X9,X12) = multiply(multiply(X9,multiply(X10,inverse(multiply(X10,inverse(multiply(X11,inverse(X11))))))),X12) [superposition 552,23]\n",
      "[SA] new: 775. multiply(X13,X16) = multiply(inverse(inverse(multiply(X13,multiply(X14,inverse(multiply(X14,inverse(multiply(X15,inverse(X15))))))))),X16) [superposition 552,49]\n",
      "[SA] new: 776. multiply(X17,multiply(inverse(X20),X21)) = multiply(inverse(multiply(X20,inverse(multiply(X17,multiply(X18,inverse(multiply(X18,inverse(multiply(X19,inverse(X19)))))))))),X21) [superposition 552,98]\n",
      "[SA] new: 777. multiply(X22,multiply(X25,X26)) = multiply(inverse(multiply(inverse(X25),inverse(multiply(X22,multiply(X23,inverse(multiply(X23,inverse(multiply(X24,inverse(X24)))))))))),X26) [superposition 552,100]\n",
      "[SA] new: 778. multiply(X27,X31) = multiply(inverse(multiply(inverse(multiply(X27,multiply(X28,inverse(multiply(X28,inverse(multiply(X29,inverse(X29)))))))),multiply(X30,inverse(X30)))),X31) [superposition 552,243]\n",
      "[SA] new: 779. multiply(X32,X36) = multiply(inverse(multiply(inverse(multiply(X32,multiply(X33,inverse(multiply(X33,inverse(multiply(X34,inverse(X34)))))))),inverse(multiply(X35,inverse(X35))))),X36) [superposition 552,486]\n",
      "[SA] new: 780. multiply(X37,X42) = multiply(inverse(multiply(inverse(multiply(X37,multiply(X38,inverse(multiply(X38,inverse(multiply(X39,inverse(X39)))))))),multiply(X40,inverse(multiply(X40,multiply(X41,inverse(X41))))))),X42) [superposition 552,366]\n",
      "[SA] new: 781. multiply(X43,X48) = multiply(inverse(multiply(inverse(multiply(X43,multiply(X44,inverse(multiply(X44,inverse(multiply(X45,inverse(X45)))))))),multiply(X46,inverse(multiply(X46,inverse(multiply(X47,inverse(X47)))))))),X48) [superposition 552,552]\n",
      "[SA] new: 782. multiply(X49,X54) = multiply(inverse(multiply(inverse(multiply(X49,multiply(X50,inverse(multiply(X50,inverse(multiply(X51,inverse(X51)))))))),multiply(X52,multiply(X53,inverse(multiply(X52,X53)))))),X54) [superposition 552,340]\n",
      "[SA] new: 783. multiply(X58,multiply(X59,multiply(inverse(multiply(inverse(multiply(X55,multiply(X56,inverse(multiply(X56,inverse(multiply(X57,inverse(X57)))))))),multiply(X58,X59))),X60))) = multiply(X55,X60) [superposition 552,121]\n",
      "[SA] new: 784. multiply(inverse(multiply(X0,inverse(multiply(X0,inverse(multiply(X1,inverse(X1))))))),X2) = X2 [superposition 552,98]\n",
      "[SA] new: 785. multiply(X10,multiply(inverse(multiply(X7,multiply(inverse(multiply(X7,multiply(X8,inverse(multiply(X8,inverse(multiply(X9,inverse(X9)))))))),X10))),X11)) = X11 [superposition 552,121]\n",
      "[SA] new: 786. multiply(inverse(multiply(identity,multiply(X12,inverse(multiply(X12,inverse(multiply(X13,inverse(X13)))))))),X14) = X14 [superposition 552,6]\n",
      "[SA] new: 787. multiply(X15,multiply(X16,multiply(inverse(multiply(multiply(X15,X16),multiply(X17,inverse(multiply(X17,inverse(multiply(X18,inverse(X18)))))))),X19))) = X19 [superposition 552,8]\n",
      "[SA] new: 788. multiply(inverse(multiply(X0,inverse(multiply(X0,inverse(multiply(X1,inverse(X1))))))),X2) = X2 [superposition 98,552]\n",
      "[SA] new: 789. multiply(X3,multiply(inverse(multiply(X3,multiply(X4,inverse(multiply(X4,inverse(multiply(X5,inverse(X5)))))))),multiply(X6,multiply(inverse(X6),X7)))) = X7 [superposition 121,552]\n",
      "[SA] new: 790. multiply(inverse(multiply(X8,inverse(multiply(X9,inverse(X9))))),multiply(X8,X10)) = X10 [superposition 121,552]\n",
      "[SA] new: 791. multiply(X14,multiply(inverse(multiply(X14,multiply(X11,inverse(multiply(X11,multiply(X12,inverse(multiply(X12,inverse(multiply(X13,inverse(X13))))))))))),X15)) = X15 [superposition 121,552]\n",
      "[SA] new: 792. multiply(X19,multiply(inverse(multiply(X16,multiply(inverse(multiply(X16,multiply(X17,inverse(multiply(X17,inverse(multiply(X18,inverse(X18)))))))),X19))),X20)) = X20 [superposition 121,552]\n",
      "[SA] new: 793. multiply(X21,multiply(inverse(multiply(X24,inverse(multiply(inverse(multiply(X21,multiply(X22,inverse(multiply(X22,inverse(multiply(X23,inverse(X23)))))))),X24)))),X25)) = X25 [superposition 340,552]\n",
      "[SA] new: 794. multiply(X29,multiply(inverse(multiply(X29,multiply(X26,inverse(inverse(inverse(multiply(X26,multiply(X27,inverse(multiply(X27,inverse(multiply(X28,inverse(X28))))))))))))),X30)) = X30 [superposition 366,552]\n",
      "[SA] new: 795. multiply(X31,multiply(inverse(inverse(inverse(multiply(X31,multiply(X32,inverse(multiply(X32,inverse(multiply(X33,inverse(X33)))))))))),X34)) = X34 [superposition 243,552]\n",
      "[SA] new: 796. multiply(X35,multiply(inverse(inverse(multiply(inverse(multiply(X35,multiply(X36,inverse(multiply(X36,inverse(multiply(X37,inverse(X37)))))))),multiply(X38,inverse(X38))))),X39)) = X39 [superposition 366,552]\n",
      "[SA] new: 797. multiply(X43,multiply(inverse(multiply(X43,inverse(multiply(X40,inverse(multiply(X40,multiply(X41,inverse(multiply(X41,inverse(multiply(X42,inverse(X42)))))))))))),X44)) = X44 [superposition 340,552]\n",
      "[SA] new: 798. multiply(inverse(multiply(identity,multiply(X45,inverse(multiply(X45,inverse(multiply(X46,inverse(X46)))))))),X47) = X47 [superposition 6,552]\n",
      "[SA] new: 799. multiply(X52,multiply(inverse(multiply(X52,multiply(X48,multiply(multiply(inverse(multiply(X48,multiply(X49,inverse(multiply(X49,inverse(multiply(X50,inverse(X50)))))))),X51),inverse(X51))))),X53)) = X53 [superposition 340,552]\n",
      "[SA] new: 800. multiply(X58,multiply(X54,multiply(multiply(inverse(multiply(X54,multiply(X55,inverse(multiply(X55,inverse(multiply(X56,inverse(X56)))))))),X57),multiply(inverse(multiply(X58,X57)),X59)))) = X59 [superposition 121,552]\n",
      "[SA] new: 801. multiply(inverse(X60),X64) = multiply(multiply(inverse(multiply(X60,multiply(X61,inverse(multiply(X61,inverse(multiply(X62,inverse(X62)))))))),X63),multiply(inverse(X63),X64)) [superposition 98,552]\n",
      "[SA] new: 802. multiply(inverse(multiply(X65,multiply(X66,inverse(multiply(X66,inverse(multiply(X67,inverse(X67)))))))),X68) = multiply(inverse(X65),X68) [superposition 23,552]\n",
      "[SA] new: 803. multiply(X72,X73) = multiply(X69,multiply(multiply(inverse(multiply(X69,multiply(X70,inverse(multiply(X70,inverse(multiply(X71,inverse(X71)))))))),X72),X73)) [superposition 8,552]\n",
      "[SA] new: 804. multiply(X74,X78) = multiply(multiply(inverse(multiply(inverse(X74),multiply(X75,inverse(multiply(X75,inverse(multiply(X76,inverse(X76)))))))),X77),multiply(inverse(X77),X78)) [superposition 100,552]\n",
      "[SA] new: 805. multiply(X79,X82) = multiply(inverse(multiply(inverse(X79),multiply(X80,inverse(multiply(X80,inverse(multiply(X81,inverse(X81)))))))),X82) [superposition 49,552]\n",
      "[SA] new: 806. multiply(X83,multiply(X84,multiply(inverse(multiply(multiply(X83,X84),multiply(X85,inverse(multiply(X85,inverse(multiply(X86,inverse(X86)))))))),X87))) = X87 [superposition 8,552]\n",
      "[SA] new: 807. multiply(inverse(X88),multiply(inverse(multiply(inverse(multiply(X88,X89)),multiply(X90,inverse(multiply(X90,inverse(multiply(X91,inverse(X91)))))))),X92)) = multiply(X89,X92) [superposition 98,552]\n",
      "[SA] new: 808. multiply(X94,X97) = multiply(X93,multiply(inverse(multiply(inverse(multiply(inverse(X93),X94)),multiply(X95,inverse(multiply(X95,inverse(multiply(X96,inverse(X96)))))))),X97)) [superposition 100,552]\n",
      "[SA] new: 809. multiply(inverse(multiply(inverse(multiply(X98,multiply(X99,inverse(X99)))),multiply(X100,inverse(multiply(X100,inverse(multiply(X101,inverse(X101)))))))),X102) = multiply(X98,X102) [superposition 243,552]\n",
      "[SA] new: 810. multiply(inverse(multiply(inverse(multiply(X103,inverse(multiply(X104,inverse(X104))))),multiply(X105,inverse(multiply(X105,inverse(multiply(X106,inverse(X106)))))))),X107) = multiply(X103,X107) [superposition 486,552]\n",
      "[SA] new: 811. multiply(inverse(multiply(inverse(multiply(X108,multiply(X109,X110))),multiply(X111,inverse(multiply(X111,inverse(multiply(X112,inverse(X112)))))))),X113) = multiply(X108,multiply(X109,multiply(X110,X113))) [superposition 121,552]\n",
      "[SA] new: 812. multiply(inverse(multiply(inverse(multiply(X114,multiply(X115,inverse(multiply(X115,multiply(X116,inverse(X116))))))),multiply(X117,inverse(multiply(X117,inverse(multiply(X118,inverse(X118)))))))),X119) = multiply(X114,X119) [superposition 366,552]\n",
      "[SA] new: 813. multiply(inverse(multiply(inverse(multiply(X120,multiply(X121,multiply(X122,inverse(multiply(X121,X122)))))),multiply(X123,inverse(multiply(X123,inverse(multiply(X124,inverse(X124)))))))),X125) = multiply(X120,X125) [superposition 340,552]\n",
      "[SA] new: 814. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X1,inverse(identity)))))),X2)) = X2 [forward demodulation 748,3]\n",
      "[SA] new: 815. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X1,identity))))),X2)) = X2 [forward demodulation 814,10]\n",
      "[SA] new: 816. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(X1)))),X2)) = X2 [forward demodulation 815,1]\n",
      "[SA] new: 817. multiply(X4,multiply(inverse(multiply(X4,multiply(X5,inverse(multiply(X5,inverse(identity)))))),X6)) = X6 [forward demodulation 749,7]\n",
      "[SA] new: 818. multiply(X4,multiply(inverse(multiply(X4,multiply(X5,inverse(multiply(X5,identity))))),X6)) = X6 [forward demodulation 817,10]\n",
      "[SA] new: 819. multiply(X4,multiply(inverse(multiply(X4,multiply(X5,inverse(X5)))),X6)) = X6 [forward demodulation 818,1]\n",
      "[SA] new: 820. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(multiply(X1,identity))))),X2)) = X2 [forward demodulation 750,50]\n",
      "[SA] new: 821. multiply(X0,multiply(inverse(multiply(X0,multiply(X1,inverse(X1)))),X2)) = X2 [forward demodulation 820,1]\n",
      "[SA] new: 822. multiply(X1,multiply(inverse(multiply(X1,multiply(inverse(inverse(multiply(X0,inverse(X0)))),identity))),X2)) = X2 [forward demodulation 752,10]\n",
      "[SA] new: 823. multiply(X1,multiply(inverse(multiply(X1,inverse(inverse(multiply(X0,inverse(X0)))))),X2)) = X2 [forward demodulation 822,1]\n",
      "[SA] new: 824. multiply(X1,multiply(inverse(multiply(X1,multiply(X0,inverse(X0)))),X2)) = X2 [forward demodulation 823,50]\n",
      "[SA] new: 825. multiply(X4,multiply(inverse(multiply(X4,multiply(inverse(multiply(X3,inverse(X3))),identity))),X5)) = X5 [forward demodulation 753,10]\n",
      "[SA] new: 826. multiply(X4,multiply(inverse(multiply(X4,inverse(multiply(X3,inverse(X3))))),X5)) = X5 [forward demodulation 825,1]\n",
      "[SA] new: 827. multiply(X7,multiply(inverse(multiply(X7,inverse(inverse(multiply(X6,inverse(X6)))))),X8)) = X8 [forward demodulation 754,6]\n",
      "[SA] new: 828. multiply(X7,multiply(inverse(multiply(X7,multiply(X6,inverse(X6)))),X8)) = X8 [forward demodulation 827,50]\n",
      "[SA] new: 829. multiply(X12,multiply(inverse(multiply(X12,multiply(X9,multiply(X10,inverse(multiply(X9,multiply(X10,inverse(multiply(X11,inverse(X11)))))))))),X13)) = X13 [forward demodulation 755,8]\n",
      "[SA] new: 830. multiply(X1,multiply(inverse(multiply(X1,inverse(inverse(multiply(X0,inverse(X0)))))),X2)) = X2 [forward demodulation 756,6]\n",
      "[SA] new: 831. multiply(X1,multiply(inverse(multiply(X1,multiply(X0,inverse(X0)))),X2)) = X2 [forward demodulation 830,50]\n",
      "[SA] new: 832. multiply(X6,multiply(inverse(multiply(X6,multiply(X3,multiply(X4,inverse(multiply(X3,multiply(X4,inverse(multiply(X5,inverse(X5)))))))))),X7)) = X7 [forward demodulation 757,8]\n",
      "[SA] new: 833. multiply(inverse(X0),multiply(multiply(X0,inverse(multiply(X1,inverse(X1)))),X2)) = X2 [forward demodulation 758,50]\n",
      "[SA] new: 834. multiply(inverse(X0),multiply(X0,multiply(inverse(multiply(X1,inverse(X1))),X2))) = X2 [forward demodulation 833,8]\n",
      "[SA] new: 835. multiply(X3,multiply(multiply(inverse(X3),inverse(multiply(X4,inverse(X4)))),X5)) = X5 [forward demodulation 759,50]\n",
      "[SA] new: 836. multiply(X3,multiply(inverse(X3),multiply(inverse(multiply(X4,inverse(X4))),X5))) = X5 [forward demodulation 835,8]\n",
      "[SA] new: 837. multiply(X14,multiply(multiply(inverse(multiply(X14,multiply(X15,inverse(X15)))),inverse(multiply(X16,inverse(X16)))),X17)) = X17 [forward demodulation 762,50]\n",
      "[SA] new: 838. multiply(X14,multiply(inverse(multiply(X14,multiply(X15,inverse(X15)))),multiply(inverse(multiply(X16,inverse(X16))),X17))) = X17 [forward demodulation 837,8]\n",
      "[SA] new: 839. multiply(X18,multiply(multiply(inverse(multiply(X18,inverse(multiply(X19,inverse(X19))))),inverse(multiply(X20,inverse(X20)))),X21)) = X21 [forward demodulation 763,50]\n",
      "[SA] new: 840. multiply(X18,multiply(inverse(multiply(X18,inverse(multiply(X19,inverse(X19))))),multiply(inverse(multiply(X20,inverse(X20))),X21))) = X21 [forward demodulation 839,8]\n",
      "[SA] new: 841. multiply(X22,multiply(multiply(inverse(multiply(X22,multiply(X23,inverse(multiply(X23,multiply(X24,inverse(X24))))))),inverse(multiply(X25,inverse(X25)))),X26)) = X26 [forward demodulation 764,50]\n",
      "[SA] new: 842. multiply(X22,multiply(inverse(multiply(X22,multiply(X23,inverse(multiply(X23,multiply(X24,inverse(X24))))))),multiply(inverse(multiply(X25,inverse(X25))),X26))) = X26 [forward demodulation 841,8]\n",
      "[SA] new: 843. multiply(X27,multiply(multiply(inverse(multiply(X27,multiply(X28,inverse(multiply(X28,inverse(multiply(X29,inverse(X29)))))))),inverse(multiply(X30,inverse(X30)))),X31)) = X31 [forward demodulation 765,50]\n",
      "[SA] new: 844. multiply(X27,multiply(inverse(multiply(X27,multiply(X28,inverse(multiply(X28,inverse(multiply(X29,inverse(X29)))))))),multiply(inverse(multiply(X30,inverse(X30))),X31))) = X31 [forward demodulation 843,8]\n",
      "[SA] new: 845. multiply(X32,multiply(multiply(inverse(multiply(X32,multiply(X33,multiply(X34,inverse(multiply(X33,X34)))))),inverse(multiply(X35,inverse(X35)))),X36)) = X36 [forward demodulation 766,50]\n",
      "[SA] new: 846. multiply(X32,multiply(inverse(multiply(X32,multiply(X33,multiply(X34,inverse(multiply(X33,X34)))))),multiply(inverse(multiply(X35,inverse(X35))),X36))) = X36 [forward demodulation 845,8]\n",
      "[SA] new: 847. multiply(inverse(multiply(X37,inverse(multiply(X37,inverse(multiply(X38,inverse(X38))))))),multiply(identity,X39)) = X39 [forward demodulation 767,10]\n",
      "[SA] new: 848. multiply(inverse(multiply(X37,inverse(multiply(X37,inverse(multiply(X38,inverse(X38))))))),X39) = X39 [forward demodulation 847,6]\n",
      "[SA] new: 849. multiply(multiply(X40,inverse(multiply(X40,inverse(multiply(X41,inverse(X41)))))),multiply(identity,X42)) = X42 [forward demodulation 768,10]\n",
      "[SA] new: 850. multiply(multiply(X40,inverse(multiply(X40,inverse(multiply(X41,inverse(X41)))))),X42) = X42 [forward demodulation 849,6]\n",
      "[SA] new: 851. multiply(X0,multiply(X1,inverse(multiply(X1,inverse(multiply(X2,inverse(X2))))))) = X0 [forward demodulation 771,1]\n",
      "[SA] new: 852. inverse(multiply(X3,multiply(X4,inverse(multiply(X4,inverse(multiply(X5,inverse(X5)))))))) = X3 [forward demodulation 772,1]\n",
      "[SA] new: 853. multiply(X13,X16) = multiply(multiply(X13,multiply(X14,inverse(multiply(X14,inverse(multiply(X15,inverse(X15))))))),X16) [forward demodulation 775,50]\n",
      "[SA] new: 854. multiply(X10,multiply(inverse(X10),X11)) = X11 [forward demodulation 785,552]\n",
      "[SA] new: 855. multiply(inverse(multiply(X12,inverse(multiply(X12,inverse(multiply(X13,inverse(X13))))))),X14) = X14 [forward demodulation 786,6]\n",
      "[SA] new: 856. multiply(X15,multiply(X16,multiply(inverse(multiply(X15,multiply(X16,multiply(X17,inverse(multiply(X17,inverse(multiply(X18,inverse(X18))))))))),X19))) = X19 [forward demodulation 787,8]\n",
      "[SA] new: 857. multiply(X3,multiply(inverse(multiply(X3,multiply(X4,inverse(multiply(X4,inverse(multiply(X5,inverse(X5)))))))),X7)) = X7 [forward demodulation 789,49]\n",
      "[SA] new: 858. multiply(X19,multiply(inverse(X19),X20)) = X20 [forward demodulation 792,552]\n",
      "[SA] new: 859. multiply(X29,multiply(inverse(multiply(X29,multiply(X26,inverse(multiply(X26,multiply(X27,inverse(multiply(X27,inverse(multiply(X28,inverse(X28))))))))))),X30)) = X30 [forward demodulation 794,50]\n",
      "[SA] new: 860. multiply(X31,multiply(inverse(multiply(X31,multiply(X32,inverse(multiply(X32,inverse(multiply(X33,inverse(X33)))))))),X34)) = X34 [forward demodulation 795,50]\n",
      "[SA] new: 861. multiply(X35,multiply(multiply(inverse(multiply(X35,multiply(X36,inverse(multiply(X36,inverse(multiply(X37,inverse(X37)))))))),multiply(X38,inverse(X38))),X39)) = X39 [forward demodulation 796,50]\n",
      "[SA] new: 862. multiply(X35,multiply(inverse(multiply(X35,multiply(X36,inverse(multiply(X36,inverse(multiply(X37,inverse(X37)))))))),multiply(multiply(X38,inverse(X38)),X39))) = X39 [forward demodulation 861,8]\n",
      "[SA] new: 863. multiply(X35,multiply(inverse(multiply(X35,multiply(X36,inverse(multiply(X36,inverse(multiply(X37,inverse(X37)))))))),multiply(X38,multiply(inverse(X38),X39)))) = X39 [forward demodulation 862,8]\n",
      "[SA] new: 864. multiply(X35,multiply(inverse(multiply(X35,multiply(X36,inverse(multiply(X36,inverse(multiply(X37,inverse(X37)))))))),X39)) = X39 [forward demodulation 863,49]\n",
      "[SA] new: 865. multiply(inverse(multiply(X45,inverse(multiply(X45,inverse(multiply(X46,inverse(X46))))))),X47) = X47 [forward demodulation 798,6]\n",
      "[SA] new: 866. multiply(X52,multiply(inverse(multiply(X52,multiply(X48,multiply(inverse(multiply(X48,multiply(X49,inverse(multiply(X49,inverse(multiply(X50,inverse(X50)))))))),multiply(X51,inverse(X51)))))),X53)) = X53 [forward demodulation 799,8]\n",
      "[SA] new: 867. multiply(X52,multiply(inverse(multiply(X52,multiply(X51,inverse(X51)))),X53)) = X53 [forward demodulation 866,552]\n",
      "[SA] new: 868. multiply(X58,multiply(X54,multiply(inverse(multiply(X54,multiply(X55,inverse(multiply(X55,inverse(multiply(X56,inverse(X56)))))))),multiply(X57,multiply(inverse(multiply(X58,X57)),X59))))) = X59 [forward demodulation 800,8]\n",
      "[SA] new: 869. multiply(X58,multiply(X57,multiply(inverse(multiply(X58,X57)),X59))) = X59 [forward demodulation 868,552]\n",
      "[SA] new: 870. multiply(X58,multiply(inverse(X58),X59)) = X59 [forward demodulation 869,98]\n",
      "[SA] new: 871. multiply(X72,X73) = multiply(X69,multiply(inverse(multiply(X69,multiply(X70,inverse(multiply(X70,inverse(multiply(X71,inverse(X71)))))))),multiply(X72,X73))) [forward demodulation 803,8]\n",
      "[SA] new: 872. multiply(X83,multiply(X84,multiply(inverse(multiply(X83,multiply(X84,multiply(X85,inverse(multiply(X85,inverse(multiply(X86,inverse(X86))))))))),X87))) = X87 [forward demodulation 806,8]\n",
      "[SA] new: 873. multiply(X1,multiply(inverse(multiply(X1,multiply(X2,inverse(multiply(X2,multiply(inverse(X0),multiply(X0,inverse(identity)))))))),X3)) = X3 [superposition 412,7]\n",
      "[SA] new: 874. multiply(X5,multiply(inverse(multiply(X5,multiply(X6,inverse(multiply(X6,multiply(X4,multiply(X4,inverse(identity)))))))),X7)) = X7 [superposition 412,3]\n",
      "[SA] new: 875. multiply(X9,multiply(inverse(multiply(X9,multiply(X10,inverse(multiply(X10,multiply(identity,multiply(X8,inverse(X8)))))))),X11)) = X11 [superposition 412,6]\n",
      "[SA] new: 876. multiply(X15,multiply(inverse(multiply(X15,multiply(X16,inverse(multiply(X16,multiply(multiply(X12,X13),multiply(X14,inverse(multiply(X12,multiply(X13,X14)))))))))),X17)) = X17 [superposition 412,8]\n",
      "[SA] new: 877. multiply(X19,multiply(inverse(multiply(X19,multiply(X20,inverse(multiply(X20,multiply(X18,multiply(identity,inverse(X18)))))))),X21)) = X21 [superposition 412,1]\n",
      "[SA] new: 878. multiply(X24,multiply(inverse(multiply(X24,multiply(X25,inverse(multiply(X25,multiply(inverse(X22),multiply(multiply(X22,X23),inverse(X23)))))))),X26)) = X26 [superposition 412,23]\n",
      "[SA] new: 879. multiply(X29,multiply(inverse(multiply(X29,multiply(X30,inverse(multiply(X30,multiply(X27,multiply(multiply(inverse(X27),X28),inverse(X28)))))))),X31)) = X31 [superposition 412,49]\n",
      "[SA] new: 880. multiply(X35,multiply(inverse(multiply(X35,multiply(X36,inverse(multiply(X36,multiply(X32,multiply(multiply(inverse(multiply(X33,X32)),X34),inverse(multiply(inverse(X33),X34))))))))),X37)) = X37 [superposition 412,98]\n",
      "[SA] new: 881. multiply(X41,multiply(inverse(multiply(X41,multiply(X42,inverse(multiply(X42,multiply(X38,multiply(multiply(inverse(multiply(inverse(X39),X38)),X40),inverse(multiply(X39,X40))))))))),X43)) = X43 [superposition 412,100]\n",
      "[SA] new: 882. multiply(X47,multiply(inverse(multiply(X47,multiply(X48,inverse(multiply(X48,multiply(X44,multiply(multiply(inverse(multiply(X44,multiply(X45,inverse(X45)))),X46),inverse(X46)))))))),X49)) = X49 [superposition 412,243]\n",
      "[SA] new: 883. multiply(X53,multiply(inverse(multiply(X53,multiply(X54,inverse(multiply(X54,multiply(X50,multiply(multiply(inverse(multiply(X50,inverse(multiply(X51,inverse(X51))))),X52),inverse(X52)))))))),X55)) = X55 [superposition 412,486]\n",
      "[SA] new: 884. multiply(X60,multiply(inverse(multiply(X60,multiply(X61,inverse(multiply(X61,multiply(X56,multiply(multiply(inverse(multiply(X56,multiply(X57,inverse(multiply(X57,multiply(X58,inverse(X58))))))),X59),inverse(X59)))))))),X62)) = X62 [superposition 412,366]\n",
      "[SA] new: 885. multiply(X67,multiply(inverse(multiply(X67,multiply(X68,inverse(multiply(X68,multiply(X63,multiply(multiply(inverse(multiply(X63,multiply(X64,inverse(multiply(X64,inverse(multiply(X65,inverse(X65)))))))),X66),inverse(X66)))))))),X69)) = X69 [superposition 412,552]\n",
      "[SA] new: 886. multiply(X74,multiply(inverse(multiply(X74,multiply(X75,inverse(multiply(X75,multiply(X70,multiply(multiply(inverse(multiply(X70,multiply(X71,multiply(X72,inverse(multiply(X71,X72)))))),X73),inverse(X73)))))))),X76)) = X76 [superposition 412,340]\n",
      "[SA] new: 887. multiply(X82,multiply(inverse(multiply(X82,multiply(X83,inverse(multiply(X83,multiply(X77,multiply(multiply(inverse(multiply(X77,multiply(X78,inverse(multiply(X78,multiply(X79,multiply(X80,inverse(multiply(X79,X80))))))))),X81),inverse(X81)))))))),X84)) = X84 [superposition 412,412]\n",
      "[SA] new: 888. multiply(X89,multiply(inverse(multiply(X89,multiply(X90,inverse(multiply(X90,multiply(X85,multiply(multiply(X86,multiply(X87,multiply(inverse(multiply(X85,multiply(X86,X87))),X88))),inverse(X88)))))))),X91)) = X91 [superposition 412,121]\n",
      "[SA] new: 889. multiply(X92,multiply(inverse(multiply(X92,multiply(X93,inverse(multiply(X93,multiply(a,multiply(b,inverse(c)))))))),X94)) = X94 [superposition 412,4]\n",
      "[SA] new: 890. multiply(X1,multiply(inverse(multiply(X1,multiply(X2,inverse(multiply(X2,multiply(X0,inverse(multiply(X0,identity)))))))),X3)) = X3 [superposition 412,6]\n",
      "[SA] new: 891. multiply(X7,multiply(inverse(multiply(X7,multiply(X8,inverse(multiply(X8,multiply(X6,multiply(X4,multiply(X5,inverse(multiply(X6,multiply(X4,X5))))))))))),X9)) = X9 [superposition 412,8]\n",
      "[SA] new: 892. multiply(X1,multiply(inverse(multiply(X1,multiply(X2,inverse(multiply(X2,inverse(multiply(inverse(X0),X0))))))),X3)) = X3 [superposition 412,23]\n",
      "[SA] new: 893. multiply(X5,multiply(inverse(multiply(X5,multiply(X6,inverse(multiply(X6,inverse(multiply(X4,inverse(X4)))))))),X7)) = X7 [superposition 412,49]\n",
      "[SA] new: 894. multiply(X10,multiply(inverse(multiply(X10,multiply(X11,inverse(multiply(X11,multiply(inverse(X9),inverse(multiply(X8,inverse(multiply(X9,X8)))))))))),X12)) = X12 [superposition 412,98]\n",
      "[SA] new: 895. multiply(X15,multiply(inverse(multiply(X15,multiply(X16,inverse(multiply(X16,multiply(X14,inverse(multiply(X13,inverse(multiply(inverse(X14),X13)))))))))),X17)) = X17 [superposition 412,100]\n",
      "[SA] new: 896. multiply(X20,multiply(inverse(multiply(X20,multiply(X21,inverse(multiply(X21,inverse(multiply(X18,inverse(multiply(X18,multiply(X19,inverse(X19))))))))))),X22)) = X22 [superposition 412,243]\n",
      "[SA] new: 897. multiply(X25,multiply(inverse(multiply(X25,multiply(X26,inverse(multiply(X26,inverse(multiply(X23,inverse(multiply(X23,inverse(multiply(X24,inverse(X24)))))))))))),X27)) = X27 [superposition 412,486]\n",
      "[SA] new: 898. multiply(X31,multiply(inverse(multiply(X31,multiply(X32,inverse(multiply(X32,inverse(multiply(X28,inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30)))))))))))))),X33)) = X33 [superposition 412,366]\n",
      "[SA] new: 899. multiply(X37,multiply(inverse(multiply(X37,multiply(X38,inverse(multiply(X38,inverse(multiply(X34,inverse(multiply(X34,multiply(X35,inverse(multiply(X35,inverse(multiply(X36,inverse(X36))))))))))))))),X39)) = X39 [superposition 412,552]\n",
      "[SA] new: 900. multiply(X43,multiply(inverse(multiply(X43,multiply(X44,inverse(multiply(X44,inverse(multiply(X40,inverse(multiply(X40,multiply(X41,multiply(X42,inverse(multiply(X41,X42))))))))))))),X45)) = X45 [superposition 412,340]\n",
      "[SA] new: 901. multiply(X50,multiply(inverse(multiply(X50,multiply(X51,inverse(multiply(X51,inverse(multiply(X46,inverse(multiply(X46,multiply(X47,inverse(multiply(X47,multiply(X48,multiply(X49,inverse(multiply(X48,X49)))))))))))))))),X52)) = X52 [superposition 412,412]\n",
      "[SA] new: 902. multiply(X54,multiply(inverse(multiply(X54,multiply(X55,inverse(multiply(X55,multiply(X53,inverse(multiply(identity,X53)))))))),X56)) = X56 [superposition 412,6]\n",
      "[SA] new: 903. multiply(X60,multiply(inverse(multiply(X60,multiply(X61,inverse(multiply(X61,multiply(X57,multiply(X58,multiply(X59,inverse(multiply(multiply(X57,X58),X59)))))))))),X62)) = X62 [superposition 412,8]\n",
      "[SA] new: 904. multiply(X2,multiply(inverse(multiply(X2,multiply(inverse(X0),inverse(multiply(X1,inverse(multiply(X0,X1))))))),X3)) = X3 [superposition 412,23]\n",
      "[SA] new: 905. multiply(X6,multiply(inverse(multiply(X6,multiply(X4,inverse(multiply(X5,inverse(multiply(inverse(X4),X5))))))),X7)) = X7 [superposition 412,49]\n",
      "[SA] new: 906. multiply(X11,multiply(inverse(multiply(X11,multiply(X8,inverse(multiply(inverse(X9),multiply(X10,inverse(multiply(inverse(multiply(X9,X8)),X10)))))))),X12)) = X12 [superposition 412,98]\n",
      "[SA] new: 907. multiply(X16,multiply(inverse(multiply(X16,multiply(X13,inverse(multiply(X14,multiply(X15,inverse(multiply(inverse(multiply(inverse(X14),X13)),X15)))))))),X17)) = X17 [superposition 412,100]\n",
      "[SA] new: 908. multiply(X21,multiply(inverse(multiply(X21,multiply(X18,inverse(multiply(X20,inverse(multiply(inverse(multiply(X18,multiply(X19,inverse(X19)))),X20))))))),X22)) = X22 [superposition 412,243]\n",
      "[SA] new: 909. multiply(X26,multiply(inverse(multiply(X26,multiply(X23,inverse(multiply(X25,inverse(multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),X25))))))),X27)) = X27 [superposition 412,486]\n",
      "[SA] new: 910. multiply(X32,multiply(inverse(multiply(X32,multiply(X28,inverse(multiply(X31,inverse(multiply(inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30))))))),X31))))))),X33)) = X33 [superposition 412,366]\n",
      "[SA] new: 911. multiply(X38,multiply(inverse(multiply(X38,multiply(X34,inverse(multiply(X37,inverse(multiply(inverse(multiply(X34,multiply(X35,inverse(multiply(X35,inverse(multiply(X36,inverse(X36)))))))),X37))))))),X39)) = X39 [superposition 412,552]\n",
      "[SA] new: 912. multiply(X44,multiply(inverse(multiply(X44,multiply(X40,inverse(multiply(X43,inverse(multiply(inverse(multiply(X40,multiply(X41,multiply(X42,inverse(multiply(X41,X42)))))),X43))))))),X45)) = X45 [superposition 412,340]\n",
      "[SA] new: 913. multiply(X51,multiply(inverse(multiply(X51,multiply(X46,inverse(multiply(X50,inverse(multiply(inverse(multiply(X46,multiply(X47,inverse(multiply(X47,multiply(X48,multiply(X49,inverse(multiply(X48,X49))))))))),X50))))))),X52)) = X52 [superposition 412,412]\n",
      "[SA] new: 914. multiply(X55,multiply(inverse(multiply(X55,multiply(inverse(multiply(X53,multiply(X54,inverse(multiply(X53,X54))))),inverse(identity)))),X56)) = X56 [superposition 412,7]\n",
      "[SA] new: 915. multiply(X59,multiply(inverse(multiply(X59,multiply(multiply(X57,multiply(X58,inverse(multiply(X57,X58)))),inverse(identity)))),X60)) = X60 [superposition 412,3]\n",
      "[SA] new: 916. multiply(X63,multiply(inverse(multiply(X63,multiply(identity,inverse(multiply(X61,multiply(X62,inverse(multiply(X61,X62)))))))),X64)) = X64 [superposition 412,6]\n",
      "[SA] new: 917. multiply(X69,multiply(inverse(multiply(X69,multiply(multiply(X65,X66),inverse(multiply(X65,multiply(X66,multiply(X67,multiply(X68,inverse(multiply(X67,X68)))))))))),X70)) = X70 [superposition 412,8]\n",
      "[SA] new: 918. multiply(X2,multiply(inverse(multiply(X2,inverse(multiply(identity,multiply(X0,multiply(X1,inverse(multiply(X0,X1)))))))),X3)) = X3 [superposition 412,6]\n",
      "[SA] new: 919. multiply(X8,multiply(inverse(multiply(X8,multiply(X4,multiply(X5,inverse(multiply(multiply(X4,X5),multiply(X6,multiply(X7,inverse(multiply(X6,X7)))))))))),X9)) = X9 [superposition 412,8]\n",
      "[SA] new: 920. multiply(inverse(X0),multiply(inverse(inverse(multiply(X0,multiply(X1,multiply(X2,inverse(multiply(X1,X2))))))),X3)) = X3 [superposition 412,23]\n",
      "[SA] new: 921. multiply(X4,multiply(inverse(inverse(multiply(inverse(X4),multiply(X5,multiply(X6,inverse(multiply(X5,X6))))))),X7)) = X7 [superposition 412,49]\n",
      "[SA] new: 922. multiply(X8,multiply(inverse(multiply(inverse(X9),inverse(multiply(inverse(multiply(X9,X8)),multiply(X10,multiply(X11,inverse(multiply(X10,X11)))))))),X12)) = X12 [superposition 412,98]\n",
      "[SA] new: 923. multiply(X13,multiply(inverse(multiply(X14,inverse(multiply(inverse(multiply(inverse(X14),X13)),multiply(X15,multiply(X16,inverse(multiply(X15,X16)))))))),X17)) = X17 [superposition 412,100]\n",
      "[SA] new: 924. multiply(X18,multiply(inverse(inverse(multiply(inverse(multiply(X18,multiply(X19,inverse(X19)))),multiply(X20,multiply(X21,inverse(multiply(X20,X21))))))),X22)) = X22 [superposition 412,243]\n",
      "[SA] new: 925. multiply(X23,multiply(inverse(inverse(multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),multiply(X25,multiply(X26,inverse(multiply(X25,X26))))))),X27)) = X27 [superposition 412,486]\n",
      "[SA] new: 926. multiply(X28,multiply(inverse(inverse(multiply(inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30))))))),multiply(X31,multiply(X32,inverse(multiply(X31,X32))))))),X33)) = X33 [superposition 412,366]\n",
      "[SA] new: 927. multiply(X34,multiply(inverse(inverse(multiply(inverse(multiply(X34,multiply(X35,inverse(multiply(X35,inverse(multiply(X36,inverse(X36)))))))),multiply(X37,multiply(X38,inverse(multiply(X37,X38))))))),X39)) = X39 [superposition 412,552]\n",
      "[SA] new: 928. multiply(X40,multiply(inverse(inverse(multiply(inverse(multiply(X40,multiply(X41,multiply(X42,inverse(multiply(X41,X42)))))),multiply(X43,multiply(X44,inverse(multiply(X43,X44))))))),X45)) = X45 [superposition 412,340]\n",
      "[SA] new: 929. multiply(X46,multiply(inverse(inverse(multiply(inverse(multiply(X46,multiply(X47,inverse(multiply(X47,multiply(X48,multiply(X49,inverse(multiply(X48,X49))))))))),multiply(X50,multiply(X51,inverse(multiply(X50,X51))))))),X52)) = X52 [superposition 412,412]\n",
      "[SA] new: 930. multiply(inverse(multiply(X53,inverse(multiply(X53,multiply(X54,multiply(X55,inverse(multiply(X54,X55)))))))),multiply(inverse(identity),X56)) = X56 [superposition 412,7]\n",
      "[SA] new: 931. multiply(multiply(X57,inverse(multiply(X57,multiply(X58,multiply(X59,inverse(multiply(X58,X59))))))),multiply(inverse(identity),X60)) = X60 [superposition 412,3]\n",
      "[SA] new: 932. multiply(identity,multiply(inverse(multiply(X61,inverse(multiply(X61,multiply(X62,multiply(X63,inverse(multiply(X62,X63)))))))),X64)) = X64 [superposition 412,6]\n",
      "[SA] new: 933. multiply(multiply(X65,X66),multiply(inverse(multiply(X65,multiply(X66,multiply(X67,inverse(multiply(X67,multiply(X68,multiply(X69,inverse(multiply(X68,X69)))))))))),X70)) = X70 [superposition 412,8]\n",
      "[SA] new: 934. multiply(X0,identity) = multiply(X0,multiply(X1,inverse(multiply(X1,multiply(X2,multiply(X3,inverse(multiply(X2,X3)))))))) [superposition 412,7]\n",
      "[SA] new: 935. multiply(X4,identity) = inverse(multiply(X4,multiply(X5,inverse(multiply(X5,multiply(X6,multiply(X7,inverse(multiply(X6,X7))))))))) [superposition 412,3]\n",
      "[SA] new: 936. identity = multiply(X8,inverse(multiply(X8,multiply(X9,inverse(multiply(X9,multiply(X10,multiply(X11,inverse(multiply(X10,X11)))))))))) [superposition 412,1]\n",
      "[SA] new: 937. multiply(multiply(X12,multiply(X13,inverse(multiply(X13,multiply(X14,multiply(X15,inverse(multiply(X14,X15)))))))),X16) = multiply(X12,X16) [superposition 412,23]\n",
      "[SA] new: 938. multiply(inverse(inverse(multiply(X17,multiply(X18,inverse(multiply(X18,multiply(X19,multiply(X20,inverse(multiply(X19,X20)))))))))),X21) = multiply(X17,X21) [superposition 412,49]\n",
      "[SA] new: 939. multiply(inverse(multiply(X26,inverse(multiply(X22,multiply(X23,inverse(multiply(X23,multiply(X24,multiply(X25,inverse(multiply(X24,X25))))))))))),X27) = multiply(X22,multiply(inverse(X26),X27)) [superposition 412,98]\n",
      "[SA] new: 940. multiply(inverse(multiply(inverse(X32),inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,multiply(X31,inverse(multiply(X30,X31))))))))))),X33) = multiply(X28,multiply(X32,X33)) [superposition 412,100]\n",
      "[SA] new: 941. multiply(inverse(multiply(inverse(multiply(X34,multiply(X35,inverse(multiply(X35,multiply(X36,multiply(X37,inverse(multiply(X36,X37))))))))),multiply(X38,inverse(X38)))),X39) = multiply(X34,X39) [superposition 412,243]\n",
      "[SA] new: 942. multiply(inverse(multiply(inverse(multiply(X40,multiply(X41,inverse(multiply(X41,multiply(X42,multiply(X43,inverse(multiply(X42,X43))))))))),inverse(multiply(X44,inverse(X44))))),X45) = multiply(X40,X45) [superposition 412,486]\n",
      "[SA] new: 943. multiply(inverse(multiply(inverse(multiply(X46,multiply(X47,inverse(multiply(X47,multiply(X48,multiply(X49,inverse(multiply(X48,X49))))))))),multiply(X50,inverse(multiply(X50,multiply(X51,inverse(X51))))))),X52) = multiply(X46,X52) [superposition 412,366]\n",
      "[SA] new: 944. multiply(inverse(multiply(inverse(multiply(X53,multiply(X54,inverse(multiply(X54,multiply(X55,multiply(X56,inverse(multiply(X55,X56))))))))),multiply(X57,inverse(multiply(X57,inverse(multiply(X58,inverse(X58)))))))),X59) = multiply(X53,X59) [superposition 412,552]\n",
      "[SA] new: 945. multiply(inverse(multiply(inverse(multiply(X60,multiply(X61,inverse(multiply(X61,multiply(X62,multiply(X63,inverse(multiply(X62,X63))))))))),multiply(X64,multiply(X65,inverse(multiply(X64,X65)))))),X66) = multiply(X60,X66) [superposition 412,340]\n",
      "[SA] new: 946. multiply(inverse(multiply(inverse(multiply(X67,multiply(X68,inverse(multiply(X68,multiply(X69,multiply(X70,inverse(multiply(X69,X70))))))))),multiply(X71,inverse(multiply(X71,multiply(X72,multiply(X73,inverse(multiply(X72,X73))))))))),X74) = multiply(X67,X74) [superposition 412,412]\n",
      "[SA] new: 947. multiply(X79,multiply(X80,multiply(inverse(multiply(inverse(multiply(X75,multiply(X76,inverse(multiply(X76,multiply(X77,multiply(X78,inverse(multiply(X77,X78))))))))),multiply(X79,X80))),X81))) = multiply(X75,X81) [superposition 412,121]\n",
      "[SA] new: 948. multiply(inverse(multiply(X0,inverse(multiply(X0,multiply(X1,multiply(X2,inverse(multiply(X1,X2)))))))),X3) = X3 [superposition 412,98]\n",
      "[SA] new: 949. multiply(X13,multiply(inverse(multiply(X9,multiply(inverse(multiply(X9,multiply(X10,inverse(multiply(X10,multiply(X11,multiply(X12,inverse(multiply(X11,X12))))))))),X13))),X14)) = X14 [superposition 412,121]\n",
      "[SA] new: 950. multiply(inverse(multiply(identity,multiply(X15,inverse(multiply(X15,multiply(X16,multiply(X17,inverse(multiply(X16,X17))))))))),X18) = X18 [superposition 412,6]\n",
      "[SA] new: 951. multiply(X19,multiply(X20,multiply(inverse(multiply(multiply(X19,X20),multiply(X21,inverse(multiply(X21,multiply(X22,multiply(X23,inverse(multiply(X22,X23))))))))),X24))) = X24 [superposition 412,8]\n",
      "[SA] new: 952. multiply(inverse(multiply(X0,inverse(multiply(X0,multiply(X1,multiply(X2,inverse(multiply(X1,X2)))))))),X3) = X3 [superposition 98,412]\n",
      "[SA] new: 953. multiply(X4,multiply(inverse(multiply(X4,multiply(X5,inverse(multiply(X5,multiply(X6,multiply(X7,inverse(multiply(X6,X7))))))))),multiply(X8,multiply(inverse(X8),X9)))) = X9 [superposition 121,412]\n",
      "[SA] new: 954. multiply(inverse(multiply(X10,multiply(X11,multiply(X12,inverse(multiply(X11,X12)))))),multiply(X10,X13)) = X13 [superposition 121,412]\n",
      "[SA] new: 955. multiply(X18,multiply(inverse(multiply(X18,multiply(X14,inverse(multiply(X14,multiply(X15,inverse(multiply(X15,multiply(X16,multiply(X17,inverse(multiply(X16,X17)))))))))))),X19)) = X19 [superposition 121,412]\n",
      "[SA] new: 956. multiply(X24,multiply(inverse(multiply(X20,multiply(inverse(multiply(X20,multiply(X21,inverse(multiply(X21,multiply(X22,multiply(X23,inverse(multiply(X22,X23))))))))),X24))),X25)) = X25 [superposition 121,412]\n",
      "[SA] new: 957. multiply(X26,multiply(inverse(multiply(X30,inverse(multiply(inverse(multiply(X26,multiply(X27,inverse(multiply(X27,multiply(X28,multiply(X29,inverse(multiply(X28,X29))))))))),X30)))),X31)) = X31 [superposition 340,412]\n",
      "[SA] new: 958. multiply(X36,multiply(inverse(multiply(X36,multiply(X32,inverse(inverse(inverse(multiply(X32,multiply(X33,inverse(multiply(X33,multiply(X34,multiply(X35,inverse(multiply(X34,X35)))))))))))))),X37)) = X37 [superposition 366,412]\n",
      "[SA] new: 959. multiply(X38,multiply(inverse(inverse(inverse(multiply(X38,multiply(X39,inverse(multiply(X39,multiply(X40,multiply(X41,inverse(multiply(X40,X41))))))))))),X42)) = X42 [superposition 243,412]\n",
      "[SA] new: 960. multiply(X43,multiply(inverse(inverse(multiply(inverse(multiply(X43,multiply(X44,inverse(multiply(X44,multiply(X45,multiply(X46,inverse(multiply(X45,X46))))))))),multiply(X47,inverse(X47))))),X48)) = X48 [superposition 366,412]\n",
      "[SA] new: 961. multiply(X49,multiply(inverse(inverse(multiply(inverse(multiply(X49,multiply(X50,inverse(multiply(X50,multiply(X51,multiply(X52,inverse(multiply(X51,X52))))))))),inverse(multiply(X53,inverse(X53)))))),X54)) = X54 [superposition 552,412]\n",
      "[SA] new: 962. multiply(X59,multiply(inverse(multiply(X59,inverse(multiply(X55,inverse(multiply(X55,multiply(X56,inverse(multiply(X56,multiply(X57,multiply(X58,inverse(multiply(X57,X58))))))))))))),X60)) = X60 [superposition 340,412]\n",
      "[SA] new: 963. multiply(inverse(multiply(identity,multiply(X61,inverse(multiply(X61,multiply(X62,multiply(X63,inverse(multiply(X62,X63))))))))),X64) = X64 [superposition 6,412]\n",
      "[SA] new: 964. multiply(X70,multiply(inverse(multiply(X70,multiply(X65,multiply(multiply(inverse(multiply(X65,multiply(X66,inverse(multiply(X66,multiply(X67,multiply(X68,inverse(multiply(X67,X68))))))))),X69),inverse(X69))))),X71)) = X71 [superposition 340,412]\n",
      "[SA] new: 965. multiply(X77,multiply(X72,multiply(multiply(inverse(multiply(X72,multiply(X73,inverse(multiply(X73,multiply(X74,multiply(X75,inverse(multiply(X74,X75))))))))),X76),multiply(inverse(multiply(X77,X76)),X78)))) = X78 [superposition 121,412]\n",
      "[SA] new: 966. multiply(inverse(X79),X84) = multiply(multiply(inverse(multiply(X79,multiply(X80,inverse(multiply(X80,multiply(X81,multiply(X82,inverse(multiply(X81,X82))))))))),X83),multiply(inverse(X83),X84)) [superposition 98,412]\n",
      "[SA] new: 967. multiply(inverse(multiply(X85,multiply(X86,inverse(multiply(X86,multiply(X87,multiply(X88,inverse(multiply(X87,X88))))))))),X89) = multiply(inverse(X85),X89) [superposition 23,412]\n",
      "[SA] new: 968. multiply(X90,multiply(multiply(inverse(multiply(X90,multiply(X91,inverse(multiply(X91,multiply(X92,multiply(X93,inverse(multiply(X92,X93))))))))),X94),X95)) = multiply(X94,X95) [superposition 8,412]\n",
      "[SA] new: 969. multiply(X96,X101) = multiply(multiply(inverse(multiply(inverse(X96),multiply(X97,inverse(multiply(X97,multiply(X98,multiply(X99,inverse(multiply(X98,X99))))))))),X100),multiply(inverse(X100),X101)) [superposition 100,412]\n",
      "[SA] new: 970. multiply(inverse(multiply(inverse(X102),multiply(X103,inverse(multiply(X103,multiply(X104,multiply(X105,inverse(multiply(X104,X105))))))))),X106) = multiply(X102,X106) [superposition 49,412]\n",
      "[SA] new: 971. multiply(X107,multiply(X108,multiply(inverse(multiply(multiply(X107,X108),multiply(X109,inverse(multiply(X109,multiply(X110,multiply(X111,inverse(multiply(X110,X111))))))))),X112))) = X112 [superposition 8,412]\n",
      "[SA] new: 972. multiply(inverse(X113),multiply(inverse(multiply(inverse(multiply(X113,X114)),multiply(X115,inverse(multiply(X115,multiply(X116,multiply(X117,inverse(multiply(X116,X117))))))))),X118)) = multiply(X114,X118) [superposition 98,412]\n",
      "[SA] new: 973. multiply(X119,multiply(inverse(multiply(inverse(multiply(inverse(X119),X120)),multiply(X121,inverse(multiply(X121,multiply(X122,multiply(X123,inverse(multiply(X122,X123))))))))),X124)) = multiply(X120,X124) [superposition 100,412]\n",
      "[SA] new: 974. multiply(inverse(multiply(inverse(multiply(X125,multiply(X126,inverse(X126)))),multiply(X127,inverse(multiply(X127,multiply(X128,multiply(X129,inverse(multiply(X128,X129))))))))),X130) = multiply(X125,X130) [superposition 243,412]\n",
      "[SA] new: 975. multiply(inverse(multiply(inverse(multiply(X131,inverse(multiply(X132,inverse(X132))))),multiply(X133,inverse(multiply(X133,multiply(X134,multiply(X135,inverse(multiply(X134,X135))))))))),X136) = multiply(X131,X136) [superposition 486,412]\n",
      "[SA] new: 976. multiply(inverse(multiply(inverse(multiply(X137,multiply(X138,X139))),multiply(X140,inverse(multiply(X140,multiply(X141,multiply(X142,inverse(multiply(X141,X142))))))))),X143) = multiply(X137,multiply(X138,multiply(X139,X143))) [superposition 121,412]\n",
      "[SA] new: 977. multiply(inverse(multiply(inverse(multiply(X144,multiply(X145,inverse(multiply(X145,multiply(X146,inverse(X146))))))),multiply(X147,inverse(multiply(X147,multiply(X148,multiply(X149,inverse(multiply(X148,X149))))))))),X150) = multiply(X144,X150) [superposition 366,412]\n",
      "[SA] new: 978. multiply(inverse(multiply(inverse(multiply(X151,multiply(X152,inverse(multiply(X152,inverse(multiply(X153,inverse(X153)))))))),multiply(X154,inverse(multiply(X154,multiply(X155,multiply(X156,inverse(multiply(X155,X156))))))))),X157) = multiply(X151,X157) [superposition 552,412]\n",
      "[SA] new: 979. multiply(inverse(multiply(inverse(multiply(X158,multiply(X159,multiply(X160,inverse(multiply(X159,X160)))))),multiply(X161,inverse(multiply(X161,multiply(X162,multiply(X163,inverse(multiply(X162,X163))))))))),X164) = multiply(X158,X164) [superposition 340,412]\n",
      "[SA] new: 980. multiply(X1,multiply(inverse(multiply(X1,multiply(X2,inverse(multiply(X2,inverse(identity)))))),X3)) = X3 [forward demodulation 873,23]\n",
      "[SA] new: 981. multiply(X1,multiply(inverse(multiply(X1,multiply(X2,inverse(multiply(X2,identity))))),X3)) = X3 [forward demodulation 980,10]\n",
      "[SA] new: 982. multiply(X1,multiply(inverse(multiply(X1,multiply(X2,inverse(X2)))),X3)) = X3 [forward demodulation 981,1]\n",
      "[SA] new: 983. multiply(X5,multiply(inverse(multiply(X5,multiply(X6,inverse(multiply(X6,multiply(X4,multiply(X4,identity))))))),X7)) = X7 [forward demodulation 874,10]\n",
      "[SA] new: 984. multiply(X5,multiply(inverse(multiply(X5,multiply(X6,inverse(multiply(X6,multiply(X4,X4)))))),X7)) = X7 [forward demodulation 983,1]\n",
      "[SA] new: 985. multiply(X5,multiply(inverse(multiply(X5,multiply(X6,inverse(multiply(X6,identity))))),X7)) = X7 [forward demodulation 984,3]\n",
      "[SA] new: 986. multiply(X5,multiply(inverse(multiply(X5,multiply(X6,inverse(X6)))),X7)) = X7 [forward demodulation 985,1]\n",
      "[SA] new: 987. multiply(X9,multiply(inverse(multiply(X9,multiply(X10,inverse(multiply(X10,multiply(X8,inverse(X8))))))),X11)) = X11 [forward demodulation 875,6]\n",
      "[SA] new: 988. multiply(X15,multiply(inverse(multiply(X15,multiply(X16,inverse(multiply(X16,multiply(X12,multiply(X13,multiply(X14,inverse(multiply(X12,multiply(X13,X14))))))))))),X17)) = X17 [forward demodulation 876,8]\n",
      "[SA] new: 989. multiply(X19,multiply(inverse(multiply(X19,multiply(X20,inverse(multiply(X20,multiply(X18,inverse(X18))))))),X21)) = X21 [forward demodulation 877,6]\n",
      "[SA] new: 990. multiply(X24,multiply(inverse(multiply(X24,multiply(X25,inverse(multiply(X25,multiply(inverse(X22),multiply(X22,multiply(X23,inverse(X23))))))))),X26)) = X26 [forward demodulation 878,8]\n",
      "[SA] new: 991. multiply(X24,multiply(inverse(multiply(X24,multiply(X25,inverse(multiply(X25,multiply(X23,inverse(X23))))))),X26)) = X26 [forward demodulation 990,23]\n",
      "[SA] new: 992. multiply(X29,multiply(inverse(multiply(X29,multiply(X30,inverse(multiply(X30,multiply(X27,multiply(inverse(X27),multiply(X28,inverse(X28))))))))),X31)) = X31 [forward demodulation 879,8]\n",
      "[SA] new: 993. multiply(X29,multiply(inverse(multiply(X29,multiply(X30,inverse(multiply(X30,multiply(X28,inverse(X28))))))),X31)) = X31 [forward demodulation 992,49]\n",
      "[SA] new: 994. multiply(X35,multiply(inverse(multiply(X35,multiply(X36,inverse(multiply(X36,multiply(X32,multiply(inverse(multiply(X33,X32)),multiply(X34,inverse(multiply(inverse(X33),X34)))))))))),X37)) = X37 [forward demodulation 880,8]\n",
      "[SA] new: 995. multiply(X35,multiply(inverse(multiply(X35,multiply(X36,inverse(multiply(X36,multiply(inverse(X33),multiply(X34,inverse(multiply(inverse(X33),X34))))))))),X37)) = X37 [forward demodulation 994,98]\n",
      "[SA] new: 996. multiply(X41,multiply(inverse(multiply(X41,multiply(X42,inverse(multiply(X42,multiply(X38,multiply(inverse(multiply(inverse(X39),X38)),multiply(X40,inverse(multiply(X39,X40)))))))))),X43)) = X43 [forward demodulation 881,8]\n",
      "[SA] new: 997. multiply(X41,multiply(inverse(multiply(X41,multiply(X42,inverse(multiply(X42,multiply(inverse(inverse(X39)),multiply(X40,inverse(multiply(X39,X40))))))))),X43)) = X43 [forward demodulation 996,98]\n",
      "[SA] new: 998. multiply(X41,multiply(inverse(multiply(X41,multiply(X42,inverse(multiply(X42,multiply(X39,multiply(X40,inverse(multiply(X39,X40))))))))),X43)) = X43 [forward demodulation 997,50]\n",
      "[SA] new: 999. multiply(X47,multiply(inverse(multiply(X47,multiply(X48,inverse(multiply(X48,multiply(X44,multiply(inverse(multiply(X44,multiply(X45,inverse(X45)))),multiply(X46,inverse(X46))))))))),X49)) = X49 [forward demodulation 882,8]\n",
      "[SA] new: 1000. multiply(X47,multiply(inverse(multiply(X47,multiply(X48,inverse(multiply(X48,multiply(X46,inverse(X46))))))),X49)) = X49 [forward demodulation 999,243]\n",
      "[SA] new: 1001. multiply(X53,multiply(inverse(multiply(X53,multiply(X54,inverse(multiply(X54,multiply(X50,multiply(inverse(multiply(X50,inverse(multiply(X51,inverse(X51))))),multiply(X52,inverse(X52))))))))),X55)) = X55 [forward demodulation 883,8]\n",
      "[SA] new: 1002. multiply(X53,multiply(inverse(multiply(X53,multiply(X54,inverse(multiply(X54,multiply(X52,inverse(X52))))))),X55)) = X55 [forward demodulation 1001,486]\n",
      "[SA] new: 1003. multiply(X60,multiply(inverse(multiply(X60,multiply(X61,inverse(multiply(X61,multiply(X56,multiply(inverse(multiply(X56,multiply(X57,inverse(multiply(X57,multiply(X58,inverse(X58))))))),multiply(X59,inverse(X59))))))))),X62)) = X62 [forward demodulation 884,8]\n",
      "[SA] new: 1004. multiply(X60,multiply(inverse(multiply(X60,multiply(X61,inverse(multiply(X61,multiply(X59,inverse(X59))))))),X62)) = X62 [forward demodulation 1003,366]\n",
      "[SA] new: 1005. multiply(X67,multiply(inverse(multiply(X67,multiply(X68,inverse(multiply(X68,multiply(X63,multiply(inverse(multiply(X63,multiply(X64,inverse(multiply(X64,inverse(multiply(X65,inverse(X65)))))))),multiply(X66,inverse(X66))))))))),X69)) = X69 [forward demodulation 885,8]\n",
      "[SA] new: 1006. multiply(X67,multiply(inverse(multiply(X67,multiply(X68,inverse(multiply(X68,multiply(X66,inverse(X66))))))),X69)) = X69 [forward demodulation 1005,552]\n",
      "[SA] new: 1007. multiply(X74,multiply(inverse(multiply(X74,multiply(X75,inverse(multiply(X75,multiply(X70,multiply(inverse(multiply(X70,multiply(X71,multiply(X72,inverse(multiply(X71,X72)))))),multiply(X73,inverse(X73))))))))),X76)) = X76 [forward demodulation 886,8]\n",
      "[SA] new: 1008. multiply(X74,multiply(inverse(multiply(X74,multiply(X75,inverse(multiply(X75,multiply(X73,inverse(X73))))))),X76)) = X76 [forward demodulation 1007,340]\n",
      "[SA] new: 1009. multiply(X82,multiply(inverse(multiply(X82,multiply(X83,inverse(multiply(X83,multiply(X77,multiply(inverse(multiply(X77,multiply(X78,inverse(multiply(X78,multiply(X79,multiply(X80,inverse(multiply(X79,X80))))))))),multiply(X81,inverse(X81))))))))),X84)) = X84 [forward demodulation 887,8]\n",
      "[SA] new: 1010. multiply(X82,multiply(inverse(multiply(X82,multiply(X83,inverse(multiply(X83,multiply(X81,inverse(X81))))))),X84)) = X84 [forward demodulation 1009,412]\n",
      "[SA] new: 1011. multiply(X89,multiply(inverse(multiply(X89,multiply(X90,inverse(multiply(X90,multiply(X85,multiply(X86,multiply(multiply(X87,multiply(inverse(multiply(X85,multiply(X86,X87))),X88)),inverse(X88))))))))),X91)) = X91 [forward demodulation 888,8]\n",
      "[SA] new: 1012. multiply(X89,multiply(inverse(multiply(X89,multiply(X90,inverse(multiply(X90,multiply(X85,multiply(X86,multiply(X87,multiply(multiply(inverse(multiply(X85,multiply(X86,X87))),X88),inverse(X88)))))))))),X91)) = X91 [forward demodulation 1011,8]\n",
      "[SA] new: 1013. multiply(X89,multiply(inverse(multiply(X89,multiply(X90,inverse(multiply(X90,multiply(X85,multiply(X86,multiply(X87,multiply(inverse(multiply(X85,multiply(X86,X87))),multiply(X88,inverse(X88))))))))))),X91)) = X91 [forward demodulation 1012,8]\n",
      "[SA] new: 1014. multiply(X89,multiply(inverse(multiply(X89,multiply(X90,inverse(multiply(X90,multiply(X88,inverse(X88))))))),X91)) = X91 [forward demodulation 1013,121]\n",
      "[SA] new: 1015. multiply(X1,multiply(inverse(multiply(X1,multiply(X2,inverse(multiply(X2,multiply(X0,inverse(X0))))))),X3)) = X3 [forward demodulation 890,1]\n",
      "[SA] new: 1016. multiply(X1,multiply(inverse(multiply(X1,multiply(X2,inverse(multiply(X2,inverse(identity)))))),X3)) = X3 [forward demodulation 892,7]\n",
      "[SA] new: 1017. multiply(X1,multiply(inverse(multiply(X1,multiply(X2,inverse(multiply(X2,identity))))),X3)) = X3 [forward demodulation 1016,10]\n",
      "[SA] new: 1018. multiply(X1,multiply(inverse(multiply(X1,multiply(X2,inverse(X2)))),X3)) = X3 [forward demodulation 1017,1]\n",
      "[SA] new: 1019. multiply(X54,multiply(inverse(multiply(X54,multiply(X55,inverse(multiply(X55,multiply(X53,inverse(X53))))))),X56)) = X56 [forward demodulation 902,6]\n",
      "[SA] new: 1020. multiply(X60,multiply(inverse(multiply(X60,multiply(X61,inverse(multiply(X61,multiply(X57,multiply(X58,multiply(X59,inverse(multiply(X57,multiply(X58,X59))))))))))),X62)) = X62 [forward demodulation 903,8]\n",
      "[SA] new: 1021. multiply(X55,multiply(inverse(multiply(X55,multiply(inverse(multiply(X53,multiply(X54,inverse(multiply(X53,X54))))),identity))),X56)) = X56 [forward demodulation 914,10]\n",
      "[SA] new: 1022. multiply(X55,multiply(inverse(multiply(X55,inverse(multiply(X53,multiply(X54,inverse(multiply(X53,X54))))))),X56)) = X56 [forward demodulation 1021,1]\n",
      "[SA] new: 1023. multiply(X59,multiply(inverse(multiply(X59,multiply(X57,multiply(multiply(X58,inverse(multiply(X57,X58))),inverse(identity))))),X60)) = X60 [forward demodulation 915,8]\n",
      "[SA] new: 1024. multiply(X59,multiply(inverse(multiply(X59,multiply(X57,multiply(X58,multiply(inverse(multiply(X57,X58)),inverse(identity)))))),X60)) = X60 [forward demodulation 1023,8]\n",
      "[SA] new: 1025. multiply(X59,multiply(inverse(multiply(X59,multiply(X57,multiply(inverse(X57),inverse(identity))))),X60)) = X60 [forward demodulation 1024,98]\n",
      "[SA] new: 1026. multiply(X59,multiply(inverse(multiply(X59,inverse(identity))),X60)) = X60 [forward demodulation 1025,49]\n",
      "[SA] new: 1027. multiply(X59,multiply(inverse(multiply(X59,identity)),X60)) = X60 [forward demodulation 1026,10]\n",
      "[SA] new: 1028. multiply(X59,multiply(inverse(X59),X60)) = X60 [forward demodulation 1027,1]\n",
      "[SA] new: 1029. multiply(X63,multiply(inverse(multiply(X63,inverse(multiply(X61,multiply(X62,inverse(multiply(X61,X62))))))),X64)) = X64 [forward demodulation 916,6]\n",
      "[SA] new: 1030. multiply(X69,multiply(inverse(multiply(X69,multiply(X65,multiply(X66,inverse(multiply(X65,multiply(X66,multiply(X67,multiply(X68,inverse(multiply(X67,X68))))))))))),X70)) = X70 [forward demodulation 917,8]\n",
      "[SA] new: 1031. multiply(X2,multiply(inverse(multiply(X2,inverse(multiply(X0,multiply(X1,inverse(multiply(X0,X1))))))),X3)) = X3 [forward demodulation 918,6]\n",
      "[SA] new: 1032. multiply(X8,multiply(inverse(multiply(X8,multiply(X4,multiply(X5,inverse(multiply(X4,multiply(X5,multiply(X6,multiply(X7,inverse(multiply(X6,X7))))))))))),X9)) = X9 [forward demodulation 919,8]\n",
      "[SA] new: 1033. multiply(inverse(X0),multiply(multiply(X0,multiply(X1,multiply(X2,inverse(multiply(X1,X2))))),X3)) = X3 [forward demodulation 920,50]\n",
      "[SA] new: 1034. multiply(inverse(X0),multiply(X0,multiply(multiply(X1,multiply(X2,inverse(multiply(X1,X2)))),X3))) = X3 [forward demodulation 1033,8]\n",
      "[SA] new: 1035. multiply(inverse(X0),multiply(X0,multiply(X1,multiply(multiply(X2,inverse(multiply(X1,X2))),X3)))) = X3 [forward demodulation 1034,8]\n",
      "[SA] new: 1036. multiply(inverse(X0),multiply(X0,multiply(X1,multiply(X2,multiply(inverse(multiply(X1,X2)),X3))))) = X3 [forward demodulation 1035,8]\n",
      "[SA] new: 1037. multiply(inverse(X0),multiply(X0,multiply(X1,multiply(inverse(X1),X3)))) = X3 [forward demodulation 1036,98]\n",
      "[SA] new: 1038. multiply(inverse(X0),multiply(X0,X3)) = X3 [forward demodulation 1037,49]\n",
      "[SA] new: 1039. multiply(X4,multiply(multiply(inverse(X4),multiply(X5,multiply(X6,inverse(multiply(X5,X6))))),X7)) = X7 [forward demodulation 921,50]\n",
      "[SA] new: 1040. multiply(X4,multiply(inverse(X4),multiply(multiply(X5,multiply(X6,inverse(multiply(X5,X6)))),X7))) = X7 [forward demodulation 1039,8]\n",
      "[SA] new: 1041. multiply(X4,multiply(inverse(X4),multiply(X5,multiply(multiply(X6,inverse(multiply(X5,X6))),X7)))) = X7 [forward demodulation 1040,8]\n",
      "[SA] new: 1042. multiply(X4,multiply(inverse(X4),multiply(X5,multiply(X6,multiply(inverse(multiply(X5,X6)),X7))))) = X7 [forward demodulation 1041,8]\n",
      "[SA] new: 1043. multiply(X4,multiply(inverse(X4),multiply(X5,multiply(inverse(X5),X7)))) = X7 [forward demodulation 1042,98]\n",
      "[SA] new: 1044. multiply(X4,multiply(inverse(X4),X7)) = X7 [forward demodulation 1043,49]\n",
      "[SA] new: 1045. multiply(X18,multiply(multiply(inverse(multiply(X18,multiply(X19,inverse(X19)))),multiply(X20,multiply(X21,inverse(multiply(X20,X21))))),X22)) = X22 [forward demodulation 924,50]\n",
      "[SA] new: 1046. multiply(X18,multiply(inverse(multiply(X18,multiply(X19,inverse(X19)))),multiply(multiply(X20,multiply(X21,inverse(multiply(X20,X21)))),X22))) = X22 [forward demodulation 1045,8]\n",
      "[SA] new: 1047. multiply(X18,multiply(inverse(multiply(X18,multiply(X19,inverse(X19)))),multiply(X20,multiply(multiply(X21,inverse(multiply(X20,X21))),X22)))) = X22 [forward demodulation 1046,8]\n",
      "[SA] new: 1048. multiply(X18,multiply(inverse(multiply(X18,multiply(X19,inverse(X19)))),multiply(X20,multiply(X21,multiply(inverse(multiply(X20,X21)),X22))))) = X22 [forward demodulation 1047,8]\n",
      "[SA] new: 1049. multiply(X18,multiply(inverse(multiply(X18,multiply(X19,inverse(X19)))),multiply(X20,multiply(inverse(X20),X22)))) = X22 [forward demodulation 1048,98]\n",
      "[SA] new: 1050. multiply(X18,multiply(inverse(multiply(X18,multiply(X19,inverse(X19)))),X22)) = X22 [forward demodulation 1049,49]\n",
      "[SA] new: 1051. multiply(X23,multiply(multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),multiply(X25,multiply(X26,inverse(multiply(X25,X26))))),X27)) = X27 [forward demodulation 925,50]\n",
      "[SA] new: 1052. multiply(X23,multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),multiply(multiply(X25,multiply(X26,inverse(multiply(X25,X26)))),X27))) = X27 [forward demodulation 1051,8]\n",
      "[SA] new: 1053. multiply(X23,multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),multiply(X25,multiply(multiply(X26,inverse(multiply(X25,X26))),X27)))) = X27 [forward demodulation 1052,8]\n",
      "[SA] new: 1054. multiply(X23,multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),multiply(X25,multiply(X26,multiply(inverse(multiply(X25,X26)),X27))))) = X27 [forward demodulation 1053,8]\n",
      "[SA] new: 1055. multiply(X23,multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),multiply(X25,multiply(inverse(X25),X27)))) = X27 [forward demodulation 1054,98]\n",
      "[SA] new: 1056. multiply(X23,multiply(inverse(multiply(X23,inverse(multiply(X24,inverse(X24))))),X27)) = X27 [forward demodulation 1055,49]\n",
      "[SA] new: 1057. multiply(X28,multiply(multiply(inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30))))))),multiply(X31,multiply(X32,inverse(multiply(X31,X32))))),X33)) = X33 [forward demodulation 926,50]\n",
      "[SA] new: 1058. multiply(X28,multiply(inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30))))))),multiply(multiply(X31,multiply(X32,inverse(multiply(X31,X32)))),X33))) = X33 [forward demodulation 1057,8]\n",
      "[SA] new: 1059. multiply(X28,multiply(inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30))))))),multiply(X31,multiply(multiply(X32,inverse(multiply(X31,X32))),X33)))) = X33 [forward demodulation 1058,8]\n",
      "[SA] new: 1060. multiply(X28,multiply(inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30))))))),multiply(X31,multiply(X32,multiply(inverse(multiply(X31,X32)),X33))))) = X33 [forward demodulation 1059,8]\n",
      "[SA] new: 1061. multiply(X28,multiply(inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30))))))),multiply(X31,multiply(inverse(X31),X33)))) = X33 [forward demodulation 1060,98]\n",
      "[SA] new: 1062. multiply(X28,multiply(inverse(multiply(X28,multiply(X29,inverse(multiply(X29,multiply(X30,inverse(X30))))))),X33)) = X33 [forward demodulation 1061,49]\n",
      "[SA] new: 1063. multiply(X34,multiply(multiply(inverse(multiply(X34,multiply(X35,inverse(multiply(X35,inverse(multiply(X36,inverse(X36)))))))),multiply(X37,multiply(X38,inverse(multiply(X37,X38))))),X39)) = X39 [forward demodulation 927,50]\n",
      "[SA] new: 1064. multiply(X34,multiply(inverse(multiply(X34,multiply(X35,inverse(multiply(X35,inverse(multiply(X36,inverse(X36)))))))),multiply(multiply(X37,multiply(X38,inverse(multiply(X37,X38)))),X39))) = X39 [forward demodulation 1063,8]\n",
      "[SA] new: 1065. multiply(X34,multiply(inverse(multiply(X34,multiply(X35,inverse(multiply(X35,inverse(multiply(X36,inverse(X36)))))))),multiply(X37,multiply(multiply(X38,inverse(multiply(X37,X38))),X39)))) = X39 [forward demodulation 1064,8]\n",
      "[SA] new: 1066. multiply(X34,multiply(inverse(multiply(X34,multiply(X35,inverse(multiply(X35,inverse(multiply(X36,inverse(X36)))))))),multiply(X37,multiply(X38,multiply(inverse(multiply(X37,X38)),X39))))) = X39 [forward demodulation 1065,8]\n",
      "[SA] new: 1067. multiply(X34,multiply(inverse(multiply(X34,multiply(X35,inverse(multiply(X35,inverse(multiply(X36,inverse(X36)))))))),multiply(X37,multiply(inverse(X37),X39)))) = X39 [forward demodulation 1066,98]\n",
      "[SA] new: 1068. multiply(X34,multiply(inverse(multiply(X34,multiply(X35,inverse(multiply(X35,inverse(multiply(X36,inverse(X36)))))))),X39)) = X39 [forward demodulation 1067,49]\n",
      "[SA] new: 1069. multiply(X40,multiply(multiply(inverse(multiply(X40,multiply(X41,multiply(X42,inverse(multiply(X41,X42)))))),multiply(X43,multiply(X44,inverse(multiply(X43,X44))))),X45)) = X45 [forward demodulation 928,50]\n",
      "[SA] new: 1070. multiply(X40,multiply(inverse(multiply(X40,multiply(X41,multiply(X42,inverse(multiply(X41,X42)))))),multiply(multiply(X43,multiply(X44,inverse(multiply(X43,X44)))),X45))) = X45 [forward demodulation 1069,8]\n",
      "[SA] new: 1071. multiply(X40,multiply(inverse(multiply(X40,multiply(X41,multiply(X42,inverse(multiply(X41,X42)))))),multiply(X43,multiply(multiply(X44,inverse(multiply(X43,X44))),X45)))) = X45 [forward demodulation 1070,8]\n",
      "[SA] new: 1072. multiply(X40,multiply(inverse(multiply(X40,multiply(X41,multiply(X42,inverse(multiply(X41,X42)))))),multiply(X43,multiply(X44,multiply(inverse(multiply(X43,X44)),X45))))) = X45 [forward demodulation 1071,8]\n",
      "[SA] new: 1073. multiply(X40,multiply(inverse(multiply(X40,multiply(X41,multiply(X42,inverse(multiply(X41,X42)))))),multiply(X43,multiply(inverse(X43),X45)))) = X45 [forward demodulation 1072,98]\n",
      "[SA] new: 1074. multiply(X40,multiply(inverse(multiply(X40,multiply(X41,multiply(X42,inverse(multiply(X41,X42)))))),X45)) = X45 [forward demodulation 1073,49]\n",
      "[SA] new: 1075. multiply(X46,multiply(multiply(inverse(multiply(X46,multiply(X47,inverse(multiply(X47,multiply(X48,multiply(X49,inverse(multiply(X48,X49))))))))),multiply(X50,multiply(X51,inverse(multiply(X50,X51))))),X52)) = X52 [forward demodulation 929,50]\n",
      "[SA] new: 1076. multiply(X46,multiply(inverse(multiply(X46,multiply(X47,inverse(multiply(X47,multiply(X48,multiply(X49,inverse(multiply(X48,X49))))))))),multiply(multiply(X50,multiply(X51,inverse(multiply(X50,X51)))),X52))) = X52 [forward demodulation 1075,8]\n",
      "[SA] new: 1077. multiply(X46,multiply(inverse(multiply(X46,multiply(X47,inverse(multiply(X47,multiply(X48,multiply(X49,inverse(multiply(X48,X49))))))))),multiply(X50,multiply(multiply(X51,inverse(multiply(X50,X51))),X52)))) = X52 [forward demodulation 1076,8]\n",
      "[SA] new: 1078. multiply(X46,multiply(inverse(multiply(X46,multiply(X47,inverse(multiply(X47,multiply(X48,multiply(X49,inverse(multiply(X48,X49))))))))),multiply(X50,multiply(X51,multiply(inverse(multiply(X50,X51)),X52))))) = X52 [forward demodulation 1077,8]\n",
      "[SA] new: 1079. multiply(X46,multiply(inverse(multiply(X46,multiply(X47,inverse(multiply(X47,multiply(X48,multiply(X49,inverse(multiply(X48,X49))))))))),multiply(X50,multiply(inverse(X50),X52)))) = X52 [forward demodulation 1078,98]\n",
      "[SA] new: 1080. multiply(X46,multiply(inverse(multiply(X46,multiply(X47,inverse(multiply(X47,multiply(X48,multiply(X49,inverse(multiply(X48,X49))))))))),X52)) = X52 [forward demodulation 1079,49]\n",
      "[SA] new: 1081. multiply(inverse(multiply(X53,inverse(multiply(X53,multiply(X54,multiply(X55,inverse(multiply(X54,X55)))))))),multiply(identity,X56)) = X56 [forward demodulation 930,10]\n",
      "[SA] new: 1082. multiply(inverse(multiply(X53,inverse(multiply(X53,multiply(X54,multiply(X55,inverse(multiply(X54,X55)))))))),X56) = X56 [forward demodulation 1081,6]\n",
      "[SA] new: 1083. multiply(multiply(X57,inverse(multiply(X57,multiply(X58,multiply(X59,inverse(multiply(X58,X59))))))),multiply(identity,X60)) = X60 [forward demodulation 931,10]\n",
      "[SA] new: 1084. multiply(multiply(X57,inverse(multiply(X57,multiply(X58,multiply(X59,inverse(multiply(X58,X59))))))),X60) = X60 [forward demodulation 1083,6]\n",
      "[SA] new: 1085. multiply(X0,multiply(X1,inverse(multiply(X1,multiply(X2,multiply(X3,inverse(multiply(X2,X3)))))))) = X0 [forward demodulation 934,1]\n",
      "[SA] new: 1086. inverse(multiply(X4,multiply(X5,inverse(multiply(X5,multiply(X6,multiply(X7,inverse(multiply(X6,X7))))))))) = X4 [forward demodulation 935,1]\n",
      "[SA] new: 1087. multiply(X17,X21) = multiply(multiply(X17,multiply(X18,inverse(multiply(X18,multiply(X19,multiply(X20,inverse(multiply(X19,X20)))))))),X21) [forward demodulation 938,50]\n",
      "[SA] new: 1088. multiply(X13,multiply(inverse(X13),X14)) = X14 [forward demodulation 949,412]\n",
      "[SA] new: 1089. multiply(inverse(multiply(X15,inverse(multiply(X15,multiply(X16,multiply(X17,inverse(multiply(X16,X17)))))))),X18) = X18 [forward demodulation 950,6]\n",
      "[SA] new: 1090. multiply(X19,multiply(X20,multiply(inverse(multiply(X19,multiply(X20,multiply(X21,inverse(multiply(X21,multiply(X22,multiply(X23,inverse(multiply(X22,X23)))))))))),X24))) = X24 [forward demodulation 951,8]\n",
      "[SA] new: 1091. multiply(X4,multiply(inverse(multiply(X4,multiply(X5,inverse(multiply(X5,multiply(X6,multiply(X7,inverse(multiply(X6,X7))))))))),X9)) = X9 [forward demodulation 953,49]\n",
      "[SA] new: 1092. multiply(X24,multiply(inverse(X24),X25)) = X25 [forward demodulation 956,412]\n",
      "[SA] new: 1093. multiply(X36,multiply(inverse(multiply(X36,multiply(X32,inverse(multiply(X32,multiply(X33,inverse(multiply(X33,multiply(X34,multiply(X35,inverse(multiply(X34,X35)))))))))))),X37)) = X37 [forward demodulation 958,50]\n",
      "[SA] new: 1094. multiply(X38,multiply(inverse(multiply(X38,multiply(X39,inverse(multiply(X39,multiply(X40,multiply(X41,inverse(multiply(X40,X41))))))))),X42)) = X42 [forward demodulation 959,50]\n",
      "[SA] new: 1095. multiply(X43,multiply(multiply(inverse(multiply(X43,multiply(X44,inverse(multiply(X44,multiply(X45,multiply(X46,inverse(multiply(X45,X46))))))))),multiply(X47,inverse(X47))),X48)) = X48 [forward demodulation 960,50]\n",
      "[SA] new: 1096. multiply(X43,multiply(inverse(multiply(X43,multiply(X44,inverse(multiply(X44,multiply(X45,multiply(X46,inverse(multiply(X45,X46))))))))),multiply(multiply(X47,inverse(X47)),X48))) = X48 [forward demodulation 1095,8]\n",
      "[SA] new: 1097. multiply(X43,multiply(inverse(multiply(X43,multiply(X44,inverse(multiply(X44,multiply(X45,multiply(X46,inverse(multiply(X45,X46))))))))),multiply(X47,multiply(inverse(X47),X48)))) = X48 [forward demodulation 1096,8]\n",
      "[SA] new: 1098. multiply(X43,multiply(inverse(multiply(X43,multiply(X44,inverse(multiply(X44,multiply(X45,multiply(X46,inverse(multiply(X45,X46))))))))),X48)) = X48 [forward demodulation 1097,49]\n",
      "[SA] new: 1099. multiply(X49,multiply(multiply(inverse(multiply(X49,multiply(X50,inverse(multiply(X50,multiply(X51,multiply(X52,inverse(multiply(X51,X52))))))))),inverse(multiply(X53,inverse(X53)))),X54)) = X54 [forward demodulation 961,50]\n",
      "[SA] new: 1100. multiply(X49,multiply(inverse(multiply(X49,multiply(X50,inverse(multiply(X50,multiply(X51,multiply(X52,inverse(multiply(X51,X52))))))))),multiply(inverse(multiply(X53,inverse(X53))),X54))) = X54 [forward demodulation 1099,8]\n",
      "[SA] new: 1101. multiply(inverse(multiply(X61,inverse(multiply(X61,multiply(X62,multiply(X63,inverse(multiply(X62,X63)))))))),X64) = X64 [forward demodulation 963,6]\n",
      "[SA] new: 1102. multiply(X70,multiply(inverse(multiply(X70,multiply(X65,multiply(inverse(multiply(X65,multiply(X66,inverse(multiply(X66,multiply(X67,multiply(X68,inverse(multiply(X67,X68))))))))),multiply(X69,inverse(X69)))))),X71)) = X71 [forward demodulation 964,8]\n",
      "[SA] new: 1103. multiply(X70,multiply(inverse(multiply(X70,multiply(X69,inverse(X69)))),X71)) = X71 [forward demodulation 1102,412]\n",
      "[SA] new: 1104. multiply(X77,multiply(X72,multiply(inverse(multiply(X72,multiply(X73,inverse(multiply(X73,multiply(X74,multiply(X75,inverse(multiply(X74,X75))))))))),multiply(X76,multiply(inverse(multiply(X77,X76)),X78))))) = X78 [forward demodulation 965,8]\n",
      "[SA] new: 1105. multiply(X77,multiply(X76,multiply(inverse(multiply(X77,X76)),X78))) = X78 [forward demodulation 1104,412]\n",
      "[SA] new: 1106. multiply(X77,multiply(inverse(X77),X78)) = X78 [forward demodulation 1105,98]\n",
      "[SA] new: 1107. multiply(X94,X95) = multiply(X90,multiply(inverse(multiply(X90,multiply(X91,inverse(multiply(X91,multiply(X92,multiply(X93,inverse(multiply(X92,X93))))))))),multiply(X94,X95))) [forward demodulation 968,8]\n",
      "[SA] new: 1108. multiply(X107,multiply(X108,multiply(inverse(multiply(X107,multiply(X108,multiply(X109,inverse(multiply(X109,multiply(X110,multiply(X111,inverse(multiply(X110,X111)))))))))),X112))) = X112 [forward demodulation 971,8]\n",
      "\n"
     ]
    }
   ],
   "source": [
    "print(interactive.proof)"
   ]
  }
 ],
 "metadata": {
  "kernelspec": {
   "display_name": "foreduce",
   "language": "python",
   "name": "python3"
  },
  "language_info": {
   "codemirror_mode": {
    "name": "ipython",
    "version": 3
   },
   "file_extension": ".py",
   "mimetype": "text/x-python",
   "name": "python",
   "nbconvert_exporter": "python",
   "pygments_lexer": "ipython3",
   "version": "3.12.2"
  }
 },
 "nbformat": 4,
 "nbformat_minor": 2
}
